<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 10: Statistical Theory in Practice

## References

* [Why the Mean Matters](https://ccsf-math-108.github.io/textbook/chapters/14/Why_the_Mean_Matters.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Distribution Bounds and Approximations

---

In each task in this section, fill in the blank with a percent that makes the statement true **without further assumptions**.

*Note:* No credit will be given for loose bounds such as "at least 0%" or "at most 100%". Give the best answer that is possible with the information given.

---

### Task 01 📍

A community has an average age of 45 years with a standard deviation of 5 years. 

Assign `task_01` to a numerical value (`int` or `float`) that completes the following statement based on information from [Chebyshev's bounds](https://ccsf-math-108.github.io/textbook/chapters/14/2/Variability.html#chebychev-s-bounds).

**Statement**: According to Chebyshev's bounds, at least ________% of the people are between 25 and 65 years old.

**Note**: Keep in mind that you do not know whether or not the ages are normally distributed.

_Points:_ 2

```python
task_01 = ...
print(f"According to Chebyshev's bounds, at least {task_01}% of the people are between 25 and 65 years old.")
```

```python
grader.check("task_01")
```

---

### Task 02 📍

A community has an average age of 45 years with a standard deviation of 5 years. 

Assign `task_02` to a numerical value (`int` or `float`) that completes the following statement based on information from [Chebyshev's bounds](https://ccsf-math-108.github.io/textbook/chapters/14/2/Variability.html#chebychev-s-bounds).

**Statement**: According to Chebyshev's bounds, at most _______% of the people have ages that are not in the range 25 years to 65 years.

**Note**: Keep in mind that you do not know whether or not the ages are normally distributed.

_Points:_ 2

```python
task_02 = ...
print(f"According to Chebyshev's bounds, at most {task_02}% of the people have ages that are not in the range 25 years to 65 years.")
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Another community also has an average age of 45 years with a standard deviation of 5 years **AND** you know the ages of that community are normally distributed.

Assign `task_03` to a numerical value (`int` or `float`) that completes the following statement based on information from column _Normal Distribution: Approximation_ in the table at the bottom of [The Standard Normal CDF section of the textbook](https://ccsf-math-108.github.io/textbook/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf).

**Statement**: Since the ages are normally distributed, about _____% of the people are between 35 and 55 years old.

_Points:_ 2

```python
task_03 = ...
print(f"Since the ages are normally distributed, about {task_03}% of the people are between 35 and 55 years old.")
```

```python
grader.check("task_03")
```

---

## Sample Size and Confidence Level

---

A data science class at the large Data 8 University wants to estimate the percentage of Facebook users among students at the school. To do this, they need to take a random sample of students. You can assume that their method of sampling is equivalent to drawing at random with replacement from students at the school.

***Please review [Section 14.6](https://ccsf-math-108.github.io/textbook/chapters/14/6/Choosing_a_Sample_Size.html#) of the textbook before proceeding with this section. Your work will go much faster that way!***

---

### Task 04 📍

Assign `smallest` to the smallest number of students (an integer value) they should sample to ensure that a **95%** confidence interval for the parameter has a width of no more than 6% from the left end to the right end.

*Hint:* How can our data be represented to show if a student in the sample is a Facebook user? Given this, what assumptions can we make for the SD of the population? [Section 14.6](https://ccsf-math-108.github.io/textbook/chapters/14/6/Choosing_a_Sample_Size.html#) might be helpful!

*Note:* You can use `np.ceil` to round a float up to the next whole number. Keep in mind that the resulting value will still be a `float` data type. You'll need to convert the float to an integer.

_Points:_ 2

```python
smallest = ...
smallest
```

```python
grader.check("task_04")
```

---

### Task 05 📍

Suppose the data science class decides to construct a 90% confidence interval instead of a 95% confidence interval, but they still require that the width of the interval is no more than 6% from the left end to the right end. 

Will they need the same sample size as in [Task 04](#Task-04-📍)? Assign the integer corresponding to the correct response from the options listed below to the name `task_05`.

1. Yes, they must use the same sample size.
2. No, a smaller sample size will work.
3. No, they will need a bigger sample.

_Points:_ 2

```python
task_05 = ...
task_05
```

```python
grader.check("task_05")
```

---

### Task 06 📍

The professor tells the class that a 90% confidence interval for the parameter is constructed exactly like a 95% confidence interval, except that you have to go only 1.65 SDs on either side of the estimate (±1.65) instead of 2 SDs on either side (±2). Assign `smallest_num` to the smallest number of students (an integer) they should sample to ensure that a **90%** confidence interval for the parameter has a width of no more than 6% from the left end to the right end.

*Note:* Make sure you round up your sample size value and make sure it is an integer.


_Points:_ 2

```python
smallest_num = ...
smallest_num
```

```python
grader.check("task_06")
```

---

For this next exercise, please consult [Section 14.3.4](https://ccsf-math-108.github.io/textbook/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf) of the textbook for similar examples.

The students are curious about how the professor came up with the value 1.65 in [Task 06](#Task-06-📍). She says she ran the following two code cells. The first one calls the `datascience` library function `plot_normal_cdf`, which displays the proportion that is at most the specified number of SDs above average under the normal curve plotted with standard units on the horizontal axis. You can find the documentation [here](http://data8.org/datascience/util.html#datascience.util.plot_normal_cdf).

*Note:* The acronym `cdf` stands for `cumulative distribution function`. It measures the proportion to the left of a specified point under a probability histogram. In the graph you get by running the cell below, it is represented by the yellow shaded area. You are not responsible for knowing the `cdf` function.

```python
plot_normal_cdf(1.65)
```

To run the second cell, the professor had to first import a Python library for probability and statistics. Note, you are not responsible for knowing the `scipy` library below, nor the `norm.cdf` function in the library:

```python
# Just run this cell
from scipy import stats
```

Then she used the `norm.cdf` method in the library to find the gold proportion above.

```python
# Just run this cell
stats.norm.cdf(1.65)
```

*Note*: You do not need to understand how the `scipy` library or how to use the method yourself.

---

### Task 07 📍

This shows that the percentage in a normal distribution that is at most 1.65 SDs above average is about 95%. Why is 1.65 the right number of SDs to use when constructing a 90% confidence interval? 

Assign `why_165` to the integer `1, 2, 3, 4` that is associated with the statement below that best answers this question.

1. 1.65 is the number of SDs needed to ensure that the area to the left of the interval is exactly 10%.
2. The area between -1.65 and 1.65 SDs includes 90% of the data, leaving 5% on each side outside this range, creating a total of 90% within the confidence interval.
3. A 90% confidence interval requires the data to extend 1.65 SDs above the average to include the upper 95% of the distribution.
4. The area to the right of -1.65 SDs contains 95% of the distribution, so 1.65 SDs is used for a 90% confidence interval.

_Points:_ 2

```python
why_165 = ...
why_165
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Run the following cell to see that the proportion that is at most 2.33 SDs above average in a normal distribution is 99%.

```python
# Just run this cell, do not change it.
stats.norm.cdf(2.33)
```

With this in mind, assign `option` to the right option `1, 2, 3, 4` to fill in the blank in the following statement:

If you start at the estimate and go 2.33 SDs on either side, then you will get a _______% confidence interval for the parameter.

1. 98
2. 98.5
3. 99
4. 99.5



_Points:_ 2

```python
option = ...
option
```

```python
grader.check("task_08")
```

---

## Polling and the Normal Distribution

---

Michelle is a statistical consultant, and she works for a group that supports Proposition 68 (which would mandate labeling of all horizontal and vertical axes) called Yes on 68.  They want to know how many Californians will vote for the proposition.

Michelle polls a uniform random sample of all California voters, and she finds that 210 of the 400 sampled voters will vote in favor of the proposition. We have provided a table for you below which has 3 columns: the first two columns are identical to `sample`. The third column contains the proportion of total voters that chose each option.

```python
sample = Table().with_columns(
    "Vote",  make_array("Yes", "No"),
    "Count", make_array(210,   190))

sample_size = sum(sample.column("Count"))
sample_with_proportions = sample.with_column("Proportion", sample.column("Count") / sample_size)
sample_with_proportions
```

---

### Task 09 📍

Michelle wants to use 10,000 bootstrap resamples to compute a confidence interval for the proportion of all California voters who will vote Yes.  

Fill in the next cell to simulate an empirical distribution of Yes proportions. Use bootstrap resampling to simulate 10,000 election outcomes, and assign `resample_yes_proportions` to contain the Yes proportion of each bootstrap resample. Then, visualize `resample_yes_proportions` with a histogram. **You should see a bell-shaped curve centered near the proportion of Yes in the original sample.**

*Hint:* `sample_proportions` may be useful here!


_Points:_ 2

```python
resample_yes_proportions = make_array()
for i in np.arange(10000):
    resample = ...
    resample_yes_proportions = ...
Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
```

```python
grader.check("task_09")
```

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

The Central Limit Theorem states: **the probability distribution of the sum or average of a large random sample drawn with replacement will be roughly normal, regardless of the distribution of the population from which the sample is drawn.**

Identify which specific instructions in the tasks above, correspond to the conditions and consequence of the Central Limit Theorem?


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In a population whose members are 0 and 1, there is a simple formula for the **standard deviation of that population**:

$$\text{standard deviation of population} = \sqrt{(\text{proportion of 0s}) \times (\text{proportion of 1s})}$$

(Figuring out this formula, starting from the definition of the standard deviation, is an fun exercise for those who enjoy algebra.)

---

### Task 11 📍

 Using only the Central Limit Theorem and the numbers of Yes and No voters in our sample of 400, *algebraically* compute the predicted standard deviation of the `resample_yes_proportions` array. Assign this number to `approximate_sd`. **Do not access the data in `resample_yes_proportions` in any way.**

Remember that the standard deviation of the sample means can be computed from the population SD and the size of the sample (the formula above might be helpful). If we do not know the population SD, we can use the sample SD as a reasonable approximation in its place. [This section](https://ccsf-math-108.github.io/textbook/chapters/14/5/Variability_of_the_Sample_Mean.html#the-sd-of-all-the-sample-means) of the textbook also may be helpful.


_Points:_ 2

```python
approx_pop_sd = ...
approximate_sd = ...
approximate_sd
```

```python
grader.check("task_11")
```

---

### Task 12 📍

Compute the standard deviation of the array `resample_yes_proportions`, which will act as an approximation to the true SD of the possible sample proportions. This will help verify whether your answer to [Task 11](#Task-11-📍) is approximately correct.


_Points:_ 2

```python
exact_sd = ...
exact_sd
```

```python
grader.check("task_12")
```

---

### Task 13 📍

**Again, without accessing `resample_yes_proportions` in any way**, compute an approximate 95% confidence interval for the proportion of Yes voters in California.

The cell below draws your interval as a red bar below the histogram of `resample_yes_proportions`; use that to verify that your answer looks right.

*Hint:* How many SDs correspond to 95% of the distribution promised by the CLT? Recall the discussion in the textbook [here](https://ccsf-math-108.github.io/textbook/chapters/14/3/SD_and_the_Normal_Curve.html).

*Hint:* The `approximate_sd` variable you previously defined may be helpful!


_Points:_ 3

```python
lower_limit = ...
upper_limit = ...
print('lower:', lower_limit, 'upper:', upper_limit)
```

```python
grader.check("task_13")
```

---

Run the following cell to plot your confidence interval.

```python
# Run this cell to plot your confidence interval.
Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
plt.plot(make_array(lower_limit, upper_limit), make_array(0, 0), c='r', lw=10);
```

Your confidence interval should overlap the number 0.5.  That means we can't be very sure whether Proposition 68 is winning, even though the sample Yes proportion is a bit above 0.5.

The Yes on 68 campaign really needs to know whether they're winning.  It's impossible to be absolutely sure without polling the whole population, but they'd be okay if the standard deviation of the sample mean were only 0.005.  They ask Michelle to run a new poll with a sample size that's large enough to achieve that.  (Polling is expensive, so the sample also shouldn't be bigger than necessary.)

Michelle consults Chapter 14 of your textbook.  Instead of making the conservative assumption that the population standard deviation is 0.5 (coding Yes voters as 1 and No voters as 0), she decides to assume that it's equal to the standard deviation of the sample,

$$\sqrt{(\text{Yes proportion in the sample}) \times (\text{No proportion in the sample})}.$$

Under that assumption, Michelle decides that a sample of 9,975 would suffice.

---

### Task 14 📍

Does Michelle's sample size achieve the desired standard deviation of sample means? What SD would you achieve with a smaller sample size? A higher sample size?

To explore this, first compute the SD of sample means obtained by using Michelle's sample size.

_Points:_ 4

```python
estimated_population_sd = ...
michelle_sample_size = ...
michelle_sample_mean_sd = ...
print("With Michelle's sample size, you would predict a sample mean SD of %f." % michelle_sample_mean_sd)
```

```python
grader.check("task_14")
```

---

### Task 15 📍

Next, compute the SD of the sample means that you would get from a smaller sample size. Ideally, you should pick a number that is significantly smaller, but any sample size smaller than Michelle's will do.*


_Points:_ 2

```python
smaller_sample_size = ...
smaller_sample_mean_sd = ...
print("With this smaller sample size, you would predict a sample mean SD of %f" % smaller_sample_mean_sd)
```

```python
grader.check("task_15")
```

---

### Task 16 📍

Finally, compute the SD of the sample means that you would get from a larger sample size. Here, a number that is significantly larger would make any difference more obvious, but any sample size larger than Michelle's will do.

_Points:_ 2

```python
larger_sample_size = ...
larger_sample_mean_sd = ...
print("With this larger sample size, you would predict a sample mean SD of %f" % larger_sample_mean_sd)
```

```python
grader.check("task_16")
```

---

### Task 17 📍

Based on this, was Michelle's sample size approximately the minimum sufficient sample, given her assumption that the sample SD is the same as the population SD? Assign `min_sufficient` to `True` if 9,975 was indeed approximately the minimum sufficient sample, and `False` if it wasn't.


_Points:_ 2

```python
min_sufficient = ...
min_sufficient
```

```python
grader.check("task_17")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> import numbers
>>> isinstance(task_01, numbers.Number)
True
```
failure_message: "❌ task_01 is not assigned to a numerical value."
success_message: "✅ task_01 is assigned to a numerical value."

```
>>> 0 <= task_01 <= 100
True
```
failure_message: "❌ task_01 is not assigned to a numerical value between 0 and 100."
success_message: "✅ task_01 is assigned to a numerical value between 0 and 100."



Tests for task_02 (points: 2)
```
>>> import numbers
>>> isinstance(task_02, numbers.Number)
True
```
failure_message: "❌ task_02 is not assigned to a numerical value."
success_message: "✅ task_02 is assigned to a numerical value."

```
>>> 0 <= task_02 <= 100
True
```
failure_message: "❌ task_02 is not assigned to a numerical value between 0 and 100."
success_message: "✅ task_02 is assigned to a numerical value between 0 and 100."



Tests for task_03 (points: 2)
```
>>> import numbers
>>> isinstance(task_03, numbers.Number)
True
```
failure_message: "❌ task_03 is not assigned to a numerical value."
success_message: "✅ task_03 is assigned to a numerical value."

```
>>> 0 <= task_03 <= 100
True
```
failure_message: "❌ task_03 is not assigned to a numerical value between 0 and 100."
success_message: "✅ task_03 is assigned to a numerical value between 0 and 100."



Tests for task_04 (points: 2)
```
>>> 1100 <= smallest <= 1150
True
```
failure_message: ❌ smallest should represent a value between 110 and 1150, inclusive.
success_message: ✅ smallest represents a value between 110 and 1150, inclusive.
```
>>> isinstance(smallest, (int, np.integer))
True
```
failure_message: ❌ smallest should be an integer.
success_message: ✅ smallest is an integer.


Tests for task_05 (points: 2)
```
>>> isinstance(task_05, (int, np.integer))
True
```
failure_message: "❌ task_05 is not assigned to an int value."
success_message: "✅ task_05 is assigned to an int value."

```
>>> task_05 in {1, 2, 3}
True
```
failure_message: "❌ task_05 is not assigned to 1, 2, or 3."
success_message: "✅ task_05 is assigned to 1, 2, or 3."



Tests for task_06 (points: 2)
```
>>> 700 <= smallest_num <= 800
True
```
failure_message: ❌ smallest_num should represent a value between 700 and 800, inclusive.
success_message: ✅ smallest_num represents a value between 700 and 800, inclusive.
```
>>> isinstance(smallest, (int, np.integer))
True
```
failure_message: ❌ smallest_num should be an integer.
success_message: ✅ smallest_num is an integer.


Tests for task_07 (points: 2)
```
>>> why_165 in {1, 2, 3, 4}
True
```
failure_message: ❌ why_165 should represent 1, 2, 3, or 4.
success_message: ✅ why_165 represents 1, 2, 3, or 4.


Tests for task_08 (points: 2)
```
>>> option in {1, 2, 3, 4}
True
```
failure_message: ❌ option should represent 1, 2, 3, or 4.
success_message: ✅ option represents 1, 2, 3, or 4.


Tests for task_09 (points: 2)
```
>>> bool(len(resample_yes_proportions) == 10000)
True
```
failure_message: ❌ resample_yes_proportions should contain 10,000 items.
success_message: ✅ resample_yes_proportions contains 10,000 items.
```
>>> bool(0 <= abs(np.mean(resample_yes_proportions) - 0.525) < 0.025 and 0 <= np.std(resample_yes_proportions) < 0.08)
True
```
failure_message: ❌ The mean and standard deviation of resample_yes_proportions does not seem possible.
success_message: ✅ The mean and standard deviation of resample_yes_proportions seems possible.


Tests for task_11 (points: 2)
```
>>> np.float64(approximate_sd) < 0.025
np.True_
```
failure_message: ❌ approximate_sd should be less than 0.025.
success_message: ✅ approximate_sd is less than 0.025.


Tests for task_12 (points: 2)
```
>>> exact_sd < 0.05
np.True_
```
failure_message: ❌ exact_sd should be less than 0.05
success_message: ✅ exact_sd is less than 0.05


Tests for task_13 (points: 3)
```
>>> 0.4 <= np.float64(lower_limit) < np.float64(upper_limit) <= 0.7
np.True_
```
failure_message: ❌ The lower_limit and upper_limit values do not seem possible.
success_message: ✅ The lower_limit and upper_limit values seem possible.


Tests for task_14 (points: 4)
```
>>> isinstance(michelle_sample_size, (int, np.integer))
True
```
failure_message: ❌ 
success_message: ✅ 


Tests for task_15 (points: 2)
```
>>> np.float64(smaller_sample_size) < np.float64(michelle_sample_size)
np.True_
```
failure_message: ❌ smaller_sample_size should be less than michelle_sample_size.
success_message: ✅ smaller_sample_size is less than michelle_sample_size.


Tests for task_16 (points: 2)
```
>>> np.float64(larger_sample_size) > np.float64(michelle_sample_size)
np.True_
```
failure_message: ❌ larger_sample_size should be greater than michelle_sample_size.
success_message: ✅ larger_sample_size is greater than michelle_sample_size.


Tests for task_17 (points: 2)
```
>>> isinstance(min_sufficient, bool)
True
```
failure_message: ❌ min_sufficient should be a Boolean value.
success_message: ✅ min_sufficient is a Boolean value.
