<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 02: Data Types

---

## Textbook References

* [Data Types](https://ccsf-math-108.github.io/textbook/chapters/04/Data_Types.html) 
* [Sequences](https://ccsf-math-108.github.io/textbook/chapters/05/Sequences.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the follwing code cell to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")

# Configure this Notebook
import numpy as np
from datascience import *
```

<!-- END QUESTION -->

---

## Creating Arrays


---

### Note About Arrays 

You might have worked with Python lists before. Python lists are different/behave differently than NumPy arrays. In this course, you will primarily use arrays because of their efficiency.

---

### Task 01 📍

Use the function `make_array` to assign `weird_numbers` to an array containing the following numbers (in the given order):

1. 3,000
2. the square root of 5
3. -2
4. $\pi$ to the power of the square root of 2

Remember that `math` module contains the square root function and the constant value $\pi$.

_Points:_ 2

```python
# You will probably want to import something to start!
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("task_01")
```

---

### Task 02 📍

Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".

**Note:** In your response, you need to match the spacing and capitalization to the above strings.

_Points:_ 2

```python
book_title_words = ...
book_title_words
```

```python
grader.check("task_02")
```

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string.

---

### Task 03 📍

Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

*Hint:* If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .


_Points:_ 3

```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

```python
grader.check("task_03")
```

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Run the following code, which makes an array and does not show an error:

```python
make_array('a string', 1234)
```

Does this mean that it is possible to make an array where one item is a string and the other item is an integer? Respond with Yes or No and provide an explanation to support your response.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Indexing Arrays


These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0, the second element is at index 1, etc.

---

### Note About Creating Arrays

When you are working with an array (`an_array` for example) in this class:

**We recommend that you use `an_array.item(...)` to access the items in an element in the array to be consistent with the course materials.** 

We are working to accept using `an_array[0]` in all situations. The bracket notation is the more standard way of indexing, but you might run into issues with data types that might lead to the auto-grader saying your work is incorrect.

---

### Task 05 📍

The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.


_Points:_ 2

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("task_05")
```

The test above checks that your answer is in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully. The same is true for most of the tests in this homework, and every other homework for this class.

---

### Task 06 📍

You'll sometimes want to find the *last* element of an array.  Suppose an array has 145 elements.  What is the index of its last element?


_Points:_ 2

```python
index_of_last_element = ...
```

```python
grader.check("task_06")
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

---

### Task 07 📍

The cell below loads an array called `president_birth_years`.  Calling `.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent birth year of any deceased president. Assign that year to `most_recent_birth_year`.


_Points:_ 2

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

last_index = ...
most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Finally, assign `sum_of_birth_years` to the sum of the third, tenth, and last birth year in `president_birth_years`. You might consider breaking up your solution into parts.

_Points:_ 2

```python
sum_of_birth_years = ...
```

```python
grader.check("task_08")
```

---

## Basic Array Arithmetic


---

### Task 09 📍

The goal here is to multiply each of the numbers 42, 4224, 42422424, and -250 by the number 157. Assign each variable below such that `first_product` is assigned to the result of $42 * 157$, `second_product` is assigned to the result of $4224 * 157$, and so on. 

For this task, **don't** use arrays. You'll use them later.


_Points:_ 2

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

```python
grader.check("task_09")
```

---

### Task 10 📍

Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.


_Points:_ 2

```python
numbers = ...
products = numbers * 157
products
```

```python
grader.check("task_10")
```

---

### Task 11 📍

Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the correct products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers.


_Points:_ 2

```python
correct_products = ...
correct_products
```

```python
grader.check("task_11")
```

---

### Task 12 📍

We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a particular day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](https://noaa.gov), all the temperatures are in Fahrenheit.

Convert all temperature values from Fahrenheit to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Make sure to **ROUND** the final result after converting to Celsius to the nearest integer using the `np.round` function. Remember to check the documentation of a function if you are not sure how it works.

_Points:_ 3

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("task_12")
```

---

### Task 13 📍

The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Pay attention to the units, give your answer in Celsius!** Make sure **NOT** to round your answer for this question!


_Points:_ 3

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("task_13")
```

---

## World Population


The cell below loads a table of estimates of the world population for from 1950 until 2024. The estimates come from the [US Census Bureau website](https://www.census.gov/en.html).

```python
world = Table.read_table("world_population.csv")
world
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column('Population')
population
```

In this question, you will apply some built-in Numpy functions to this array. Numpy is a module that is often used in Data Science!

<img src="array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

<img src="array_cumsum.png" style="width: 700px;"/>

The cumulative sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

---

### Task 14 📍

Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.


_Points:_ 2

```python
largest_population_change = ...
largest_population_change
```

```python
grader.check("task_14")
```

---

### Task 15 📍

Run the following code cell and then answer the question below. 

```python
np.cumsum(np.diff(population))
```

What do the values in the resulting array represent? Assign `cumulative_sum_answer` to `1`, `2`, or `3` indicating the correct answer from the choices below.

1. The total population change between consecutive years, starting at 1951.
2. The total population change between 1950 and each later year, starting at 1951.
3. The total population change between 1950 and each later year, starting inclusively at 1950.

_Points:_ 2

```python
cumulative_sum_answer = ...
```

```python
grader.check("task_15")
```

---

## Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

---

### Task 16 📍

The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` to values computed using useful functions applied to the array of waiting times so that the print statement is correct.

_Points:_ 2

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("task_16")
```

---

### Task 17 📍

Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was 74 - 62 = 12 minutes. 

*Hint 1*: You'll need an array arithmetic function [mentioned in the textbook](https://ccsf-math-108.github.io/textbook/chapters/05/1/Arrays.html#functions-on-arrays). You have also seen this function earlier in the homework!

*Hint 2*: We want to return the absolute value of the biggest decrease.

*Hint 3*: If a value goes from `87` to `80`, then notice that the difference `80 - 87` is negative!

_Points:_ 2

```python
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("task_17")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

1. **Review the Rubric:** View the rubric on the Canvas assignment page to understand the scoring criteria.  
2. **Review Manually Graded Tasks:** Verify that you have responded to all manually graded tasks marked with 📍🔎.  
3. **Review Auto-Graded Tasks:** Inspect the output from the `grade.check_all()` command above. Ensure you have passed all available tests for the auto-graded tasks in this notebook marked with 📍.  
   - This command executes all auto-grader tests sequentially.  
   - Note that Homework and Project assignments include hidden auto-grader tests that are not visible to you. Passing all visible tests does **not** guarantee that your solution is fully correct.  
4. **Save Your Work:** In the notebook toolbar, go to `File → Save Notebook` to save your work and create a checkpoint.  
5. **Download the Notebook:** In the notebook toolbar, go to `File → Download IPYNB` to download the notebook (`.ipynb`) file.  
6. **Upload to Canvas:** Upload the downloaded `.ipynb` file to the appropriate Canvas assignment page.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> import numpy
>>> isinstance(weird_numbers, numpy.ndarray) and len(weird_numbers) == 4
True
```
failure_message: ❌ You didn't provide an array of 4 values. Consider using make_array.
success_message: ✅ You provided an array of 4 values.
```
>>> np.allclose(weird_numbers, np.array([3000.0, 2.23606798, -2, 5.04749727]), rtol=0.001, atol=0.001) == True
True
```


Tests for task_02 (points: 2)
```
>>> import numpy
>>> isinstance(book_title_words, numpy.ndarray) and len(book_title_words) == 3
True
```
failure_message: ❌ You didn't provide an array of 3 values. Consider using make_array.
success_message: ✅ You provided an array of 3 values.
```
>>> not any([',' in text for text in book_title_words])
True
```
failure_message: ❌ It looks like you included commas in the text. The three pieces of text in the array should be: 'Eats' 'Shoots' 'Leaves'
success_message: ✅ Great work not including the commas.
```
>>> 'and ' in book_title_words.item(2)
True
```
failure_message: ❌ It looks like you didn't include both words in the last piece of text. It should be the actual string: 'and Leaves'.
success_message: ✅ Great work including 'and' with the last string.


Tests for task_03 (points: 3)
```
>>> isinstance(with_commas, str) and isinstance(without_commas, str)
True
```
failure_message: ❌ I doesn't seem that with_commas and without_commas are strings. Make sure to use the join method on a string. Use ', ' for a comma and ' ' for a space.
success_message: ✅ Nice! with_commas and without_commas are both strings.


Tests for task_05 (points: 2)
```
>>> third_element != -10
True
```
failure_message: ❌ It looks like you wrote some_numbers.item(3), but the third element has index 2, not 3.
success_message: ✅ Nice work not using the wrong index.


Tests for task_06 (points: 2)
```
>>> isinstance(index_of_last_element, (int, np.integer))
True
```
failure_message: ❌ It doesn't seem like you've assigned an integer to index_of_last_element.
success_message: ✅ You assigned an integer to index_of_last_element.
```
>>> index_of_last_element <= 145
True
```
failure_message: ❌ index_of_last_element doesn't represent a possible index value.
success_message: ✅ index_of_last_element represents a possible index value.


Tests for task_07 (points: 2)
```
>>> import numpy as np
>>> isinstance(most_recent_birth_year, (int, np.integer))
True
```
failure_message: ❌ Make sure you are assigning an integer from the president_birth_years array. Try using .item(...)
success_message: ✅ You successfully submitted an integer.


Tests for task_08 (points: 2)
```
>>> import numpy as np
>>> isinstance(sum_of_birth_years, (int, np.integer))
True
```
failure_message: ❌ Make sure you are assigning a sum of 3 integers from the sum_of_birth_years array.
success_message: ✅ You successfully submitted an integer.


Tests for task_09 (points: 2)
```
>>> all((x % 157 == 0 for x in (first_product, second_product, third_product, fourth_product)))
True
```
failure_message: ❌ Not all of your products are divisible by 157, so it doesn't seem that you've multiplied the 4 numbers by 157.
success_message: ✅ It seems like you've multiplied the 4 numbers by 157.


Tests for task_10 (points: 2)
```
>>> import numpy
>>> isinstance(numbers, numpy.ndarray)
True
```
failure_message: ❌ Make sure you are assigning an array to numbers. Try using make_array()
success_message: ✅ Good work. It seems like you've assigned an array to numbers.


Tests for task_11 (points: 2)
```
>>> import numpy as np
>>> isinstance(correct_products, np.ndarray) and all((x % 1577 == 0 for x in correct_products))
True
```
failure_message: ❌ All the values in your array should be divible by 1577.
success_message: ✅ Good work. It seems like all the values in your array are divible by 1577.


Tests for task_12 (points: 3)
```
>>> sum(celsius_max_temperatures) != 356705.0
np.True_
```
failure_message: ❌ Looks like you multiplied and subtracted in the wrong order.
success_message: ✅ Nice work. It is easy to multiply and subtract in the wrong order by mistake.
```
>>> import numpy as np
>>> celsius_max_temperatures[0] - np.floor(celsius_max_temperatures[0]) == 0 and celsius_max_temperatures[0] != 25
np.True_
```
failure_message: ❌ Make sure you round your results to the nearest integer. Try using np.round.
success_message: ✅ It seems like you've rounded your results to the nearest integer.


Tests for task_13 (points: 3)
```
>>> abs(celsius_temperature_ranges[0]) != 12 and celsius_temperature_ranges[0] >= 0
np.True_
```
failure_message: ❌ Make sure that you are not submitting Fahrenheit values.
success_message: ✅ It seems like you are not submitting Fahrenheit values. Nice work!
```
>>> import numpy as np
>>> celsius_temperature_ranges[0] - np.floor(celsius_temperature_ranges[0]) != 0
np.True_
```
failure_message: ❌ Make sure to not round your results on this task.
success_message: ✅ It doesn't seem that you rounded the results. Great work catching that detail in the instructions.


Tests for task_14 (points: 2)
```
>>> import numpy as np
>>> isinstance(largest_population_change, (int, np.integer))
True
```
failure_message: ❌ It doesn't seem that you are using the numpy functions.
success_message: ✅ Nice work using numpy functions.


Tests for task_15 (points: 2)
```
>>> cumulative_sum_answer in {1, 2, 3}
True
```
failure_message: ❌ Make sure to assign 1, 2, or 3 to cumulative_sum_answer.
success_message: ✅ It seems like you've entered a possible answer.


Tests for task_16 (points: 2)
```
>>> shortest <= average <= longest
np.True_
```
failure_message: ❌ The average value is between the shortest and the longest values. Make sure to use functions like min, max, and np.average (or np.mean).
success_message: ✅ Your average value is between your shortest and longest values.


Tests for task_17 (points: 2)
```
>>> biggest_decrease != 47
np.True_
```
failure_message: ❌ If you are getting 47 as your answer, you might be computing the biggest change rather than biggest decrease.
success_message: ✅ It seems that you didn't compute the biggest change. Nice work.
```
>>> 30 <= biggest_decrease < 47
np.True_
```
failure_message: ❌ The biggest decrease is above 30, but not 47
success_message: ✅ Your value seems to be in a reasonable interval of values.
