<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 6: Testing Hypotheses

## References

* [Chapter 10](https://ccsf-math-108.github.io/textbook/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Sections 11.0 -11.2](https://ccsf-math-108.github.io/textbook/chapters/11/Testing_Hypotheses.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Simulation with Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ..., 36. Pockets 0 and 00 are green, and the other pockets are alternating red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="./roulette_wheel.jpeg" alt="roulette wheel" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

---

### Betting on Red

If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

---

#### Task 01 📍

Define a function `dollar_bet_on_red`. The function definition should:
1. Have one argument `color` that is a `str` for the name of a color.
2. Return your gain in dollars as an `int` if that color had won and you had placed a one-dollar bet on red. 

Consider the following as you work:
* Remember that the gain can be negative.
* Make sure your function returns an integer.
* You can assume that the only colors that will be passed as arguments are `'red'`, `'black'`, and `'green'`. Your function doesn't have to check that the input is correct.

_Points:_ 3

```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("task_01")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

---

#### Task 02 📍

1. Add a column labeled `'Winnings: Red'` as the last column in the table `wheel` such that, for each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. 
2. Your code should apply the function `dollar_bet_on_red` on the `wheel` table to create an array `red_winnings`.

_Points:_ 3

```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("task_02")
```

---

### Simulating 10 Bets on Red

Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

---

#### Task 03 📍

1. Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. 
2. Your table should have the same three column labels as in `wheel`. 
3. Once you've created that table, set `sum_bets` to your net gain, an integer, in all 10 bets, assuming that you bet one dollar on red each time.

While you are working on this, it may be helpful to print out `ten_bets` after you create it!

_Points:_ 5

```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("task_03")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

---

#### Task 04 📍

Let's see what would happen if you made more bets. 

Define a function `net_gain_red` with one argument that takes the number of bets (an integer) and returns the net gain in that number of one-dollar bets on red (an integer).

Reference the `wheel` table that you've defined in the notebook within your function.

_Points:_ 2

```python
def net_gain_red(...):
    ...
    
net_gain_red(10)
```

```python
grader.check("task_04")
```

---

#### Task 05 📍

1. Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. 
2. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

_Points:_ 2

```python
num_bets = ...
repetitions = ...

all_gains_red = ...
...

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

```python
grader.check("task_05")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

---

#### Task 06 📍

Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.

_Points:_ 2

```python
loss_more_than_50 = ...
```

```python
grader.check("task_06")
```

---

### Betting on a Split

If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bets pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

---

#### Task 07 📍

Define a function `dollar_bet_on_5_6_split`.
1. The function should have 1 argument (type `str`) that represents the pocket number
2. The function should return the gain in dollars if that pocket won and they had bet one dollar on the 5-6 split.

Remember that the pockets are represented as strings.

_Points:_ 4

```python
def dollar_bet_on_5_6_split(...):
    ...
```

```python
grader.check("task_07")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_5_6_split('5'))
print(dollar_bet_on_5_6_split('6'))
print(dollar_bet_on_5_6_split('00'))
print(dollar_bet_on_5_6_split('23'))
```

---

#### Task 08 📍

Add a column `'Winnings: 5-6 Split'` to the end of the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.

_Points:_ 3

```python
split_winnings = ...
wheel = ...
```

```python
grader.check("task_08")
```

---

#### Task 09 📍

1. Simulate the net gain in 200 one-dollar bets on the 5-6 split.
2. Repeat the simulation for a total of 10,000 times.
3. Store your gains in the array `all_gains_split`.

_Points:_ 2

```python
all_gains_split = ...
...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', all_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("task_09")
```

---

#### Task 10 📍

Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram.

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

Notice that you've already seen one of these statements in a prior question.

_Points:_ 5

```python
histogram_statements = ...
```

```python
grader.check("task_10")
```

If this exercise has put you off playing roulette, it has done its job. If you are still curious about other bets, [here](https://en.wikipedia.org/wiki/Roulette#Bet_odds_table) they all are, and [here](https://en.wikipedia.org/wiki/Roulette#House_edge) is the bad news. The house – that is, the casino – always has an edge over the gambler.

---

## Using TVD as a Test Statistic

Before beginning this section, please read [this section](https://inferentialthinking.com/chapters/11/2/Multiple_Categories.html#a-new-statistic-the-distance-between-two-distributions) of the textbook on TVD!

**Total variation distance (TVD)** is a special type of test statistic that we use when we want to compare two distributions of *categorical data*. It is often used when we observe that a set of observed proportions/probabilities is different than what we expect under the null model. 

Consider a six-sided die that we roll 6,000 times. If the die is fair, we would expect that each face comes up $\frac{1}{6}$ of the time. By random chance, a fair die won't always result in equal proportions (that is, we won't get exactly 1,000 of each face). However, if we suspect that the die might be unfair based on the data, we can conduct a hypothesis test using TVD to compare the expected [$\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$] distribution to what is actually observed.

In this part of the homework, we'll look at how we can use TVD to determine the effect that different factors have on happiness. 

We will be working with data from the [Gallup World Poll](https://www.gallup.com/analytics/349487/gallup-global-happiness-center.aspx#:~:text=World%20Happiness%20Report&text=Using%20the%20Gallup%20World%20Poll,about%20the%20World%20Happiness%20Report.) that is presented in the World Happiness Report, a survey of the state of global happiness. The survey ranked 155 countries by overall happiness and estimated the influence that economic production, social support, life expectancy, freedom, absence of corruption, and generosity had on population happiness. The study has been repeated for several years, but we'll be looking at data from the 2016 survey.

Run the cell below to load in the `happiness_scores` table.

```python
happiness_scores = Table.read_table("happiness_scores.csv")
happiness_scores.show(5)
```

Participants in the study were asked to evaluate their life satisfaction from a scale of 0 (worst possible life) to 10 (best possible life). The responses for each country were averaged to create the `Happiness Score`.

The columns `Economy (GDP per Capita)`, `Family`, `Health (Life Expectancy)`, `Freedom`, `Trust (Government Corruption)`, and `Generosity` estimate the extent to which each factor influences happiness, both for better or for worse. The happiness score is the sum of these factors; the larger a factor is, the more it contributes to overall happiness. [In other words, if you add up all the factors (in addition to a "Difference from Dystopia" value we excluded in the dataset), you get the happiness score.]

Let's look at the different factors that affect happiness in the United States. Run the cell below to view the row in `us_happiness` that contains data for the United States.

```python
us_happiness = happiness_scores.where("Country", "United States")
us_happiness
```

**To compare the different factors, we'll look at the proportion of the happiness score that is attributed to each variable. 
You can find these proportions in the table `us_happiness_factors` after running the cell below.**

*Note:* The factors shown in `us_happiness` don't add up exactly to the happiness score, so we adjusted the proportions to  only account for the data we have access to. The proportions were found by dividing each Happiness Factor value by the sum of all Happiness Factor values in `us_happiness`.

```python
us_happiness_factors = Table().read_table("us_happiness_factors.csv")
us_happiness_factors
```

---

### Task 11 📍

Suppose we want to test whether or not each factor contributes the same amount to the overall Happiness Score. If we want to run a simulation, under which model below (choose one) must we simulate? Assign `task_11_answer` to the integer value of the best choice.

1. Some factors contribute more to the happines score than other factors.
2. Each factor contributes an equal amount to the happiness score. Any deviation is due to random chance.

_Points:_ 2

```python
task_11_answer = ...
```

```python
grader.check("task_11")
```

---

### Task 12 📍

Write a function `calculate_tvd` that takes in the observed distribution (`obs_dist`) and expected distribution under the null hypothesis (`null_dist`) and calculates the total variation distance. Use this function to set `observed_tvd` to be equal to the observed test statistic.

_Points:_ 2

```python
null_distribution = make_array(1/6, 1/6, 1/6, 1/6, 1/6, 1/6)

def calculate_tvd(obs_dist, null_dist):
    ...
    
observed_tvd = ...
observed_tvd
```

```python
grader.check("task_12")
```

---

### Task 13 📍

Create an array called `simulated_tvds` that contains 10,000 simulated values under the null hypothesis. Assume that the original sample consisted of 1,000 individuals.

*Hint:* The `sample_proportions` function may be helpful to you. Refer to the Python Reference Guide in Canvas.

```python
simulated_tvds = ...

...
```

Run the cell below to plot a histogram of your simulated test statistics, as well as a red dot representing the observed value of the test statistic.

_Points:_ 4

```python
Table().with_column("Simulated TVDs", simulated_tvds).hist()
plt.scatter(observed_tvd, 0, color='red', s=70, zorder=2);
plt.show();
```

```python
grader.check("task_13")
```

---

### Task 14 📍

Use your simulated statistics to calculate the proporion of the simulation TVDs of your test that are more at least as big as the observed TVD. Make sure that this number is consistent with what you observed in the histogram above.

_Points:_ 1

```python
proportion_tvd = ...
proportion_tvd
```

```python
grader.check("task_14")
```

---

### Task 15 📍

Using the results of our simulation, what can you conclude about how each factor contributes to the overall happiness score in the US? Assign `conclusion` to the integer value of the best choice below.

1. The data are more consistent with the claim that the factors do not equally contribute to the overall happiness score of a country.
2. The data are more consistent with the claim that the factors do equally contribute to the overall happiness score of a country.

_Points:_ 2

```python
conclusion = ...
```

```python
grader.check("task_15")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

1. **Review the Rubric:** View the rubric on the Canvas assignment page to understand the scoring criteria.  
2. **Review Manually Graded Tasks:** Verify that you have responded to all manually graded tasks marked with 📍🔎.  
3. **Review Auto-Graded Tasks:** Inspect the output from the `grade.check_all()` command above. Ensure you have passed all available tests for the auto-graded tasks in this notebook marked with 📍.  
   - This command executes all auto-grader tests sequentially.  
   - Note that Homework and Project assignments include hidden auto-grader tests that are not visible to you. Passing all visible tests does **not** guarantee that your solution is fully correct.  
4. **Save Your Work:** In the notebook toolbar, go to `File → Save Notebook` to save your work and create a checkpoint.  
5. **Download the Notebook:** In the notebook toolbar, go to `File → Download IPYNB` to download the notebook (`.ipynb`) file.  
6. **Upload to Canvas:** Upload the downloaded `.ipynb` file to the appropriate Canvas assignment page.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 3)
```
>>> dollar_bet_on_red('red') == 1
True
```
failure_message: ❌ Your function should return 1 when it is called on 'red'.
success_message: ✅ Your function correctly returns 1 when it is called on 'red'.


Tests for task_02 (points: 3)
```
>>> isinstance(red_winnings, np.ndarray)
True
```
failure_message: ❌ red_winnings should be an array.
success_message: ✅ It seems that red_winnings is an array.
```
>>> wheel.labels[2] == 'Winnings: Red'
True
```
failure_message: ❌ wheel is not a table with the correct labels.
success_message: ✅ wheel is a table with the correct labels.


Tests for task_03 (points: 5)
```
>>> ten_bets.num_rows == 10
True
```
failure_message: ❌ ten_bets should have 10 rows.
success_message: ✅ ten_bets has 10 rows.
```
>>> set(['Pocket', 'Color', 'Winnings: Red']) == set(ten_bets.labels)
True
```
failure_message: ❌ ten_bets should have the same column labels as wheels.
success_message: ✅ ten_bets has the same column labels as wheels.
```
>>> isinstance(sum_bets, int) or isinstance(sum_bets, np.int64)
True
```
failure_message: ❌ sum_bets should be an integer.
success_message: ✅ sum_bets is an integer.


Tests for task_04 (points: 2)
```
>>> net_gain_red(10000) != net_gain_red(10000)
np.True_
```
failure_message: ❌ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.
success_message: ✅ There is essentially no chance that the net gain for 2 rounds of 10,000 bets should the same.


Tests for task_05 (points: 2)
```
>>> len(all_gains_red) == 10000
True
```
failure_message: ❌ all_gains_red should be an array with 10,000 items.
success_message: ✅ all_gains_red is an array with 10,000 items.


Tests for task_06 (points: 2)
```
>>> isinstance(loss_more_than_50, bool)
True
```
failure_message: ❌ loss_more_than_50 should be a bool value.
success_message: ✅ loss_more_than_50 is a bool value.


Tests for task_07 (points: 4)
```
>>> dollar_bet_on_5_6_split('5') == 17
True
```
failure_message: ❌ Your function didn't return the correct value for the 5 pocket.
success_message: ✅ Your function returned the correct value for the 5 pocket.
```
>>> dollar_bet_on_5_6_split('00') == -1
True
```
failure_message: ❌ Your function didn't return the correct value for the 00 pocket.
success_message: ✅ Your function returned the correct value for the 00 pocket.


Tests for task_08 (points: 3)
```
>>> wheel.labels[3] == 'Winnings: 5-6 Split'
True
```
failure_message: ❌ Your last column label should be 'Winnings: 5-6 Split'.
success_message: ✅ Your last column label is 'Winnings: 5-6 Split'.
```
>>> wheel.where('Pocket', '5').column('Winnings: 5-6 Split').item(0) == 17
True
```
failure_message: ❌ The added column value associated with the 5 pocket is not correct.
success_message: ✅ The added column value associated with the 5 pocket is correct.


Tests for task_09 (points: 2)
```
>>> len(all_gains_split) == 10000
True
```
failure_message: ❌ There should be 10,000 items in all_gains_split.
success_message: ✅ There are 10,000 items in all_gains_split.


Tests for task_10 (points: 5)
```
>>> isinstance(histogram_statements, np.ndarray)
True
```
failure_message: ❌ histogram_statements should be an array.
success_message: ✅ histogram_statements is an array.
```
>>> isinstance(histogram_statements.item(0), int)
True
```
failure_message: ❌ histogram_statements should contain at least one integer.
success_message: ✅ histogram_statements contains at least one integer.


Tests for task_11 (points: 2)
```
>>> task_11_answer in [1, 2]
True
```
failure_message: ❌ It doesn't seem like you have entered a valid answer.
success_message: ✅ You have entered a valid choice.


Tests for task_12 (points: 2)
```
>>> isinstance(observed_tvd, (float, np.float64, np.float32))
True
```
failure_message: ❌ observed_tvd should be a float.
success_message: ✅ observed_tvd is a float.
```
>>> np.isclose(round(observed_tvd, 6), 0.262086)
np.True_
```
failure_message: ❌ observed_tvd doesn't seem correct.
success_message: ✅ observed_tvd seems correct.


Tests for task_13 (points: 4)
```
>>> len(simulated_tvds) == 10000
True
```
failure_message: ❌ simulated_tvds should have 10,000 items in it.
success_message: ✅ simulated_tvds has 10,000 items in it.
```
>>> all(simulated_tvds >= 0)
True
```
failure_message: ❌ Each value in simulated_tvds should be 0 or larger.
success_message: ✅ Each value in simulated_tvds is 0 or larger.
```
>>> len(np.unique(simulated_tvds)) != 1
True
```
failure_message: ❌ None of the values in simulated_tvds should be the same.
success_message: ✅ None of the values in simulated_tvds are the same.
```
>>> np.mean(simulated_tvds) < 0.1
np.True_
```
failure_message: ❌ It doesn't seem like you are simulating under the null hypothesis.
success_message: ✅ It seems like you are simulating under the null hypothesis.


Tests for task_14 (points: 1)
```
>>> 0 <= proportion_tvd <= 1
np.True_
```
failure_message: ❌ proportion_tvd should be a number between 0 and 1.
success_message: ✅ proportion_tvd is a number between 0 and 1.


Tests for task_15 (points: 2)
```
>>> conclusion in [1, 2]
True
```
failure_message: ❌ It doesn't seem like you have entered a valid answer.
success_message: ✅ You have entered a valid choice.
