<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 05: Simulation and Chance

## References

* [Sections 8.4 - 8.5](https://ccsf-math-108.github.io/textbook/chapters/08/4/Joining_Tables_by_Columns.html)  
* [Sections 9.0 - 9.5](https://ccsf-math-108.github.io/textbook/chapters/09/Randomness.html)
* [datascience Documentation](https://datascience.readthedocs.io/)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab05.ipynb')
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Joining Tables

There are a  variety of reasons why data might be stored across multiple tables, and it might be necessary to combine data from 2 or more sources into one table. Think of joining tables like mixing ingredients to create a delicious recipe! When we use methods like `join` in the `datascience` Python library, we're basically combining different sets of data based on something they have in common, just like adding various ingredients to make a yummy dish. This helps us get a bigger picture of the data and discover exciting insights by bringing together information from different sources, making data analysis a bit like cooking up a data stew!

For example, you might have customer data in two different tables:

* A table called `customers` containing contact information.
* A table called `purchase_history` containing information on the customer's latest purchases.

Run the following code cell to import and preview the two tables.

```python
customers = Table().read_table('customers.csv')

purchase_history = Table().read_table('purchase_history.csv')

display(customers)
display(purchase_history)
```

You might want to put all this information in one table to more easily analyze the data or follow up with a customer about one of their purchases. The key to joining data is that there must be a clearly defined relation between the two tables. In this case, the link between the tables is the customer identification number. 

To use the join method from the `datascience` library, you'll typically have two tables, let's call them `tbl1` and `tbl2`, and you want to combine them based on a common key. Let's say that the key values are located in `'Column 1'` for `tbl1` and `'Column 2'` for `tbl2`.

You could use `tbl1.join('Column 1', tbl2, 'Column 2')` to join these tables together, aligning rows with matching key values in `'Column 1'` and `'Column 2'` into a single table where the focus is placed on the column of `tbl1`.
If you used `tbl2.join('Column 2', tb1, 'Column 1')`, you would also create a joined table, but the output would likely be different as the table would be created from the perspective of the column in `tbl2`.

---

### Task 01 📍

Complete the following code to join the data in the `purchase_history` table with the data in the `customers` table. The resulting table should have 6 rows and the 6 columns `'Customer ID', 'Name', 'Email', 'Product', 'Ad', 'Amount ($)'`.

```python
customer_data = customers.join(..., ..., ...)
customer_data
```

```python
grader.check("task_01")
```

---

There are a few things to notice about how this `join` method worked. 
* Since you started with `customers`, it kept all the column labels of the `customers` table.
* The `'Customer ID'` column was brought to the front of the table.
* The columns in `purchase_history` that were not used to join the data (`'Product'`, `'Ad'`, and `'Amount ($)'`) were added to the end of the table.
* The label `'Customer Number'` doesn't appear in the resulting table.
* If any customer was missing an email then that information would be missing from the joined table.
* If any customer had not made any purchases then they would be missing from the joined table.

---

## Nachos and Conditionals

In Python, the boolean data type contains only two unique values:  `True` and `False`. Expressions containing comparison operators such as `<` (less than), `>` (greater than), and `==` (equal to) evaluate to Boolean values. A list of common comparison operators can be found below!

| Comparison            | Operator | True example | False example |
|-----------------------|----------|--------------|---------------|
| Less than             | <        | 2 < 3        | 2 < 2         |
| Great than            | >        | 3 > 2        | 3 > 3         |
| Less than or equal    | <=       | 2 <= 2       | 3 <= 2        |
| Greater than or equal | >=       | 3 >= 3       | 2 >= 3        |
| Equal                 | ==       | 3 == 3       | 3 == 2        |
| Not equal             | !=       | 3 != 2       | 2 != 2        |

Run the cell below to see an example of a comparison operator in action.

```python
3 > 1 + 1
```

We can even assign the result of a comparison operation to a variable.

```python
result = 10 / 2 == 5
result
```

Arrays are compatible with comparison operators. The output is an array of boolean values.

```python
make_array(1, 5, 7, 8, 3, -1) > 3
```

One day, when you come home after a long week, you see a hot bowl of nachos waiting on the dining table! Let's say that whenever you take a nacho from the bowl, it will either have only **cheese**, only **salsa**, **both** cheese and salsa, or **neither** cheese nor salsa (a sad tortilla chip indeed). 

Let's try and simulate taking nachos from the bowl at random using the function, `np.random.choice(...)`.

---

### `np.random.choice`

`np.random.choice` picks one item at random from the given array. It is equally likely to pick any of the items. Run the cell below several times, and observe how the results change.

```python
nachos = make_array('cheese', 'salsa', 'both', 'neither')
np.random.choice(nachos)
```

To repeat this process multiple times, pass in an int `n` as the second argument to return `n` different random choices. By default, `np.random.choice` samples **with replacement** and returns an *array* of items. 

Run the next cell to see an example of sampling with replacement 10 times from the `nachos` array.

```python
np.random.choice(nachos, 10)
```

To count the number of times a certain type of nacho is randomly chosen, we can use `np.count_nonzero`

---

### `np.count_nonzero`

`np.count_nonzero` counts the number of non-zero values that appear in an array. When an array of boolean values are passed through the function, it will count the number of `True` values (remember that in Python, `True` is coded as 1 and `False` is coded as 0.)

Run the next cell to see an example that uses `np.count_nonzero`.

```python
np.count_nonzero(make_array(True, False, False, True, True))
```

---

### Task 02 📍

Assume we took ten nachos at random and stored the results in an array called `ten_nachos` as done below. Find the number of nachos with only cheese using code (do not hardcode the answer).  

*Hint:* Our solution involves a comparison operator (e.g. `==`, `<`, ...) and the `np.count_nonzero` method.


```python
ten_nachos = make_array('neither', 'cheese', 'both', 'both', 'cheese', 'salsa', 'both', 'neither', 'cheese', 'both')
number_cheese = ...
number_cheese
```

```python
grader.check("task_02")
```

---

### Conditional Statements

A conditional statement is a multi-line statement that allows Python to choose among different alternatives based on the truth value of an expression.

Here is a basic example.

```
def sign(x):
    if x > 0:
        return 'Positive'
    else:
        return 'Negative'
```

If the input `x` is greater than `0`, we return the string `'Positive'`. Otherwise, we return `'Negative'`.

If we want to test multiple conditions at once, we use the following general format.

```
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```

Only the body for the first conditional expression that is true will be evaluated. Each `if` and `elif` expression is evaluated and considered in order, starting at the top. As soon as a true value is found, the corresponding body is executed, and the rest of the conditional statement is skipped. If none of the `if` or `elif` expressions are true, then the `else body` is executed. 

For more examples and explanation, refer to the [section on conditional statements in the course textbook](https://ccsf-math-108.github.io/textbook/chapters/09/1/Conditional_Statements.html).

---

### Task 03 📍

Complete the following conditional statement so that the string `'More please'` is assigned to the variable `say_please` if the number of nachos with cheese in `ten_nachos` is less than `5`.

*Hint*: You should be using `number_cheese` from a previous task.


```python
say_please = '?'

if ...:
    say_please = 'More please' 

say_please
```

```python
grader.check("task_03")
```

---

### Task 04 📍

Write a function called `nacho_reaction` that returns a reaction (as a string) based on the type of nacho passed in as an argument. Use the table below to match the nacho type to the appropriate reaction.

<img src="./nacho_reactions.png">

*Hint:* If you're failing the test, double check the spelling of your reactions.


```python
def nacho_reaction(nacho):
    if nacho == "cheese":
        ...
    ... :
        ...
    ... :
        ...
    ... :
        ...


spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

```python
grader.check("task_04")
```

---

### Task 05 📍

Create a table `ten_nachos_reactions` that consists of the nachos in `ten_nachos` as well as the reactions for each of those nachos. The columns should be called `Nachos` and `Reactions`.

*Hint:* Use the `apply` method. 


```python
ten_nachos_tbl = Table().with_column('Nachos', ten_nachos)
...
ten_nachos_reactions
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Using code, find the number of 'WOW' reactions for the nachos in `ten_nachos_reactions`.


```python
number_wow_reactions = ...
number_wow_reactions
```

```python
grader.check("task_06")
```

---

## Simulations and For Loops

Using a `for` statement, we can perform a task multiple times. This is known as iteration.

One use of iteration is to loop through a set of values. For instance, we can print out all of the colors of the rainbow.

```python
rainbow = make_array("red", "orange", "yellow", "green", "blue", "indigo", "violet")

for color in rainbow:
    print(color)
```

We can see that the indented part of the `for` loop, known as the body, is executed once for each item in `rainbow`. The name `color` is assigned to the next value in `rainbow` at the start of each iteration. Note that the name `color` is arbitrary; we could easily have named it something else. The important thing is we stay consistent throughout the `for` loop. 

```python
for another_name in rainbow:
    print(another_name)
```

In general, however, we would like the variable name to be somewhat informative. 

---

### Task 07 📍

In the following cell, we've loaded the text of _Pride and Prejudice_ by Jane Austen, split it into individual words, and stored these words in an array `p_and_p_words`. Using a `for` loop, assign `longer_than_five` to the number of words in the novel that are more than 5 characters long.

*Hint*: You can find the number of letters in a word with the `len` function.

*Note*: You should expect "words" like `"About` to be included in this collection of words that are more than 5 characters long because the quotation mark is not removed with the `split` function.


```python
austen_string = open('Austen_PrideAndPrejudice.txt', encoding='utf-8').read()
p_and_p_words = np.array(austen_string.split())

longer_than_five = ...

# a for loop would be useful here



longer_than_five
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Using a simulation with 10,000 trials, assign num_different to the number of times, in 10,000 trials, that two words picked uniformly at random (with replacement) from Pride and Prejudice have different lengths. 

*Hint 1*: What function did we use in section 1 to sample at random with replacement from an array? 

*Hint 2*: Remember that `!=` checks for non-equality between two items.


```python
trials = 10000
num_different = ...

for ... in ...:
    ...

num_different
```

```python
grader.check("task_08")
```

---

## Probability

### Electric Board

An electronic board with lights contains 10 equal-sized zones labeled with values from 1 to 10.

<img src="./number_grid.png" alt="a 2x5 grid of numbers 1 - 10" width=40%></img>

The board illuminates a sequence of numbers at random. First, a number is lit up with a pink color, then a number is lit up with an orange color where the second number is different from the first. For example, the following image shows 6 lit up with a pink color followed by 5 lit up with an orange color.

<img src="./number_grid_color.png" alt="a 2x5 grid of numbers 1 - 10 where 6 is pink and 5 is orange" width=40%></img>

After the sequence of two numbers is shown, the board resets.

<img src="./number_grid.png" alt="a 2x5 grid of numbers 1 - 10" width=40%></img>

Use this situation to practice with the basic ideas of probability that you've learned so far.

### A Definition of Probability

Given an experiment where all possible outcomes are equally likely, the probability of an event occurring is calculated by dividing the number of ways that event can occur by the total number of possible outcomes.

---

### Task 09 📍

A sequence of two numbers is illuminated on the board as described above. What is the probability that the first number in the sequence is 6?

Assign `chance_6` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_6 = ...
chance_6
```

```python
grader.check("task_09")
```

---

### The Multiplication Rule

To calculate the probability of one event happening and then another event happening, multiply the probability of the first event by the probability of the second event given that the first event happened. This is generally referred to as the multiplication rule.

---

### Task 10 📍

Consider the sequence of two numbers illuminated on this board as described above (first the 6, and then the 5). What is the probability that this sequence (6, 5) occurs?

Assign `chance_65` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_65 = ...
chance_65
```

```python
grader.check("task_10")
```

---

### Task 11📍

Suppose another sequence of two numbers is illuminated on this board. What is the probability that the sequence is (6, 6)?

Assign `chance_66` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_66 = ...
chance_66
```

```python
grader.check("task_11")
```

---

### Task 12 📍

If all possible sequences are equally likely, what is the chance that any sequence of two numbers (where the two numbers are different) is illuminated on the board?

Assign `chance_of_any_sequence` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_of_any_sequence = ...
chance_of_any_sequence
```

```python
grader.check("task_12")
```

---

### The Addition Rule for Disjoint Events

To calculate the probability of one event or another event happening, where the two events are disjoint (or mutually exclusive of one another), we calculate the probability of each event happening and then add the values. This is called the addition rule for disjoint events.

---

### Task 13 📍

Suppose another sequence of two numbers is illuminated on this board. What is the probability that the sequence is either (6, 5) or (5, 6)?

Assign `chance_65_56` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_65_56 = ...
chance_65_56
```

```python
grader.check("task_13")
```

---

### The Complement Rule

One way to determine the probability of an event not occurring is to subtract the probability that it will occur from 1. This is generally referred to as the complement rule.

---

### Task 14 📍

Twenty sequences of two numbers are illuminated on the board. What is the probability that at least one of the sequences is (5, 6)?

Assign `chance_at_least_one_56` to a value or expression that represents this probability value. The value should be a number between `0` and `1`, inclusive.

```python
chance_not_56 = ...
chance_at_least_one_56 = ...
chance_at_least_one_56
```

```python
grader.check("task_14")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: None)
```
>>> isinstance(customer_data, Table)
True
```
failure_message: ❌ customer_data should be a Table. Are you using the join method?
success_message: ✅ customer_data is a Table.
```
>>> customer_data.labels == ('Customer ID', 'Name', 'Email', 'Product', 'Ad', 'Amount ($)')
True
```
failure_message: ❌ customer_data doesn't have the correct labels. Make sure you start your join statement with customers.join because purchase_history.join would produce a different collection of labels.
success_message: ✅ customer_data has the correct labels.
```
>>> customer_data
Customer ID | Name          | Email             | Product     | Ad   | Amount ($)
1           | Alice Johnson | alice@example.com | Gadget      | No   | 28
1           | Alice Johnson | alice@example.com | Doodad      | No   | 16
1           | Alice Johnson | alice@example.com | Widget      | No   | 28
1           | Alice Johnson | alice@example.com | Doodad      | No   | 21
1           | Alice Johnson | alice@example.com | Thingamajig | No   | 27
1           | Alice Johnson | alice@example.com | Widget      | No   | 21
1           | Alice Johnson | alice@example.com | Gadget      | No   | 20
1           | Alice Johnson | alice@example.com | Widget      | No   | 30
1           | Alice Johnson | alice@example.com | Thingamajig | No   | 15
1           | Alice Johnson | alice@example.com | Thingamajig | No   | 17
... (389 rows omitted)
```
failure_message: ❌ customer_data doesn't look correct. Check with someone before moving on.
success_message: ✅ customer_data seems correct!


Tests for task_02 (points: None)
```
>>> number_cheese == 3
np.True_
```
failure_message: ❌ There are 3 nachos with only cheese. Try using np.count_nonzero.
success_message: ✅ There are 3 nachos with only cheese.


Tests for task_03 (points: None)
```
>>> say_please == 'More please'
True
```
failure_message: ❌ Since number_cheese < 5, say_please should be assigned to 'More please'. Check if number_cheese < 5.
success_message: ✅ say_please is assigned to 'More please'.


Tests for task_04 (points: None)
```
>>> nacho_reaction('salsa') == 'SPICY'
True
```
failure_message: ❌ nacho_reaction('salsa') should display 'SPICY' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('salsa') should display 'SPICY'.
```
>>> nacho_reaction('cheese') == 'CHEESY'
True
```
failure_message: ❌ nacho_reaction('cheese') should display 'CHEESY' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('cheese') should display 'CHEESY'.
```
>>> nacho_reaction('both') == 'WOW'
True
```
failure_message: ❌ nacho_reaction('both') should display 'WOW' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('both') should display 'WOW'.
```
>>> nacho_reaction('neither') == 'MEH'
True
```
failure_message: ❌ nacho_reaction('neither') should display 'MEH' You might need to check your spelling/punctuation.
success_message: ✅ nacho_reaction('neither') should display 'MEH'.


Tests for task_05 (points: None)
```
>>> np.count_nonzero(ten_nachos_reactions.column('Reactions') == make_array('MEH', 'CHEESY', 'WOW', 'WOW', 'CHEESY', 'SPICY', 'WOW', 'MEH', 'CHEESY', 'WOW')) == 10
np.True_
```
failure_message: ❌ It seems like one or more of the reaction results could be incorrect. To get the values for the table, trying using the apply method with the function you just made.
success_message: ✅ The reaction values in the table seem correct.


Tests for task_06 (points: None)
```
>>> 2 < number_wow_reactions < 6
np.True_
```
failure_message: ❌ number_wow_reactions should be between 2 and 6, non-inclusive. Try using np.count_nonzero.
success_message: ✅ number_wow_reactions is between 2 and 6, non-inclusive.
```
>>> number_wow_reactions == 4
np.True_
```
failure_message: ❌ You have the incorrect value for number_wow_reactions.
success_message: ✅ You have the correct value for number_wow_reactions.


Tests for task_07 (points: None)
```
>>> longer_than_five == 35453
True
```
failure_message: ❌ The value assigned to longer_than_five is not correct. Make sure to initilalize longer_than_five to 0 and check if the length of the word is great than 5 as you loop through each word in p_and_p_words.
success_message: ✅ The value assigned to longer_than_five is correct.


Tests for task_08 (points: None)
```
>>> 8100 <= num_different <= 9100
True
```
failure_message: ❌ num_different should be between 8100 and 9100. Try initializing num_different to 0, iterating through the number of trials, randomly choosing 2 words from p_and_p_words, checking the length of the 2 words, and updating num_different depending on the results of that check.
success_message: ✅ num_different represents a value betwen 8100 and 9100.


Tests for task_09 (points: None)
```
>>> 0 <= chance_6 <= 1
True
```
failure_message: ❌ chance_6 should be a number between 0 and 1, inclusive.
success_message: ✅ chance_6 is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(chance_6, 4))
'cb5ae17636e975f9bf71ddf5bc542075'
```
failure_message: ❌ Your value for chance_6 is not correct! How many squares are there? How many squares are a 6?
success_message: ✅ Your value for chance_6 is correct!


Tests for task_10 (points: None)
```
>>> 0 <= chance_65 <= 1
True
```
failure_message: ❌ chance_65 should be a number between 0 and 1, inclusive.
success_message: ✅ chance_65 is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(chance_65, 4))
'2e55d74f5c78981f6b877b198bdc61ba'
```
failure_message: ❌ Your value for chance_65 is not correct! What is the chance the first number is 6 AND the chance the second number is 5? Remember that the second number cannot be the same as the first.
success_message: ✅ Your value for chance_65 is correct!


Tests for task_11 (points: None)
```
>>> 0 <= chance_66 <= 1
True
```
failure_message: ❌ chance_66 should be a number between 0 and 1, inclusive.
success_message: ✅ chance_66 is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(float(round(chance_66, 4)))
'30565a8911a6bb487e3745c0ea3c8224'
```
failure_message: "❌ Your value for chance_66 is not correct! Remember that the second number cannot be the same as the first."
success_message: "✅ Your value for chance_66 is correct!"



Tests for task_12 (points: None)
```
>>> 0 <= chance_of_any_sequence <= 1
True
```
failure_message: ❌ chance_of_any_sequence should be a number between 0 and 1, inclusive.
success_message: ✅ chance_of_any_sequence is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(chance_of_any_sequence, 4))
'2e55d74f5c78981f6b877b198bdc61ba'
```
failure_message: "❌ Your value for chance_of_any_sequence is not correct! Remember that the second number cannot be the same as the first, so the probabilities change for the second number in the sequence."
success_message: "✅ Your value for chance_of_any_sequence is correct!"



Tests for task_13 (points: None)
```
>>> 0 <= chance_65_56 <= 1
True
```
failure_message: ❌ chance_65_56 should be a number between 0 and 1, inclusive.
success_message: ✅ chance_65_56 is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(chance_65_56, 4))
'a3a4f732d111b1d6e677b91943de2625'
```
failure_message: ❌ Your value for chance_65_56 is not correct! This is like the previous task, but there is another possible option for a sequence.
success_message: ✅ Your value for chance_65_56 is correct!


Tests for task_14 (points: None)
```
>>> 0 <= chance_at_least_one_56 <= 1
True
```
failure_message: ❌ chance_at_least_one_56 should be a number between 0 and 1, inclusive.
success_message: ✅ chance_at_least_one_56 is a number between 0 and 1, inclusive.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(chance_at_least_one_56, 4))
'59990f2f6d79a1d4bf0fd216a24a87db'
```
failure_message: ❌ Your value for chance_at_least_one_56 is not correct! Consider the complementary event. First, try to determine the chance that (5, 6) will not be seen in one sequence.
success_message: ✅ Your value for chance_at_least_one_56 is correct!
