<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 11: Regression

## References
* [The Regression Line](https://ccsf-math-108.github.io/textbook/chapters/15/2/Regression_Line.html)
* [Method of Least Squares](https://ccsf-math-108.github.io/textbook/chapters/15/3/Method_of_Least_Squares.html)
* [Least Squares Regression](https://ccsf-math-108.github.io/textbook/chapters/15/4/Least_Squares_Regression.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw11.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Linear Regression Setup

---

When performing linear regression in this course, you will need to compute several important quantities which will be used throughout our analysis. Unless otherwise specified when asked to make a prediction please assume we are predicting $y$ from $x$ throughout this assignment. To help with our later analysis, we will begin by writing some of these functions and understanding what they can do for us.

---

### Task 01 📍

Define a function `standard_units` that converts a given array to standard units.

**Hint:** You may find the `np.mean` and `np.std` functions helpful.

_Points:_ 2

```python
def standard_units(data):
    ...
```

```python
grader.check("task_01")
```

---

### Task 02 📍

Which of the following are true about standard units? Assume we have converted an array of data into standard units using the function above.

1. The units of the data converted into standard units are the same as the units of the original data.
2. The sum of the data converted into standard units is 0.
3. The standard deviation of the data converted into standard units is 1.
4. Adding 5 to the original data has no impact on the resultant data when converted to standard units.
5. Multiplying the data by 5 has no impact on the resultant data when converted to standard units.

Assign `standard_array` to an array of your selections.

_Points:_ 2

```python
standard_array = make_array(...)
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Define a function `correlation` that computes the correlation between 2 arrays of data in original units.

**Hint:** Consider using the `standard_units` function you have defined previously.

_Points:_ 2

```python
def correlation(x, y):
    ...
```

```python
grader.check("task_03")
```

---

### Task 04 📍

Which of the following statements are true about the correlation coefficient?

1. The correlation coefficient measures the strength of a linear relationship.
2. A correlation coefficient of 1.0 means an increase in one variable always means an increase in the other variable by the same amount.
3. The correlation coefficient is the slope of the regression line in standard units.
4. The correlation coefficient stays the same if we swap our x-axis and y-axis.
5. If we add 5 to every value of our original data, our correlation coefficient will increase by 5.

Assign `r_array` to an array of your selections.

_Points:_ 2

```python
r_array = make_array(...)
```

```python
grader.check("task_04")
```

---

### Task 05 📍

Define a function `slope` that computes the slope of our line of best fit (to predict y given x), given two arrays of data in original units. Assume we want to create a line of best fit in original units.

**Hint:** Consider using the `correlation` function you have defined previously.

_Points:_ 2

```python
def slope(x, y):
    r = ...
    ...
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Which of the following are true statements about the slope of our line of best fit? Assume $x$ refers to the value of one variable that we use to predict the value of $y$.

1. In original units, the slope has as units: units of $x$ / units of $y$.
2. In standard units, the slope is unitless.
3. In original units, the slope is unchanged by swapping $x$ and $y$.
4. In standard units, a slope of 1 means our data is perfectly linearly correlated.
5. In both original units and standard units, the slope has the same sign (positive/negative).
 
Assign `slope_array` to an array of your selections.

_Points:_ 2

```python
slope_array = make_array(...)
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Define a function `intercept` that computes the intercept of our line of best fit (to predict `y` given `x`), given 2 arrays of data in original units. Assume we want to create a line of best fit in original units.

**Hint:** Consider using slope function you have defined previously.

_Points:_ 2

```python
def intercept(x, y):
    ...
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Which of the following are true about the intercept of our line of best fit? Assume $x$ refers to the value of one variable that we use to predict the value of $y$.

1. In original units, the intercept has the same unit as the $y$ values.
2. In original units, the intercept has the same unit as the $x$ values.
3. In original units, the slope and intercept have the same unit.
4. In standard units, the intercept for the regression line is 0.
5. The intercept value in original units is numerically the same as the intercept value in standard units.

Assign `intercept_array` to an array of your selections.

_Points:_ 2

```python
intercept_array = make_array(...)
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Define a function `predict` that takes in a table and 2 column names, and returns an array of predictions. Keep the following in mind when you make this function:
* The predictions should be created using a fitted regression line.
* We are predicting `y_col` from `x_col`, both in original units.

**Hint 1:** Consider using the `slope` and `intercept` functions you have defined previously.

**Hint 2:** Re-reading [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html) might be helpful here.

**Note:** The public tests are quite comprehensive for this question, so passing them means that your function most likely works correctly.

_Points:_ 3

```python
def predict(tbl, x_col, y_col):
    x = ...
    y = ...
    ...
```

```python
grader.check("task_09")
```

---

## Electronic Arts (EA) FC Predictions

---

The following data was [scraped](https://en.wikipedia.org/wiki/Data_scraping) from [sofifa.com](https://sofifa.com/), a website dedicated to collecting information from EA FC video games. The dataset consists of all players in [FC 26](https://en.wikipedia.org/wiki/EA_Sports_FC_26) and their corresponding attributes. We have truncated the dataset to a limited number of rows (top 100 based on overall rating) to ease with our visualizations and analysis. 

**Since we're learning about linear regression, we will look specifically for a linear association between various player attributes. To help with understanding where the line of best fit generated in linear regression comes from, please do not use the `.fit_line` argument in `.scatter` at any point unless the code is provided for you.**

Run the following cell to load the data and reduce the original data to a more manageable amount of information. (There are 110 columns in the original dataset!) 

```python
fc = Table.read_table('ea_fc_26.csv')
fc = fc.select("short_name", "overall", "value_eur", 
                   "wage_eur", "age", "pace", "shooting", 
                   "passing", "attacking_finishing")
fc.show(5)
```

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Before jumping into any statistical techniques, it's important to see what the data looks like, because data visualizations allow us to uncover patterns in our data that would have otherwise been much more difficult to see.

Create a scatter plot with age on the x-axis (`"age"`), and the player's value in Euros (`"value_eur"`) on the y-axis.

_Points:_ 2

```python
...
```

<!-- END QUESTION -->

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Create a scatter plot with player age (`"age"`) along the x-axis and both real player value (`"value_eur"`) and predicted player value along the y-axis. 

* The predictions should be created using a fitted regression line.
* The color of the dots for the real player values should be different from the color for the predicted player values.

**Hint 1:** Consider using the `predict` function you have defined previously.

**Hint 2:** [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html) and [Section 7.3](https://inferentialthinking.com/chapters/07/3/Overlaid_Graphs.html) have examples of creating such scatter plots.

_Points:_ 2

```python
predictions = ...
fc_with_predictions = ...
...
```

<!-- END QUESTION -->

---

### Task 12 📍

Define a function `rmse` that takes in two arguments: a slope and an intercept for a potential regression line. The function should return the root mean squared error between the values predicted by a regression line with the given slope and intercept and the actual outcomes.

**Note:** Assume we are still predicting `"value_eur"` from `"age"` in original units from the `fc` table.

_Points:_ 2

```python
def rmse(slope, intercept):
    predictions = ...
    errors = ...
    ...
```

```python
grader.check("task_12")
```

---

### Task 13 📍🔎

<!-- BEGIN QUESTION -->

1. Use the `rmse` function you defined along with `minimize` to find the least-squares regression parameters predicting player value from player age. Review the [example of using the `minimize` function from the textbook](https://inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html#numerical-optimization).
2. Then set `lsq_slope` and `lsq_intercept` to be the least-squares regression line slope and intercept, respectively.
3. Finally, create a scatter plot like you did in Task 11 with player age (`"age"`) along the x-axis and both real player value (`"value_eur"`) and predicted player value along the y-axis. Be sure to use your least-squares regression line to compute the predicted values. The color of the dots for the real player values should be different from the color for the predicted player values.

**Note:** Your solution should **not** make any calls to the slope or intercept functions defined earlier. You are demonstrating a different method for creating the line in this task.

**Hint:** Your call to `minimize` will return an array of argument values that minimize the return value of the function passed to minimize.

_Points:_ 2

```python
minimized_parameters = ...
lsq_slope = ...
lsq_intercept = ...

# This just prints your slope and intercept
print("Slope: {:g} | Intercept: {:g}".format(lsq_slope, lsq_intercept))

lsq_predictions = ...
fc_with_lsq_predictions = ...
...
```

<!-- END QUESTION -->

---

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

What do you notice about the visuals produced in Task 11 and Task 13, and what can you conclude about the two different methods you used to create those two lines of gold dots, the predicted values? In your response, make sure to reference [Section 15.3](https://ccsf-math-108.github.io/textbook/chapters/15/3/Method_of_Least_Squares.html) of the textbook to justify what you write.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> standard_units(make_array(1, 2, 3, 4, 5))
array([-1.41421356, -0.70710678,  0.        ,  0.70710678,  1.41421356])
```
failure_message: ❌ standard_units doesn't work for the test data we provided.
success_message: ✅ standard_units works for the test data we provided.
```
>>> np.isclose(np.mean(standard_units([1, 2, 3, 4, 5])), 0)
np.True_
```
failure_message: ❌ The mean of the standardized test data should be 0.
success_message: ✅ The mean of the standardized test data is 0.
```
>>> np.isclose(np.std(standard_units([1, 2, 3, 4, 5])), 0.9999999999999999)
np.True_
```
failure_message: ❌ The standard deviation of the standardized test data should be 1.
success_message: ✅ The standard deviation of the standardized test data is 1.


Tests for task_02 (points: 2)
```
>>> set(standard_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ standard_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ standard_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_03 (points: 2)
```
>>> np.isclose(correlation([1, 2, 3], [4, 5, 6]), 0.9999999999999999)
np.True_
```
failure_message: ❌ correlation should show that linear data has a coefficient of 1.
success_message: ✅ correlation shows that linear data has a coefficient of 1.


Tests for task_04 (points: 2)
```
>>> set(r_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ r_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ r_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_05 (points: 2)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> np.isclose(slope(test_data, 2 * test_data), 2)
np.True_
```
failure_message: ❌ slope did not produce the correct result for our test data set.
success_message: ✅ slope produced the correct result for our test data set.


Tests for task_06 (points: 2)
```
>>> set(slope_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ slope_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ slope_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_07 (points: 2)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> np.isclose(intercept(test_data, 2 * test_data), 0)
np.True_
```
failure_message: ❌ intercept did not produce the correct result for our test data set.
success_message: ✅ intercept produced the correct result for our test data set.


Tests for task_08 (points: 2)
```
>>> set(intercept_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ intercept_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ intercept_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_09 (points: 3)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', 2 * test_data)
>>> test_prediction = predict(test_tbl, 'x_col', 'y_col')
>>> isinstance(test_prediction, np.ndarray)
True
```
failure_message: ❌ predict should return an array.
success_message: ✅ predict returns an array.
```
>>> test_data = make_array(1, 2, 3, 4)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', 2 * test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([2., 4., 6., 8.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.
```
>>> test_data = make_array(-2, -1, 1, 2)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([-2., -1.,  1.,  2.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.
```
>>> test_data = make_array(-2, -1, 1, 2)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', -1 * test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([ 2.,  1., -1., -2.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.


Tests for task_12 (points: 2)
```
>>> max_value_eur = np.max(fc.column('value_eur'))
>>> min_value_eur = np.min(fc.column('value_eur'))
>>> max_age = np.max(fc.column('age'))
>>> min_age = np.min(fc.column('age'))
>>> test_slope = (max_value_eur - min_value_eur) / (max_age - min_age)
>>> test_intercept = 1.2 * 100000000.0
>>> test_rmse = rmse(test_slope, test_intercept)
>>> isinstance(test_rmse, (float, np.float32, np.float64))
True
```
failure_message: ❌ rmse should return a float.
success_message: ✅ rmse returns a float.
```
>>> max_value_eur = np.max(fc.column('value_eur'))
>>> min_value_eur = np.min(fc.column('value_eur'))
>>> max_age = np.max(fc.column('age'))
>>> min_age = np.min(fc.column('age'))
>>> test_slope = (max_value_eur - min_value_eur) / (max_age - min_age)
>>> test_intercept = 1.2 * 100000000.0
>>> np.isclose(rmse(test_slope, test_intercept), 267464004.02653342)
np.True_
```
failure_message: ❌ rmse does not seem to work on our test data.
success_message: ✅ rmse seems to work on our test data.
