<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 04: Visualizations and Functions

## Textbook References

* [Chapter 7](https://ccsf-math-108.github.io/textbook/chapters/07/Visualization.html)
* [Sections 8.0 - 8.3](https://ccsf-math-108.github.io/textbook/chapters/08/Functions_and_Tables.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Unemployment


The [Federal Reserve Bank of St. Louis](https://fred.stlouisfed.org/categories/33509) publishes data about jobs in the US. We've imported data from this site on unemployment rates in the United States. There are many ways of defining unemployment rates and our dataset includes two.
1. Non-Employment Index (NEI): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job. 
2. NEI+PTER (Non-Employment Index and 'Part-Time for Economic Reasons'): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job or are only working at a part-time job.
The data that you will use in this assignment contains [quarterly average NEI percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIM156SFRBRIC) and [quarterly average NEI+PTER percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIPTERM156SFRBRIC)

Run the following code cell (which you are not responsible for understanding) to load a table called `unemployment` that contains recent NEI and PTER values.

```python
from datetime import datetime, date

def update_date_format(date_string):
    try:
        return datetime.strptime(date_string, '%Y-%m-%d').date()
    except ValueError:
        return None

unemployment = Table.read_table('unemployment.csv')

# Check if the 'DATE' column contains strings, then convert to datetime.date
if not isinstance(unemployment.column('DATE').item(0), date):
    unemployment = unemployment.with_column(
        'DATE', unemployment.apply(update_date_format, 'DATE')
    )

unemployment
```

---

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

Using the `unemployment` table, create an overlaid line plot showing two lines for NEI and PTER unemployment percentages over the dates. The dates should be on the horizontal axes and the percentages should be on the vertical axes.

**Note:** Our code below will format the graph for you so that the x-axis shows the year portion of the date for every 3 years starting from 1994.

_Points:_ 2

```python
...

# Leave the following code to improve the readability of the horizontal axis tick marks
start_date = min(unemployment.column("DATE"))
end_date = max(unemployment.column("DATE"))
years = [datetime(year, 1, 1) for year in range(start_date.year, end_date.year + 3, 3)]
plt.gca().set_xticks(years)
plt.gca().set_xticklabels([year.strftime('%Y') for year in years])
plt.xticks(rotation=45)
plt.title("Unemployment Percentages (1994 - 2024)")
plt.show()
```

<!-- END QUESTION -->

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

As you saw early on in the course, visuals produced from data can reveal a story about the data. For this task, review the line graphs above for patterns, identify 3 major events in US history evident in the graph and hypothesize how the events and the patterns in the unemployment data are related.

In your response:
1. Mention 3 events that occurred between 1994 and 2024.
2. Express how those 3 events show up in the patterns you are observing.
3. Hypothesize how those events are associated with unemployment data.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Birth Rates


The CSV file `census.csv` contains census-based population estimates for each state on both July 1, 2022 and July 1, 2023. The data was taken from [the US Census 2020-2023 national totals data set](https://www2.census.gov/programs-surveys/popest/datasets/2020-2023/state/totals/NST-EST2023-ALLDATA.csv). 

Run the following code cell to load that data into a table called `pop`.

```python
pop = Table.read_table('census.csv')
pop
```

Here is a brief explanation of the column labels:

* `REGION`: Census Region code
* `NAME`: State name
* `'2022'`: 7/1/2022 resident total population estimate
* `'2023'`" 7/1/2023 resident total population estimate
* `'BIRTHS'`: Births in period 7/1/2022 to 6/30/2023			
* `'DEATHS'`: Deaths in period 7/1/2022 to 6/30/2023
* `'MIGRATION'`: Net migration in period 7/1/2022 to 6/30/2023
* `'OTHER'`: Residual for period 7/1/2022 to 6/30/2023

The last four columns describe the components of the estimated change in population during this time interval. 

**Note:** For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.

---

### Task 03 📍

Assign `us_birth_rate` to the total US annual birth rate during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period.

**Hint:** Which year corresponds to the start of the time period?


_Points:_ 2

```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("task_03")
```

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval among the states. 

It may be helpful to create intermediate variables. In our template, we've introduced the names `birth_rates_2022` and `death_rates_2022` as suggestions. We will not test those names, so you do not need to use them. We will only score the visualization you produce.


_Points:_ 2

```python
birth_rates_2022 = ...
death_rates_2022 = ...

# Leave the following code to add a title to your graph
plt.title('2022 Death Rate vs. Birth Rate')
plt.show()
```

<!-- END QUESTION -->

---

## Uber Movement

According to the former [Uber](https://www.uber.com) Movement project page:
> Planning great cities requires great data. Uber gathers trip data in more than 10,000 cities across the world. So why not share it? Enter Uber Movement, which gives urban planners access to Uber’s aggregated data to help make informed decisions about our cities.

The Uber Movement project and data access have ended, but we still have data from Boston and Manila. Below we load tables containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the Uber Movement project. The `'sourceid'` and `'dstid'` columns contain codes corresponding to start and end locations of each ride. The `'hod'` column contains codes corresponding to the hour of the day the ride took place. The ride time column contains the length of the ride in minutes.

Run the following code cell to create the table `uber` which contains the available Uber ride data for Boston and Manila.

```python
boston = Table.read_table("boston.csv").with_column('city', ['Boston']*200_000)
manila = Table.read_table("manila.csv").with_column('city', ['Manila']*200_000)
uber = boston.append(manila)
uber
```

---

### Task 05 📍🔎

<!-- BEGIN QUESTION -->

Using the `uber` table, produce an overlaid histogram that visualizes the distributions of all ride times in Boston and Manila. Use the `group='city'` argument with the [`hist` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.hist.html#datascience.tables.Table.hist) to accomplish this. Additionally, use the given bins in `equal_bins` by utilizing the `bins` argument for the `hist` table method.

_Points:_ 2

```python
equal_bins = np.arange(0, 120, 5)
...

# Leave the following code to add a title to your histogram
plt.title('Distribution of Boston and Manila Ride Times')
plt.show()
```

<!-- END QUESTION -->

---

### Task 06📍

Based on the context surrounding this data and the histogram you created in Task 05, which of the following hypotheses is best supported by the data? Assign the corresponding integer based on the choices below to the variable `boston_manila_hypothesis`.

1. Manila has shorter ride times because it is a smaller city
2. Boston has shorter ride times because it has better traffic infrastructure
3. Manila has longer ride times despite being smaller due to higher traffic congestion
4. Both cities have similar ride time distributions

_Points:_ 2

```python
boston_manila_hypothesis = ...
boston_manila_hypothesis
```

```python
grader.check("task_06")
```

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

From the histograms, it looks like there are more 20 to 40-minute Uber rides in Manila compared to Boston. Histograms reflect density, so be careful of interpreting the heights of the bars as counts. It is okay to compare the histograms directly because they both represent 200,000 data points. For this task, produce a bar chart showing that Manila does have more 20 to 40-minute Uber rides than Boston. 

**Note:** For this task, it doesn't matter if you include 40-minute rides in your count or not.

**Hint:** Consider using the [`group` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.group.html).

_Points:_ 2

```python
...

# Leave the following code to add a title to your histogram
plt.title('Number of Uber Rides between 20 and 40 Minutes')
plt.show()
```

<!-- END QUESTION -->

---

## NYC Motor Vehicle Collisions - Crashes

The data in `'nyc_crashes_2024_sample.csv'` contains information from a random sample of 1,000 police-reported motor vehicle collisions in NYC during 2024. The police report ([MV104-AN](https://www.nhtsa.gov/sites/nhtsa.dot.gov/files/documents/ny_overlay_mv-104an_rev05_2004.pdf)) is required to be filled out for collisions where someone is injured or killed, or where there is at least $1000 worth of damage.

You can read more about the data at the [NYC OpenData page: Motor Vehicle Collisions - Crashes](https://data.cityofnewyork.us/Public-Safety/Motor-Vehicle-Collisions-Crashes/h9gi-nx95/about_data).

Run the following code cell to create the table `nyc` from that `CSV` file.

```python
nyc = Table.read_table('nyc_crashes_2024_sample.csv')
nyc
```

A major part of a data analyst's workload involves prepping the data for analysis. Some refer to this process as cleaning the data. In the next few tasks, you'll use what you've learned about making and applying functions to start to clean up these data.

---

It is very common for data sets to be incomplete. Notice the many `nan` values throughout the table. `nan` refers to "not a number" and is a generic placeholder for missing data. You might see this expressed in a variety of other ways such as `NaN`. For example, in the `'borough'` and `'latitude'` columns, you should see `nan` expressed in the data. A tricky thing about cleaning up the `nan` values in `nyc` is that the data type of the `nan` value depends on the data type of the data in the related column. In general, you should know the data type of a column before working with its values.

### Task 08 📍

What is the data type for the `nan` values in the `'borough'` column? Assign `borough_nan_type` to the integer 1, 2, 3, or 4 corresponding to your response option from below:

1. String
2. Integer
3. Float
4. None of the above

**Hint**: You can use the `type` function to check the data type of an item from an array.

_Points:_ 2

```python
borough_nan_type = ...
borough_nan_type
```

```python
grader.check("task_08")
```

---

### Task 09 📍

What is the data type for the `nan` values in the `'latitude'` column? Assign `latitude_nan_type` to the integer 1, 2, 3, or 4 corresponding to your response option from below:

1. String (`str`)
2. Integer (`int`)
3. Float (`float`)
4. None of the above

**Hint**: You can use the `type` function to check the data type of an item from an array.

_Points:_ 2

```python
latitude_nan_type = ...
latitude_nan_type
```

```python
grader.check("task_09")
```

---

Aside from the `nan` values, the following code shows that the `'latitude'` and `'longitude'` columns contain `0` values, which doesn't make sense because [New York City is bounded](https://en.wikipedia.org/wiki/Module:Location_map/data/USA_New_York_City) by the latitude values of 40.49 to 40.92 and the longitude values of -74.27 to -73.68. The center of the city is at (40.705, -73.975).

```python
nyc.sort('latitude').show(5)
```

One way to handle missing values is to remove the rows for which certain columns have missing values. 

We created the following function for you that determines if a latitude and longitude combination could possibly be in New York City. The function returns a `bool` value depending on whether or not the coordinates are within the NYC boundaries. You'll learn about the code in this function soon!

```python
def in_nyc(lat, long):
    """
    This function checks if the latitude (float) and longitude (float) provided fall within 
    the approximate boundaries of New York City. The function returns True if the latitude is between 
    40.49 and 40.92 and the longitude is between -74.27 and -73.68. Otherwise, it returns False.

    Examples:
    >>> in_nyc(40.705, -73.975)
    True

    >>> in_nyc(27.000, -73.975)
    False
    """
    
    if 40.49 <= lat <= 40.92 and -74.27 <= long <= -73.68:
        return True
    else:
        return False
```

Notice how the function handles `nan` values:

```python
in_nyc(np.nan, -73.975)
```

Let's use this tool to clean up the data regarding the latitude and longitude values.

### Task 10 📍

Apply the `in_nyc` function to `nyc` to create the table `nyc_clean_1` that contains only the rows in `nyc` where the latitude and longitude values are possible. In the end, `nyc_clean_1` should have the same columns as `nyc` but with fewer rows.

**Notes**: 
1. Remember from a recent lecture that you can apply a function with multiple arguments.
2. To remove the unwanted rows, try adding a column of the `bool` values to the table and using `where` to filter based on the flawed latitude and longitude values. Don't forget to remove that added column in your final table.

_Points:_ 4

```python
in_nyc_bools = ...
nyc_clean_1 = ...
nyc_clean_1
```

```python
grader.check("task_10")
```

```python
nyc_clean_1.pivot('borough','contributing_factor_vehicle_1')
```

---

### PARKWAY vs. PKWY

A classic data cleaning problem involves handling multiple representations of the same thing. For example, run the following code to see that Grand Central Parkway is expressed 2 times as `'GRAND CENTRAL PARKWAY'` and 9 times as 'GRAND CENTRAL PKWY' in the `'on_street_name'` column. For reference, the `'on_street_name'` column gives the street on which the collision occurred according to the [NYC OpenData reference](https://data.cityofnewyork.us/Public-Safety/Motor-Vehicle-Collisions-Crashes/h9gi-nx95/about_data).
 

```python
(nyc_clean_1
 .where('on_street_name', are.containing('GRAND CENTRAL'))
 .group('on_street_name'))
```

`'PKWY'` is an abbreviation for `'PARKWAY'`. Now, we'll work on cleaning up this issue in the data, because in the end, we should be acknowledging that there are a total of 12 incidents on this street.

### Task 11 📍

Write a function called `update_PKWY` that takes in a string as input, replaces any instance (sub-string) of `'PKWY'` with `'PARKWAY'`, and returns the updated string.

**Notes:**
* Remember that `replace` function that you've used before in this class when working with strings? That could be helpful here!
* As a test case, running `update_PKWY('GRAND CENTRAL PKWY')` should return `'GRAND CENTRAL PARKWAY'`.

_Points:_ 2

```python
...
```

```python
grader.check("task_11")
```

---

### Task 12 📍

Using `update_PKWY`, replace every instance of `PKWY` in the `'on_street_name'`, `'off_street_name'`, and `'cross_street_name`' columns of `nyc_clean_1` with `'PARKWAY'`. Save the updated version of `nyc_clean_1` as `nyc_cleanish`.

**Notes**: 
* We recommend that you break up this task into multiple steps (lines of code).
* Try using the `apply` table method to update the street names. If you have any coding experience, you might be tempted to use a loop of some kind. We don't recommend this because loops are generally inefficient for working with large data sets. If you don't know what a loop is, don't worry! You'll learn about those in this class.
* Also, remember that the `with_columns` table method will replace the information in an existing column with new information if you use the same column label.

_Points:_ 4

```python
nyc_cleanish = ...
```

```python
grader.check("task_12")
```

---

There is still a lot to clean up in this data set and a lot of ways we could fill in missing information. For example, it turns out that the entire route of the [Grand Central Parkway]((https://en.wikipedia.org/wiki/Grand_Central_Parkway)) is contained within the borough of Queens. This means that we can fill in some of the missing `'borough'` values using this observation. Or, as the next tasks will ask you to do, you can ignore those rows.

---

### Task 13 📍

For our last tasks, we want to understand the distribution of crashes relative to various factors. Let's start by looking at a distribution among the boroughs in the `nyc_cleanish` table created above. Assign the table `crash_by_borough` to a table that shows us the number of crashes in each borough.

In addition, we want to make sure that the table does NOT include any crashes where the borough is `nan`, so first run the cell below to remove such rows. Then, in the cell below that, write your solution.

HINT: The final table should have 5 rows and 2 columns.

_Points:_ 2

```python
nyc_cleanish = nyc_cleanish.where(nyc_cleanish.column('borough') != 'nan')
```

```python
crash_by_borough = ...
crash_by_borough
```

```python
grader.check("task_13")
```

---

### Task 14 📍

In this task, we want to understand the distribution of crashes relative to the borough AND the contributing factor. Start with the `nyc_cleanish` table and use the `borough` and `contributing_factor_vehicle_1` columns. Assign the table `crash_by_borough_and_factor` to a table that shows us the number of crashes for each borough and contributing factor category. Make sure that you create a table that has a column for each of the boroughs, and a row for each contributing factor.

_Points:_ 2

```python
crash_by_borough_and_factor = ...
crash_by_borough_and_factor
```

```python
grader.check("task_14")
```

---

## A Map of Crashes by Borough

To wrap up, here is a map of the crashes by borough using the latitude and longitude values you helped clean up! We just dropped the rows in the table with missing borough information. There is nothing you need to do with this except preview some upcoming content in the code.

```python
nyc_cleanish = nyc_cleanish.where(nyc_cleanish.column('borough') != 'nan')

colors = (nyc_cleanish
          .group('borough')
          .drop('count')
          .with_column('color', 
                       make_array('blue', 'red', 'green', 'orange', 'purple')))

nyc_map = (nyc_cleanish
           .join('borough', colors)
           .select('latitude', 'longitude', 'borough', 'color')
           .relabeled('latitude', 'lat')
           .relabeled('longitude', 'long')
           .relabeled('borough', 'labels'))

Marker.map_table(nyc_map)
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

1. **Review the Rubric:** View the rubric on the Canvas assignment page to understand the scoring criteria.  
2. **Review Manually Graded Tasks:** Verify that you have responded to all manually graded tasks marked with 📍🔎.  
3. **Review Auto-Graded Tasks:** Inspect the output from the `grade.check_all()` command above. Ensure you have passed all available tests for the auto-graded tasks in this notebook marked with 📍.  
   - This command executes all auto-grader tests sequentially.  
   - Note that Homework and Project assignments include hidden auto-grader tests that are not visible to you. Passing all visible tests does **not** guarantee that your solution is fully correct.  
4. **Save Your Work:** In the notebook toolbar, go to `File → Save Notebook` to save your work and create a checkpoint.  
5. **Download the Notebook:** In the notebook toolbar, go to `File → Download IPYNB` to download the notebook (`.ipynb`) file.  
6. **Upload to Canvas:** Upload the downloaded `.ipynb` file to the appropriate Canvas assignment page.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_03 (points: 2)
```
>>> 0 < us_birth_rate < 1
np.True_
```
failure_message: ❌ Make sure you provide a value betwen 0 and 1.
success_message: ✅ You have provided a possible value.


Tests for task_06 (points: 2)
```
>>> boston_manila_hypothesis in {1, 2, 3, 4}
True
```
failure_message: ❌ boston_manila_hypothesis should be assigned to an integer 1, 2, 3, or 4.
success_message: ✅ boston_manila_hypothesis is assigned to an integer 1, 2, 3, or 4.


Tests for task_08 (points: 2)
```
>>> borough_nan_type in {1, 2, 3, 4}
True
```
failure_message: ❌ borough_nan_type should be assigned to an integer 1, 2, 3, or 4.
success_message: ✅ borough_nan_type is assigned to an integer 1, 2, 3, or 4.


Tests for task_09 (points: 2)
```
>>> latitude_nan_type in {1, 2, 3, 4}
True
```
failure_message: ❌ latitude_nan_type should be assigned to an integer 1, 2, 3, or 4.
success_message: ✅ latitude_nan_type is assigned to an integer 1, 2, 3, or 4.


Tests for task_10 (points: 4)
```
>>> isinstance(nyc_clean_1, Table) and set(nyc_clean_1.labels) == set(nyc.labels)
True
```
failure_message: ❌ nyc_clean_1 is not a table with the same columns as nyc.
success_message: ✅ nyc_clean_1 is a table with the same columns as nyc.


Tests for task_11 (points: 2)
```
>>> 'update_PKWY' in globals() and callable(update_PKWY)
True
```
failure_message: ❌ update_PKWY is not callable.
success_message: ✅ update_PKWY is callable.
```
>>> update_PKWY('GRAND CENTRAL PKWY') == 'GRAND CENTRAL PARKWAY'
True
```
failure_message: ❌ update_PKWY doesn't work correctly with GRAND CENTRAL PKWY.
success_message: ✅ update_PKWY works correctly with GRAND CENTRAL PKWY.


Tests for task_12 (points: 4)
```
>>> isinstance(nyc_cleanish, Table) and set(nyc_cleanish.labels) == set(nyc.labels)
True
```
failure_message: ❌ nyc_cleanish is not a Table with the same columns as nyc.
success_message: ✅ nyc_cleanish is a Table with the same columns as nyc.
```
>>> 'GRAND CENTRAL PKWY' not in nyc_cleanish.column('on_street_name')
True
```
failure_message: ❌ 'GRAND CENTRAL PKWY' is in the on_street_name column of nyc_cleanish.
success_message: ✅ 'GRAND CENTRAL PKWY' is not in the on_street_name column of nyc_cleanish.


Tests for task_13 (points: 2)
```
>>> isinstance(crash_by_borough, Table)
True
```
failure_message: ❌ crash_by_borough is not a Table.
success_message: ✅ crash_by_borough is a Table.


Tests for task_14 (points: 2)
```
>>> isinstance(crash_by_borough_and_factor, Table)
True
```
failure_message: ❌ crash_by_borough_and_factor is not a Table.
success_message: ✅ crash_by_borough_and_factor is a Table.
