<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 14: Conditional Probability

## References

* [Sections 18.0 - 18.2 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/18/Updating_Predictions.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab14.ipynb')
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
# This line loads the visualization code for this lab.
import visualizations
```

<!-- END QUESTION -->

---

## What is conditional probability good for?

Suppose we have a known population, like all dogs in California.  So far, we've seen three ways of *predicting* something about an individual in that population, given incomplete knowledge about the identity of the individual:

* If we know nothing about the individual dog, we could predict that its speed is the *average* or *median* of all the speeds in the population.
* If we know the dog's height but not its speed, we could use *linear regression* to predict its speed from its height.  The resulting prediction is still imperfect, but it might be more accurate than the population average.

We can also compute conditional probabilities to make predictions about individuals or events. This technique is different from the previous methods we've examined because 

1. our prediction for each outcome is described by a probability, and 
2. each probability can be exactly calculated from assumptions, as opposed to estimated from data.

---

## Marbles

In this lab, you will work with a more complex example about disease, but first, let's start with a simple example.

Imagine you and Samantha are playing a game in which you are given a marble and tasked to determine the marble's texture and size. You don't know anything about the marble you're given, but you know that Samantha drew it **uniformly at random** from a bag that contained the following marbles:
* 4 large shiny marbles,
* 1 large dull marble,
* 6 small shiny marbles,
* 2 small dull marbles.

---

### Task 01 📍

Knowing only what we've told you so far, what's the probability that the marble you're given is a large shiny marble? Assign `probability_large_shiny` to this numerical probability value.

```python
probability_large_shiny = ...
probability_large_shiny
```

```python
grader.check("task_01")
```

---

Here's a table with those marbles:

```python
marbles = Table.read_table("marbles.csv")
marbles.show()
```

Here are the counts of each type of marble in a pivot table.

```python
marbles.pivot('surface', 'size')
```

Here are all the different combinations of surface and size, with the count for each surface-size combination. Each type of marble appears in its own row.

```python
marbles.group(['surface', 'size'])
```

Knowing nothing else about the marble, it's equally likely to be any of the marbles accounted above; this is because we've assumed that the marble is selected **uniformly at random** from the bag.

---

### Task 02 📍

What's the probability that the marble that you've been given is a shiny marble? Assign `probability_shiny` to this numerical value.

```python
probability_shiny = ...
probability_shiny
```

```python
grader.check("task_02")
```

---

### Conditional probability

Suppose you overhear Samantha say that you were given a large marble. Does this somehow change the chance that your marble is shiny?  Let's find out.

In Task 02, we assumed that the marble you were given was equally likely to be any of the marbles, simply because we didn't know any better.  That's why we looked at all the marbles to compute the probability that your marble was shiny.

But assuming that you've been given a large marble, we can eliminate some of these possibilities. In particular, you can't have been given a small shiny marble or a small dull marble.

You're still equally likely to have been given any of the remaining marbles, because you don't know any other information.

---

### Task 03 📍

What's the probability Samantha gave you a shiny marble, knowing that she gave you a large marble? Assign `probability_shiny_given_large` to that numerical value.

```python
probability_shiny_given_large = ...
probability_shiny_given_large
```

```python
grader.check("task_03")
```

---

You should have found that this is different from the probability that Samantha gave you a shiny marble, given no size information, which you computed earlier.  The distribution of surfaces among the large marbles is a little different from the distribution of surfaces among all the marbles. We have calculated a conditional probability with the additional assumption of a given size

---

### Task 04 📍

Suppose instead Samantha had said she gave you a shiny marble (hooray!). What's the probability that the marble given to you is large?  Assign `probability_large_given_shiny` to that numerical value.

```python
probability_large_given_shiny = ...
probability_large_given_shiny
```

```python
grader.check("task_04")
```

---

## Cancer screening

Hopefully the previous example helped you build intuition for why conditional probabilities can be helpful. Now, let's look at a real life application.

---

### Background

Medical testing is very important, especially for cancer. A basic cancer screening involves looking for cancer before a person has any symptoms. 

Cancer screening is not 100% reliable, and it can have errors. There are cases in which someone without cancer or with a benign tumor sees a positive test result. There are also cases where someone with cancer can receive a negative result. ("Positive" implies an indication of cancer in this context.) The first case, called a false positive, could cause anxiety for a patient and lead to further testing that may be risky and involve unnecessary radiation. In the second case, called a false negative, a person would not receive the necessary information and possible treatment for their condition.

Conditional probability can provide insight into the accuracy of initial cancer tests. For example, you can compute the chance that a person has cancer given the results of a diagnostic test by combining information from different probability distributions. You'll see that the chance Person X does have cancer even if they receive a positive result is not necessarily 100%.

It is important to note that conditional probabilities offer just one line of insight to understanding a test result and that this statistical approach does not take into account important factors like heredity or environment that come to affect a person's health and diagnoses. The [Centers for Disease Control and Prevention](https://www.cdc.gov/cancer/dcpc/prevention/screening.htm) (CDC) recommends getting preventative cancer screenings, and one should always consult a doctor if they are concerned about their health. 

---

### Hypothetical Cancer Statistics

**This is a hypothetical scenario.**

Suppose that in a representative group of 10,000 people who are tested for cancer ("representative" meaning that the frequencies of different events are the same as the frequencies in the whole population):
1. 100 people have cancer.
2. Among the 100 people that have cancer, 90 have positive results on a cancer test and 10 have negative results.  (So 10 people receive false negative results.)
3. The other 9,900 people don't have cancer.
4. Among these 9,900 people, 198 have positive results on a cancer test and the other 9,702 have negative results.  (So 198 see "false positive" results.)

Below we've generated a table with data from these 10,000 hypothetical people.

*Note: These statistics are made up and don't necessarily reflect the actual state of the world.*

```python
people = Table().with_columns(
    "status", ["cancer", "cancer", "no cancer", "no cancer"],
    "test status", ["positive", "negative", "positive", "negative"],
    "count", [90, 10, 198, 9702])
people
```

One way to visualize this data set is with a contingency table, which you've seen before.

---

### Task 05 📍

Using the `people` table defined above, create a contingency table that looks like this:

|status|negative|positive|
|-:|-:|-:|
|cancer|||
|no cancer||||

with the **count** of each group filled in, according to what we've told you above.  The counts in the 4 boxes should sum to 10,000.

```python
cancer = ...
cancer
```

```python
grader.check("task_05")
```

---

Now let's think about how you can use this kind of information when Person X is tested for cancer.

Before you know any information about Person X, you could imagine Person X as a **uniform random sample** of one of the 10,000 people in this imaginary population of people who have been tested.

What's the chance that Person X has cancer, knowing nothing else about them? If we add together the counts from the negative and positive of the cancer row, we can find that the chance is $\frac{100}{10000}$, or 1%.

---

### Task 06 📍

What's the chance that Person X has a positive test result, knowing nothing else about them?
Assign `probability_positive_test` to that numerical value.

```python
probability_positive_test = ...
probability_positive_test
```

```python
grader.check("task_06")
```

---

### Interpreting Test Results

Suppose Person X has a positive test result.  This means that you can now eliminate the options of people with a negative test result and narrow them down to being part of just one of the two following groups:
1. The people with cancer who have a positive test result.
2. The people without cancer who have a positive test result.

The *conditional probability* that Person X **has cancer given their positive test result** is the chance that they're in the first group (cancer), assuming they have a positive test result.

---

### Task 07 📍

What is the conditional probability that Person X has cancer given their positive test result?

Assign `prob_cancer_given_positive` to your answer.

```python
prob_cancer_given_positive = ...
prob_cancer_given_positive
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Run the next code cell to see the `cancer` table.

```python
cancer.show()
```

Write code to calculate the conditional probability that Person X has cancer **given their positive test result** using the `cancer` table. Assign the value to `prob_cancer_given_positive_code`. NOTE: Do NOT give a numerical calculation as you did in the previous task - you must use the table to find the answer.

**Hint**: Review [Section 18.1](https://ccsf-math-108.github.io/textbook/chapters/18/1/More_Likely_than_Not_Binary_Classifier.html#).

```python
prob_cancer_given_positive_code = ...
prob_cancer_given_positive_code
```

```python
grader.check("task_08")
```

---

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

Is your answer to Task 08 larger than the overall proportion of people in the population who have cancer (given as 1%)?  If so, does this make sense? Why or why not? Check your response with a classmate, a tutor, or the instructor before moving on.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## A Subjective Prior

Being right isn't always satisfying. Classifying a Positive patient as not having cancer still seems somehow wrong, for such an accurate test. In this section, we'll take a deeper look at the basis of our probability calculation: the assumption of randomness.

Our assumption was that a randomly chosen person was tested from our population of 10,000 individuals. This doesn't happen in reality. People go in to get tested because they think they might have the disease, or because their doctor thinks they might have the disease. People getting tested are not randomly chosen members of the population.

That is why our intuition about people getting tested was not fitting well with the answer that we got. In Task 09, we calculated that the probability of cancer given a positive test result was **0.3125**, or **31.25%**. This seems unusually low, especially given what we know about the efficacy of cancer screenings. We were imagining a realistic situation of a patient going in to get tested because there was some reason for them to do so, whereas the calculation was based on a randomly chosen person being tested.

So let's redo our calculation under the more realistic assumption that the patient is getting tested because the doctor thinks there's a chance the patient has the disease and compare the two probabilities.

**Aside:** For more information about this part of the lab, check out [Section 18.2.2](https://ccsf-math-108.github.io/textbook/chapters/18/2/Making_Decisions.html#a-subjective-prior) from the textbook.

---

### A New Look at the (Hypothetical) Population

Suppose now that our population of 10,000 individuals takes into account the fact that choosing to take a medical test is not done at random; there must be some motive for an individual to take one. Thus, our modified population has the following properties:
1. 1,000 people have cancer.
1. Among the 1,000 people that have cancer, 900 have positive results on a cancer test and 100 have negative results. 
3. The other 9,000 people don't have cancer.
4. Among these 9,000 people, 180 have positive results on a cancer test and the other 8,820 have negative results.  (So 180 see "false positive" results.)

Below we've generated a population table called `people_new` and a contingency table with data from these 10,000 hypothetical people called `cancer_new`.

```python
# Just run this cell to load the table
people_new = Table().with_columns(
    "status", ["cancer", "cancer", "no cancer", "no cancer"],
    "test status", ["positive", "negative", "positive", "negative"],
    "count", [900, 100, 180, 8820])
people_new
```

```python
# Just run this cell to load the table
cancer_new = Table().with_columns(
    "status", ["cancer", "no cancer"],
    "negative", [100, 8820],
    "positive", [900, 180])
cancer_new
```

---

### Interpretation and Comparison

Again, let's suppose Person X has a positive test result.  This means that you can now narrow them down to being part of just one of the two following groups:
1. The people with cancer who have a positive test result.
2. The people without cancer who have a positive test result.

---

### Task 10 📍

Using the new contingency table, `cancer_new`, that was provided at the beginning of this section, write code to calculate the exact probability that an individual has cancer given that they've received a positive test result. Assign `prob_cancer_given_positive_new` to that probability.

```python
prob_cancer_given_positive_new = ...
prob_cancer_given_positive_new
```

```python
grader.check("task_10")
```

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

How does your probability in Task 10 compare to your answer from Task 08? Check your response with a classmate, a tutor, or the instructor.

*Hint:* Review [Section 18.2.2](https://ccsf-math-108.github.io/textbook/chapters/18/2/Making_Decisions.html#a-subjective-prior).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: None)
```
>>> import numbers
>>> bool(isinstance(probability_large_shiny, numbers.Number))
True
```
failure_message: "❌ probability_large_shiny is not assigned to a number."
success_message: "✅ probability_large_shiny is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * probability_large_shiny))
'8e98d81f8217304975ccb23337bb5761'
```
failure_message: "❌ probability_large_shiny is not correct. How many large shiny marbles are there? Also, how many marbles are there in total? Divide those two numbers."
success_message: "✅ probability_large_shiny is correct."



Tests for task_02 (points: None)
```
>>> import numbers
>>> bool(isinstance(probability_shiny, numbers.Number))
True
```
failure_message: "❌ probability_shiny is not assigned to a number."
success_message: "✅ probability_shiny is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * probability_shiny))
'288cc0ff022877bd3df94bc9360b9c5d'
```
failure_message: "❌ probability_shiny is not correct. How many large shiny marbles are there? Also, how many marbles are there in total? Divide those two numbers."
success_message: "✅ probability_shiny is correct."



Tests for task_03 (points: None)
```
>>> import numbers
>>> bool(isinstance(probability_shiny_given_large, numbers.Number))
True
```
failure_message: "❌ probability_shiny_given_large is not assigned to a number."
success_message: "✅ probability_shiny_given_large is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * probability_shiny_given_large))
'7a53928fa4dd31e82c6ef826f341daec'
```
failure_message: "❌ probability_shiny_given_large is not correct. How many large shiny marbles are there? Also, how many large marbles are there in total? Divide those two numbers."
success_message: "✅ probability_shiny_given_large is correct."



Tests for task_04 (points: None)
```
>>> import numbers
>>> bool(isinstance(probability_large_given_shiny, numbers.Number))
True
```
failure_message: "❌ probability_large_given_shiny is not assigned to a number."
success_message: "✅ probability_large_given_shiny is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * probability_large_given_shiny))
'18d8042386b79e2c279fd162df0205c8'
```
failure_message: "❌ probability_large_given_shiny is not correct. How many large shiny marbles are there? Also, how many shiny marbles are there in total? Divide those two numbers."
success_message: "✅ probability_large_given_shiny is correct."



Tests for task_05 (points: None)
```
>>> isinstance(cancer, Table)
True
```
failure_message: "❌ cancer is not assigned to a Table. Try using the people table."
success_message: "✅ cancer is assigned to a Table."

```
>>> set(cancer.labels) == {'negative', 'positive', 'status'}
True
```
failure_message: "❌ cancer does not have the correct labels."
success_message: "✅ cancer has the correct labels."

```
>>> cancer.num_rows
2
```
failure_message: "❌ cancer does not have the correct number of rows."
success_message: "✅ cancer has the correct number of rows."

```
>>> cancer
status    | negative | positive
cancer    | 10       | 90
no cancer | 9702     | 198
```
failure_message: "❌ cancer is not correct. Try using the people table with the pivot command."
success_message: "✅ cancer is correct."



Tests for task_06 (points: None)
```
>>> import numbers
>>> bool(isinstance(probability_positive_test, numbers.Number))
True
```
failure_message: "❌ probability_positive_test is not assigned to a number."
success_message: "✅ probability_positive_test is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * probability_positive_test))
'33e75ff09dd601bbe69f351039152189'
```
failure_message: "❌ probability_positive_test is not correct."
success_message: "✅ probability_positive_test is correct."



Tests for task_07 (points: None)
```
>>> isinstance(prob_cancer_given_positive, (float, np.float64))
True
```
failure_message: "❌ prob_cancer_given_positive is a float."
success_message: "✅ prob_cancer_given_positive is a float."

```
>>> bool(prob_cancer_given_positive == 0.3125)
True
```
failure_message: "❌ prob_cancer_given_positive is not correct."
success_message: "✅ prob_cancer_given_positive is correct."



Tests for task_08 (points: None)
```
>>> import numbers
>>> bool(isinstance(prob_cancer_given_positive_code, numbers.Number))
True
```
failure_message: "❌ prob_cancer_given_positive_code is not assigned to a number."
success_message: "✅ prob_cancer_given_positive_code is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * prob_cancer_given_positive_code))
'950a4152c2b4aa3ad78bdd6b366cc179'
```
failure_message: "❌ prob_cancer_given_positive_code is not correct. How many have have the cancer status and a positive test result. Also, how many have a positive test result? Divide those two numbers."
success_message: "✅ prob_cancer_given_positive_code is correct."



Tests for task_10 (points: None)
```
>>> import numbers
>>> bool(isinstance(prob_cancer_given_positive_new, numbers.Number))
True
```
failure_message: "❌ prob_cancer_given_positive_new is not assigned to a number."
success_message: "✅ prob_cancer_given_positive_new is assigned to a number."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(int(1000 * prob_cancer_given_positive_new))
'013a006f03dbc5392effeb8f18fda755'
```
failure_message: "❌ prob_cancer_given_positive_new is not correct. How many have have the cancer status and a positive test result in the new table. Also, how many have a positive test result in the new table.? Divide those two numbers."
success_message: "✅ prob_cancer_given_positive_new is correct."

