<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 13: Classification

## References

* [Classification](https://ccsf-math-108.github.io/textbook/chapters/17/Classification.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw13.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Bay Area School Coordinates with Classification


Carol is trying classify students as either attendees of UC Berkeley or as attendees of Stanford. To classify the students, Carol has access to the coordinates of the location they live during the school year. First, load in the `coordinates` table.

```python
# Just run this cell!
coordinates = Table.read_table('coordinates.csv')
coordinates.show(5)
```

As usual, let's investigate our data visually before performing any kind of numerical analysis.

```python
# Just run this cell!
coordinates.scatter("longitude", "latitude", group="school")
```

The locations of the points on this scatter plot might be familiar - run the following cell to see what they correspond to.

```python
# Just run this cell!
colors = {"Berkeley":"blue", "Stanford":"red"}
t = Table().with_columns("lat", coordinates.column(0), 
                                      "lon", coordinates.column(1), 
                                      "color", coordinates.apply(colors.get, 2)
                        )
Circle.map_table(t, area=5, fill_opacity=1)
```

---

### Task 01 📍

Let's begin implementing the k-Nearest Neighbors algorithm. Define the `distance` function, which takes in two arguments: an array of numerical features, and a different array of numerical features. The function should return the [Euclidean distance](https://en.wikipedia.org/wiki/Euclidean_distance) between the two arrays. Euclidean distance is often referred to as the straight-line distance formula that you may have learned previously.



_Points:_ 2

```python
def distance(arr1, arr2):
    ...

# Don't change/delete the code below in this cell
distance_example = distance(make_array(1, 2, 3), make_array(4, 5, 6))
distance_example
```

```python
grader.check("task_01")
```

---

### Splitting the dataset

We'll do 2 different kinds of things with the `coordinates` dataset:
1. We'll build a classifier using coordinates for which we know the associated label; this will teach it to recognize labels of similar coordinate values. This process is known as *training*.
2. We'll evaluate or *test* the accuracy of the classifier we build on data we haven't seen before.

For reasons discussed in lecture and the textbook, we want to use separate datasets for these two purposes.  So we split up our one dataset into two.

NOTE: As we learned in the lectures and lab, we want to also have a validation set that is used to find the optimal `k` value, but for the purposes of this homework we will skip that part of the process.

---

#### Task 02 📍

Next, let's split our dataset into a training set and a test set. Since `coordinates` has $100$ rows, let's create a training set with the first 75 rows and a test set with the remaining 25 rows. Remember that assignment to each group should be random, so we should shuffle the table first. Assign the training set table to `train` and the test set table to `test`.

**Hint**: Try using the [`split` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.split.html#datascience.tables.Table.split) or try shuffling the rows of the table and using the `take` table method.

_Points:_ 3

```python
...
print("Training set:\t",   train.num_rows, "examples")
print("Test set:\t",       test.num_rows, "examples")
train.show(5), test.show(5);
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Assign `features` to an array of the labels of the features from the `coordinates` table.

*Hint: which of the column labels in the `coordinates` table are the features, and which of the column labels correspond to the class we're trying to predict?*


_Points:_ 2

```python
features = ...
features
```

```python
grader.check("task_03")
```

---

### Task 04 📍

Now define the `classify` function. This function should take in a `row` from a table like `test` and classify it based on the data in `train` using the `k`-Nearest Neighbors based on the correct `features`.

*Hint: use the `row_to_array` function we defined for you to convert rows to arrays of features so that you can use the `distance` function you defined earlier.*

*Hint 2: the skeleton code we provided iterates through each row in the training set*


_Points:_ 2

```python
def row_to_array(row, features):
    arr = make_array()
    for feature in features:
        arr = np.append(arr, row.item(feature))
    return arr

def classify(row, k, train):
    test_row_features_array = row_to_array(row, features)
    distances = make_array()
    for train_row in train.rows:
        train_row_features_array = ...
        row_distance = ...
        distances = np.append(distances, row_distance)
    train_with_distances = ...
    nearest_neighbors = ...
    most_common_label = ...
    return most_common_label

# Don't modify/delete the code below
first_test = classify(test.row(0), 5, train)
first_test
```

```python
grader.check("task_04")
```

---

### Task 05 📍

1. Define the function `three_classify` that takes a `row` from `test` as an argument and classifies the row based on using 3-Nearest Neighbors.
2. Use this function to find the `accuracy` of a 3-NN classifier on the `test` set. `accuracy` should be a proportion (not a percentage) of the schools that were correctly predicted.

*Hint: you should be using a function you just created!*

*Note: Usually before using a classifier on a test set, we'd classify first on a "validation" set, which we then can modify our training set again if need be, before actually testing on the test set. You don't need to do that for this question, but you will learn about this more in an advanced course.*


_Points:_ 3

```python
def three_classify(row):
    ...

test_with_prediction = ...
labels_correct = ...
accuracy = ...
accuracy
```

```python
grader.check("task_05")
```

---

### Task 06 📍

There are 77 rows of Berkeley students and 23 rows of Stanford students in the `coordinates` table. If we used the entire `coordinates` table as the train set, what is the smallest value of k that would ensure that a k-Nearest Neighbor classifier would always predict Berkeley as the class? Assign the value to `k`.


_Points:_ 2

```python
k = ...
k
```

```python
grader.check("task_06")
```

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Why do we use an odd-numbered `k` in k-NN? Explain.



_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## k-NN for Non-Binary Classification

In this class, we have taught you how to use the kNN algorithm to classify data as one of two classes. However, much of the data you will encounter in the real world will not fall nicely into one of two categories. 

How can we classify data with non-binary classes? It turns out we can still use kNN! That is, we find the distance between a point and all its neighbors, find the nearest neighbors, and take a majority vote among the neighbors to determine this point's class. 

The only difference is that now the neighboring points have more than two possible classes. This does introduce difficulty because now we have no way of guaranteeing that we will not encounter ties between classes. In the case that we do encounter a tie, we can just arbitrarily choose one of the classes.

In fact, you don't even have to modify the code you wrote before at all to enable multi-class classification!

Let's add some more data to our train table, this time for another class of students, students at San Jose Community College (SJCC).

```python
coordinates_multi = coordinates.with_rows([
                              [37.304346, -121.915401, "SJCC"],
                              [37.316275, -121.913879, "SJCC"],
                              [37.409435, -121.951379, "SJCC"],
                              [37.349387, -121.960771, "SJCC"],
                              [37.329083, -121.928479, "SJCC"],
                              [37.313017, -121.866730, "SJCC"],
                              [37.346525, -121.894767, "SJCC"],
                              [37.364157, -121.955717, "SJCC"],
                              [37.383362, -121.925776, "SJCC"],
                              [37.329545, -121.880639, "SJCC"]                             
])
```

```python
classify(coordinates_multi.row(0), 5, coordinates_multi)
```

```python
classify(coordinates_multi.row(91), 5, coordinates_multi)
```

```python
classify(coordinates_multi.row(105), 5, coordinates_multi)
```

Our classifier can classify rows as belonging to one of three classes!

Classification is one of the most important fields in statistics, data science, and machine learning. There are thousands of different classification algorithms and modifications of algorithms! There are many that you'll learn if you continue down the path of becoming a data scientist!

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> type(distance_example) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ distance_example should be a float.
success_message: ✅ distance_example is a float.
```
>>> np.round(distance_example, 3) == 5.196
np.True_
```
failure_message: ❌ distance_example doesn't seem correct.
success_message: ✅ distance_example seems correct.


Tests for task_02 (points: 3)
```
>>> train.num_rows == 75
True
```
failure_message: ❌ train does not have the correct number of rows.
success_message: ✅ train has the correct number of rows.
```
>>> test.num_rows == 25
True
```
failure_message: ❌ test does not have the correct number of rows.
success_message: ✅ test has the correct number of rows.
```
>>> train.num_rows + test.num_rows == coordinates.num_rows
True
```
failure_message: ❌ The sum of the number of rows in the training set and test set should give the total number of rows in coordinates.
success_message: ✅ The sum of the number of rows in the training set and test set gives the total number of rows in coordinates.


Tests for task_03 (points: 2)
```
>>> isinstance(features, np.ndarray) and len(features) == 2
True
```
failure_message: ❌ features should be an array with 2 items.
success_message: ✅ features is an array with 2 items.


Tests for task_04 (points: 2)
```
>>> type(first_test) == str
True
```
failure_message: ❌ first_test should be a string.
success_message: ✅ first_test is a string.
```
>>> sorted_coordinates = coordinates.sort('school')
>>> classify(sorted_coordinates.row(85), 3, sorted_coordinates.take(np.arange(50, 100))) == 'Stanford'
True
```
failure_message: ❌ classify did not correctly classify the row we tested.
success_message: ✅ classify correctly classified the row we tested.


Tests for task_05 (points: 3)
```
>>> sorted_coordinates = coordinates.sort('school')
>>> classify(sorted_coordinates.row(29), 3, train) == three_classify(sorted_coordinates.row(29))
True
```
failure_message: ❌ classify and three_classify should provide the same result for the row we tested.
success_message: ✅ classify and three_classify seem to provide the same result for the row we tested.
```
>>> 0 <= accuracy <= 1
np.True_
```
failure_message: ❌ accuracy should be a number between 0 and 1. Make sure you didn't try and calculate a percentage.
success_message: ✅ accuracy is a number between 0 and 1.


Tests for task_06 (points: 2)
```
>>> type(k) == int
True
```
failure_message: ❌ k should be an integer.
success_message: ✅ k is an integer.
