<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 04: Visualizations and Functions

## References

* [Sections 7.0 - 7.3](https://ccsf-math-108.github.io/textbook/chapters/07/Visualization.html)
* [Sections 8.0 - 8.1](https://ccsf-math-108.github.io/textbook/chapters/08/Functions_and_Tables.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Markdown Cheat Sheet](https://www.markdownguide.org/cheat-sheet/)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the follwing code cells to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab04.ipynb')
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Defining functions

Start this lab by learning how to create a very simple function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

---

### `def`

It always starts with `def` (short for **def**ine):

``` python
def
```

---

### Name

Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:

``` python
def to_percentage
```

---

### Signature

Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

``` python
def to_percentage(proportion): 
```

If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after the name. 

We put a colon after the signature to tell Python it's over. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

``` python
def to_percentage(proportion): 
```

---

### Documentation

Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
```

---

### Body

Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab or 4 spaces**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
    factor = 100
```

---

### `return`

The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

``` python
def to_percentage(proportion):
    """Converts a proportion to a percentage."""
    factor = 100
    return proportion * factor
```
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. If present, `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement.

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

`return` and `print` are **very** different. 

---

### Task 01 📍

1. Define `to_percentage` in the cell below.
2. Call your function to convert the proportion 0.2 to a percentage and name the output `twenty_percent`.


```python
def ...
    """"
    Converts a proportion to a percentage.
    
    >>> to_percentage(0.25)
    25.0
    """
    return ...
    
twenty_percent = ...
twenty_percent

```

```python
grader.check("task_01")
```

Like you’ve done with built-in functions in previous labs (max, abs, etc.), you can pass in named values as arguments to your function.

---

### Task 02 📍

Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

*Note:* You don't need to define `to_percentage` again!  Like other named values, functions stick around after you define them.


```python
a_proportion = 2 ** (0.5) / 2
a_percentage = ...
a_percentage
```

```python
grader.check("task_02")
```

Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`.

As we've seen with built-in functions, functions can also take strings (or arrays, or tables) as arguments, and they can return those things, too.

---

### Task 03 📍

Define a function called `disemvowel`.  It should take a single string as its argument.  (You can call that argument whatever you want.)  It should return a copy of that string, but with all the characters that are vowels removed.  (In English, the vowels are the characters "a", "e", "i", "o", and "u".) You can use as many lines inside of the function to do this as you’d like.

*Hint:* To remove all the "a"s from a string, you can use `that_string.replace("a", "")`.  The `.replace` method for strings returns a new string, so you can call `replace` multiple times, one after the other. 


```python
def disemvowel(a_string):
    """
    Removes all vowels from a string.
    
    >>> disemvowel('datascience')
    'dtscnc'
    """ 
    ...

# Checking to see if your function works for an example string.
disemvowel("Can you read this without vowels?")
```

```python
grader.check("task_03")
```

---

### Calls on calls on calls

Just as you can write a series of lines to build up a complex computation, it is useful to define a series of small functions that build on each other. Since you can write any code inside the body of a function, you can call other functions previously defined there.

If we think of a function as a recipe, defining a function in terms of other functions is like having a recipe for cake that calls on another recipe to make the frosting, and another to make the jam filling. This makes the cake recipe shorter and clearer, and it avoids having a bunch of duplicated frosting recipes. It's a foundation of productive programming.

For example, suppose you want to count the number of characters that are not vowels in a piece of text. One way to do this is to remove all of the vowels and find the length of the remaining string.

---

### Task 04 📍

Write a function called `num_nonvowels`.  It should take a string as its argument and return a number.  That number should be the number of characters in the argument string that aren't vowels. You should use the `disemvowel` function you wrote above inside of the `num_nonvowels` function.

*Hint:* The function `len` takes a string as its argument and returns the number of characters in it.


```python
def num_nonvowels(a_string):
    """
    The number of characters in a string, minus the vowels.
    
    >>> num_nonvowels('datascience')
    
    6
    """
    ...

# Try running an example function call of num_nonvowels below to test out your code.
...
```

```python
grader.check("task_04")
```

Functions can also encapsulate code that *displays output* instead of computing a value. For example, if you call `print` inside a function, and then call that function, something will get printed.

---

## Functions and Salaries

Next, we'll have you look at the most recent (2023) salaries for the top 100 highest-paid employees for the State of California according to [OpenPayrolls](https://openpayrolls.com/rank/highest-paid-employees/california). The salaries are stored in a file called `top_100_salaries.csv`.

```python
salaries_raw = Table.read_table('top_100_salaries.csv')
salaries_raw
```

Computing the average of the salaries might seem like a simple task, but there are some issues with this data that make that task a bit challenging. View the error produced by the following code:

<img src="./ave_error.png" alt="The UFuncTypeError the occurs from running np.average(salaries.column('Pay'))">

Let's examine why this error occurred by looking at the values in the `"Pay"` column. The first value in the `"Pay"` column is Vernon Steiner's pay in 2023, in dollars. Look at that value by running the following code cell:

```python
steiner_pay_string = salaries_raw.column("Pay").item(0)
steiner_pay_string
```

You should see quotation marks! Run the following code cell to check the datatype of the information in the `"Pay"` column.

```python
type(steiner_pay_string)
```

The data is actually stored in the table as a string. It doesn't make sense to take the average of string values, so we need to convert them to numbers if we want to do this. How can you do all of that? Here is a guide:

- The string method `replace` will be useful for removing the dollar sign and comma; it replaces a specified substring from a string with another substring. For example, the value of `"100%".replace("%", "")` is the string `"100"`.  
- You'll also need the function `float`, which converts a string that looks like a number to an actual number.  

Run the following cell to put those steps in one command:

```python
steiner_pay = float(steiner_pay_string.replace("$", "").replace(",", "")) # SOLUTION
steiner_pay
```

To compute the average pay, we need to do this for every employee.  But that looks like it would involve copying this code 100 times.

This is where functions come in.  First, we'll define a new function, giving a name to the expression that converts "Pay" strings to numeric values.  Later in this lab, we'll see the payoff: we can call that function on every pay string in the dataset at once.

---

### Task 05 📍

Complete the following function that should take in a `pay_string` value that is a `str` and return the converted pay string value in dollars as a `float`. Use the expression above (`float(steiner_pay_string.replace("$", "").replace(",", ""))`) that we used to convert `steiner_pay` in the body of the function template we've provided, but make sure you replace the specific `steiner_pay_string` with the generic `pay_string` name specified in the first line in the `def` statement.

*Hint*: When dealing with functions, you should generally not be referencing any variable outside of the function. Usually, you want to be working with the arguments that are passed into it, such as `pay_string` for this function. If you're using `steiner_pay_string` within your function, you're referencing an outside variable!


```python
def convert_pay_string_to_number(pay_string):
    """
    Converts a pay string like '$100,000.00' to a number of
    dollars.
    
    >>> convert_pay_string_to_number("$100,000.00")
    100000.00
    """
    converted_pay_string = ...
    return converted_pay_string
```

```python
grader.check("task_05")
```

Running that cell doesn't convert any particular pay string. Instead, it creates a function called `convert_pay_string_to_number` that can convert *any* string with the right format to a number representing the provided salaries.

We can call our function just like we call the built-in functions we've seen. It takes one argument, a string, and it returns a float.

```python
convert_pay_string_to_number('$639,325.06')
```

```python
convert_pay_string_to_number(steiner_pay_string)
```

We can also compute Christopher Ailman's pay in the same way. Run the following cell to see the results.

```python
convert_pay_string_to_number(
    salaries_raw.where("Name", are.containing("Christopher Ailman"))
                .column("Pay")
                .item(0)
)
```

So, what have we gained by defining the `convert_pay_string_to_number` function? 
Well, without it, we'd have to copy the code `float(some_pay_string.replace("$", "").replace(",", ""))` each time we wanted to convert a pay string.  Now we just call a function whose name says exactly what it's doing.

---

## `apply`ing functions

The `apply` table method is used to call a function many times; once on each element in a column of a table. It produces an array of the function's outputs. Below, we use `apply` to convert every salaries (as a string) to a number, using the function `convert_pay_string_to_number` defined above.

```python
salaries_raw.apply(convert_pay_string_to_number, "Pay")
```

#### Caution Using `apply`

Notice that we passed the **name** of the function `convert_pay_string_to_number` as an argument to the `apply` table method. Giving `apply` the name of the function and the column(s) of the table that contain the relevant data to apply the function to is all that is needed. We didn't call the function using parentheses. For example, here are a few common errors:

<img src="./type_error.png" alt="TypeError message that occurs when providing a function call as an argument for the apply table method.">
<img src="./value_error.png" alt="ValueError message that occurs when providing a function call as an argument for the apply table method.">

---

### `salaries` Table

### Task 06 📍

Using `apply`, make a table that's a copy of `salaries_raw` with the values in the `Pay` column, containing the result of applying `convert_pay_string_to_number` to the `Pay` column (as we did above).  Call the new table `salaries`.


```python
salaries = salaries_raw.with_column(
    "Pay",
    ...
    ) 
salaries
```

```python
grader.check("task_06")
```

Now that we have all the pays expressed as numbers, we can finally calculate the average salary from earlier.

---

### Task 07 📍

Compute the average pay of the employees in the dataset.

```python
average_total_pay = ...
average_total_pay
```

```python
grader.check("task_07")
```

---

### Why is `apply` useful?

For operations like arithmetic, or the functions in the NumPy library, you don't need to use `apply`, because they automatically work on each element of an array.  But there are many things that don't.  The string manipulation we did in today's lab is one example.  Since you can write any code you want in a function, `apply` gives you total control over how you operate on data.

---

## Aggregation

Aggregation refers to the process of summarizing and condensing large datasets into more manageable and meaningful insights. You've seen this in action through using functions like `np.average` and `np.sum`. In the context of the `datascience` library, more advanced aggregation can be achieved using table methods like `group` and `pivot`. 

* The `group` method allows you to group data by specific criteria, such as categories or attributes, and then apply aggregation functions (e.g., `sum`, `mean`, etc.) to calculate summary statistics within each group.
* Pivot tables with the `pivot` method, on the other hand, enable you to restructure data to create a compact summary table, making it easier to analyze relationships between variables.

These methods are essential for extracting valuable information and patterns from complex datasets.

For the remaining tasks you will use a hypothetical data set about customers of a store called `customer_data`. Run the following cell to load the table to preview it.

```python
customer_data  = Table().read_table('customer_data.csv')
customer_data
```

---

### Group

The `group` table method has the general format `tbl.group(column_or_label, collect)`. You've worked with the group method without providing a function name for `collect`. The `collect` argument gives you a place to specify a function that you want to apply to all the values associated with the grouped labels.  

By default, the `group` method counts up the number of rows in the table associated with the grouped values for the specified column. If you wanted to do something else besides counting the number of rows, then you'd provide a function name for `collect`. For example, `tbl.group('Column Label', np.min)` would apply the minimum function to all the grouped data based on the category values in the column `'Column Label'`. 

Run the following code to see how this works when we apply `np.min` to the grouped data from the previous task based on `'Name'` values. 

```python
customer_data.group('Name', np.min)
```

It is important to know a few things about how using `group` with a collect function works.
* The `np.min` function attempted to apply to every column of the table other than `'Name'`. In this case, if the column contains numerical values, it will return the minimum of the grouped values for each row. If the column doesn't contain numerical values, then it will return an empty value (The empty string `''` for example).
* The columns of the resulting table will be `'Name'` and a column for each of the other columns in `customer_data` with the function name `min` added to the end of the label.

You can reduce the number of columns using `select` or `drop` before you use the group method to get a more presentable table.

---

### Task 08 📍

Using the `customer_data` table, create a table called `customer_averages` that shows the name and average purchase amount for each customer in the `customer_data` table. Your resulting table should have 3 rows and 2 columns. The column labels should be `'Name'` and `'Amount ($) average'`.

```python
customer_averages = ...
customer_averages
```

```python
grader.check("task_08")
```

---

Now that you have the purchase amount averages by customer, you might want to take this a step further by doing the same and considering whether or not the purchases were made using advertisements.

You can group by more than one column using the `group` method by providing a list of column labels or indexes for the first argument of `group`. The format is `tbl.group(['Label 1', 'Label 2'], collect)`. This creates a table with a row for each unique grouping of values from both columns and calculates the averages for the grouped data (if `average` were in the collect argument).

---

### Task 09 📍

Using the `customer_data` table, create a 4-row and 3-column table called `customer_ad_averages` that contains a row for each combination of customer name and ad (Yes/No) pairing along with the average purchase amount. The labels should be `'Name', 'Ad', and 'Amount ($) average'`.

**Tip**: Don't forget that you will probably want to reduce the `customer_data` table to the relevant columns first.

```python
customer_ad_averages = ...
customer_ad_averages
```

```python
grader.check("task_09")
```

You should see that customer Alice Johnson spends about 8 dollars more on average for purchases not made through advertisements compared to purchases made through advertisements. This is a small example of the more detailed analysis you can provide when aggregating your data across multiple categories.

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Using the `customer_data` table, produce an overlaid histogram that visualizes the distributions of all purchases made through advertisements and those not made through advertisements. Use the `group` argument with the [`hist` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.hist.html#datascience.tables.Table.hist) to accomplish this. Additionally, use the given bins in `useful_bins` by utilizing the `bins` argument for the `hist` table method.

```python
equal_bins = np.arange(5, 25, 1)
...

# Leave the following code to add a title to your histogram
plt.title('Distribution of Amount ($)')
plt.show()
```

<!-- END QUESTION -->

---

### Pivot

A pivot table provides the same information you get from grouping by two columns, but the format is slightly different.

Here are a few examples where using pivot over group can be advantageous:

* Pivot tables are excellent for creating summary tables. If your goal is to create a structured summary table with row and column headers that represent specific categories, pivot is often more straightforward to use.
* If your primary goal is to visualize data, pivot tables can be more useful. The resulting grid structure from a pivot table can be directly used in data visualization tools or libraries for creating bar charts or other visual representations of your data.
* Pivot tables in the `datascience` library handle missing values differently compared with `group`. When using `pivot`, a placeholder of some kind is provided for missing combinations of values. On the other hand, `group` will leave out any missing combinations of values from the table.

The first two arguments of `pivot` are important. 
* The first argument specifies the column where the pivot table column values will come from.
* The second argument will specify the column where the pivot table row values in the first column come from.

By default, the values inside the table come from the counts of how often the pair of column and row values occur in the data set. 

For example, run the following cell  to see the pivot table showing the number of purchases made by each customer through an advertisement or not. Notice that the `'Ad'` values appear as column labels.

```python
customer_data.pivot('Ad', 'Name')
```

You can go further by specifying a third and fourth argument: `values` and `collect`. `values` identifies the column in the table to which the `collect` function will be applied in order to create the values inside the pivot table.

---

### Task 11 📍

Using the `customer_data` table, create a pivot table `ad_name_pivot` that has a row for each customer name, a column for the `'Yes'` and `'No'` `'Ad'` values, and the average purchase amount for each pairing of customers and purchases made through advertisements or not. You won't need to reduce the `customer_data` table like you needed to do in previous tasks. This is another benefit of the `pivot` method!

**Hint**: Use `'Amount ($)'` for the values and `'np.average'` for the collect function.

```python
ad_name_pivot = ...
ad_name_pivot
```

```python
grader.check("task_11")
```

As you can see there are a lot of ways to organize and summarize data in a table. You'll gain a lot of experience with using the tools in this lab throughout the semester.

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: None)
```
>>> to_percentage(0.25) == 25.0
True
```
failure_message: ❌ Make sure you are mulitplying the input proportion 100.
success_message: ✅ It seems like your function is working!


Tests for task_02 (points: None)
```
>>> a_percentage == 70.71067811865476
True
```
failure_message: ❌ It doesn't seem like you are calling the function to_percentage on a_proportion.
success_message: ✅ Great work calling the function on the value named a_proportion.


Tests for task_03 (points: None)
```
>>> disemvowel('Datascience rules!') == 'Dtscnc rls!'
True
```
failure_message: ❌ It doesn't seem like disemvowel is wokring correctly. a_string.replace('t', '') will remove any t's in the string named a_string.
success_message: ✅ It looks like disemvowel is working correctly.


Tests for task_04 (points: None)
```
>>> num_nonvowels('datascience') == 6
True
```
failure_message: ❌ It doesn't seem like your function num_nonvowels is working correctly. Try using len on the output of disemvowel.
success_message: ✅ It seems like your function num_nonvowels is working correctly!


Tests for task_05 (points: None)
```
>>> convert_pay_string_to_number('$100,000.00') == 100000
True
```
failure_message: ❌ convert_pay_string_to_number is not working for our example. Try to redo the steps from a previous task, but use the name pay_string for the string.
success_message: ✅ It looks like your function convert_pay_string_to_number is working with the input '$100,000.00'.
```
>>> convert_pay_string_to_number('$639,325.06') == 639325.06
True
```
failure_message: ❌ convert_pay_string_to_number is not working for our example.
success_message: ✅ It looks like your function convert_pay_string_to_number is working with the input '$639,325.06'.


Tests for task_06 (points: None)
```
>>> salaries.num_rows == salaries.num_rows
True
```
failure_message: ❌ salaries doesn't have the same number of rows as salaries.
success_message: ✅ salaries has the same number of rows as salaries.
```
>>> salaries.labels == salaries.labels
True
```
failure_message: ❌ salaries doesn't have the same column labels as salaries. Make sure you are updating the values in the 'Pay' column, not creating a new column.
success_message: ✅ salaries has the same column labels as salaries.
```
>>> t = salaries.sort('Pay', descending=True)
>>> t.column('Pay').item(0) == 639325.06
True
```
failure_message: ❌ You have the correct column in your table, but the values may be wrong. Try using the apply method with convert_pay_string_to_number function. You'll need to add that information to the table using with_column.
success_message: ✅ We checked a value in the table and it looks correct!


Tests for task_07 (points: None)
```
>>> import math
>>> math.isclose(average_total_pay, 447967.3620999999, rel_tol=0.1)
True
```
failure_message: ❌ It doesn't seem like you have the correct value. Try using np.mean or np.average.
success_message: ✅ Great work!


Tests for task_08 (points: None)
```
>>> isinstance(customer_averages, Table)
True
```
failure_message: ❌ customer_averages should be a Table. Make sure to use the group method.
success_message: ✅ customer_averages is a Table
```
>>> customer_averages.num_rows == 14
True
```
failure_message: ❌ customer_averages does not have the correct number of rows. Double check that you are using the customer_data table.
success_message: ✅ customer_averages has the correct number of rows.
```
>>> customer_averages.num_columns == 2
True
```
failure_message: ❌ customer_averages doesn't have the correct number of columns. Try using select or drop first before grouping the data.
success_message: ✅ customer_averages has the correct number of columns
```
>>> customer_averages
Name            | Amount ($) average
Alice Johnson   | 17.75
Bob Grant       | 16.9706
Carl Burger     | 15.2121
Chuck Brown     | 18.3043
David Evans     | 17.9394
Deborah Jones   | 17.6207
Fritz Johnson   | 16.3333
Maria Garcia    | 16.3913
Marianne Bogart | 15.6207
Mohammed Khan   | 14.5909
... (4 rows omitted)
```
failure_message: ❌ customer_averages doesn't seem correct. Check with someone before moving on.
success_message: ✅ customer_averages seems correct.


Tests for task_09 (points: None)
```
>>> isinstance(customer_ad_averages, Table)
True
```
failure_message: ❌ customer_ad_averages should be a Table. Are you using the group method with customer_data?
success_message: ✅ customer_ad_averages is a Table.
```
>>> customer_ad_averages.num_rows == 28
True
```
failure_message: ❌ customer_ad_averages should have 4 rows. Make sure you are working from the customer_data table.
success_message: ✅ customer_ad_averages has the correct number of rows.
```
>>> customer_ad_averages.labels == ('Name', 'Ad', 'Amount ($) average')
True
```
failure_message: ❌ customer_ad_averages doesn't have the correct column labels. Try using select or drop with customer_data before grouping.?
success_message: ✅ customer_ad_averages has the correct column labels.
```
>>> customer_ad_averages
Name          | Ad   | Amount ($) average
Alice Johnson | No   | 22.2857
Alice Johnson | Yes  | 14.2222
Bob Grant     | No   | 19.6667
Bob Grant     | Yes  | 13.9375
Carl Burger   | No   | 16.5714
Carl Burger   | Yes  | 14.2105
Chuck Brown   | No   | 23.1
Chuck Brown   | Yes  | 14.6154
David Evans   | No   | 20.4286
David Evans   | Yes  | 13.5833
... (18 rows omitted)
```
failure_message: ❌ customer_ad_averages doesn't seem correct. Check with someone before moving on.
success_message: ✅ customer_ad_averages seems correct.


Tests for task_11 (points: None)
```
>>> isinstance(ad_name_pivot, Table)
True
```
failure_message: ❌ ad_name_pivot should be a Table. Are you using the pivot method?
success_message: ✅ ad_name_pivot is a Table.
```
>>> ad_name_pivot.labels == ('Name', 'No', 'Yes')
True
```
failure_message: ❌ ad_name_pivot doesn't have the correct column labels. The first argument of pivot should be 'Ad'.
success_message: ✅ ad_name_pivot has the right column labels
```
>>> tuple(ad_name_pivot.column(0)) == ('Alice Johnson', 'Bob Grant', 'Carl Burger', 'Chuck Brown', 'David Evans', 'Deborah Jones', 'Fritz Johnson', 'Maria Garcia', 'Marianne Bogart', 'Mohammed Khan', 'Myrtle Olson', 'Rosie Catenheim', 'Shane Willis', 'Victor Price')
True
```
failure_message: ❌ ad_name_pivot doesn't have the correct values in the first column. The second argument of pivot should be 'Name'.
success_message: ✅ ad_name_pivot has the correct values in the first column to represent the customer names.
```
>>> ad_name_pivot
Name            | No      | Yes
Alice Johnson   | 22.2857 | 14.2222
Bob Grant       | 19.6667 | 13.9375
Carl Burger     | 16.5714 | 14.2105
Chuck Brown     | 23.1    | 14.6154
David Evans     | 20.4286 | 13.5833
Deborah Jones   | 17.7059 | 17.5
Fritz Johnson   | 17.8571 | 15
Maria Garcia    | 18.2    | 15
Marianne Bogart | 17.5333 | 13.5714
Mohammed Khan   | 18.2222 | 12.0769
... (4 rows omitted)
```
failure_message: ❌ ad_name_pivot doesn't seem correct. Check with someone before moving on.
success_message: ✅ ad_name_pivot seems correct.
