<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 13: k-NN

## References

* [Sections 17.0 - 17.6 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/17/Classification.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab13.ipynb')
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## 🐧 Penguins of the Palmer Archipelago

<a href="https://allisonhorst.github.io/palmerpenguins/reference/figures/lter_penguins.png" title="The Palmer Archipelago penguins. Artwork by @allison_horst."><img src="lter_penguins.png" width="50%" alt="The Palmer Archipelago penguins. Artwork by @allison_horst."></a>

[Dr. Kristen Gorman](https://www.uaf.edu/cfos/people/faculty/detail/kristen-gorman.php) is a renowned marine biologist whose research often focuses on penguins, particularly in Antarctica. She was with the [Palmer Station Long Term Ecological Research Program](https://pallter.marine.rutgers.edu/), part of the [US Long Term Ecological Research Network](https://lternet.edu/). From 2007 through 2009, Gorman gathered several measurements from three species of penguins, Adélie [[1](#Citations)], chinstrap [[2](#Citations)], and gentoo [[3](#Citations)], along the Palmer Archipelago near Palmer Station.

<a href="https://en.wikipedia.org/wiki/Palmer_Archipelago" title="Palmer Archipelago
"><img src="./palmer_archipelago.jpeg" alt="Map of the Palmer archipelago" width=30%></a>

We imported the data sets for each penguin species directly from the [Environmental Data Initiative (EDI) Data Portal](https://portal.edirepository.org/nis/home.jsp) and combined them into one file `penguins.csv`. These data are available for use by [CC0 license](https://creativecommons.org/public-domain/cc0/) ("No Rights Reserved").

Run the following code cell to load the data as `penguins`.

```python
penguins = Table.read_table('penguins.csv')
penguins
```

---

## Features

---

In this lab, you will build a kNN classifier to predict a penguin's species based on two of their numerical attributes (features) using Gorman's data set.

---

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

The first thing you want to do is visualize the relationship between several of the pairs of numerical measurements (features) for the penguins to see which pair would make the best classifier. The features you will focus on initially are `'Culmen Length (mm)'`, `'Culmen Depth (mm)'`, `'Flipper Length (mm)'`, `'Body Mass (g)'`. 

The culmen refers to the ridge along the penguin's bill.

<a href="https://twitter.com/allison_horst/status/1270046411002753025" title="Allison Horst X post"><img src="./culmen.jpeg" alt="An illustration of a penguins culmen." width=40%></a>

For this task, create a scatter plot for each of the pairs of features listed above. To help you out, there are 6 ways to pair up the listed features, and we've stored those in an array called `feature_pairs`. 
* Each item in `feature_pairs` contains two strings, corresponding to the relevant column labels in the `penguins` table.
* The code template loops through all the pairs and starts to make the scatter plots by creating the titles.

To complete this task, you need to add the single line of code that makes the scatter plot from the `penguins` table. Use `group='Species'` to overlay the data from the three species on the same graph.

**Notes:** 
* This task does not have an auto-grader, so check your scatter plots with a classmate, a tutor, or the instructor to make sure you're on the right track.
* The code in the template uses Python's `combinations` function from the [`itertools` module](https://docs.python.org/3/library/itertools.html). You are not expected to know about `itertools` for this class. That module simply provides us with an easy way to determine all the unique pairs of the listed features.

```python
from itertools import combinations

features = ['Culmen Length (mm)', 'Culmen Depth (mm)', 'Flipper Length (mm)', 'Body Mass (g)']
feature_pairs = np.array(list(combinations(features, 2)))

for feature_pair in feature_pairs:
    ...
    plt.title(f'{feature_pair.item(0)} and {feature_pair.item(1)}')
    plt.show()
```

<!-- END QUESTION -->

---

### Task 02 📍

Based on the scatter plots, which pair of features do you think best shows separation in the species of penguins?

Assign `a_good_pair` to one of the following integers that represent a good pair of features for showing separation of penguin species (especially for Andélie penguins):

1. Culmen Length (mm) and Culmen Depth (mm)
2. Culmen Length (mm) and Flipper Length (mm)
3. Culmen Length (mm) and Body Mass (g)
4. Culmen Depth (mm) and Flipper Length (mm)
5. Culmen Depth (mm) and Body Mass (g)
6. Flipper Length (mm) and Body Mass (g)

```python
a_good_pair = ...
```

```python
grader.check("task_02")
```

---

## Binary Classification

---

There are three islands (Biscoe, Dream, and Torgersen) represented in this data set. 

Run the following code cell to the distribution of island and species combinations.

```python
(penguins.select('Island', 'Species')
         .group(['Island', 'Species']))
```

Adélie penguins can be found on all three islands, but gentoo and chinstrap are not found on all three islands. In fact, the data shows that gentoo penguins inhabit Biscoe and chinstrap penguins inhabit Dream. This means we could easily differentiate between gentoo and chinstrap penguins based on what island the penguin was observed. However, we cannot just identify an Adélie penguin based on the island they are on. 

For this reason, we will reduce this classification problem to a binary classification and guide you to build a [kNN classifier](https://ccsf-math-108.github.io/textbook/chapters/17/1/Nearest_Neighbors.html#k-nearest-neighbors) that will classify a penguin as Adélie or not Adélie. 

Since you are learning about how kNN works, we'll have you build the classifier in pieces.

---

### Task 03 📍

First, we need to create our binary labels. There are 3 species, but we are just focusing on whether a penguin is an Adélie penguin or not.

1. Write a function `is_adelie` that takes as input the species of a penguin (a `str` value as listed in the `penguins` table) and returns a `bool` value of `True` or `False` depending on whether it is an Adélie penguin or not.
2. Add a column called `Adelie` to the end of the `penguins` table that contains the correct `True`/`False` value for the given penguin.

```python
def is_adelie(species):
    ...

penguins = penguins.with_column(
    'Adelie', ...)

penguins
```

```python
grader.check("task_03")
```

---

## Feature Selection

---

We want you to use the features `'Culmen Length (mm)'` and `'Flipper Length (mm)'` to classify Adélie penguins. 

Run the following code cell to reduce the `penguins` table down to the relevant information for the rest of this lab.

```python
penguins = penguins.select('Adelie', 
                           'Culmen Length (mm)', 
                           'Flipper Length (mm)')
penguins
```

---

## Distance

---

The kNN classifier works by identifying the majority label of the _nearest_ labeled data to the data you are trying to label. To determine the nearest data, we need to measure how far all the data pairs are from the unlabeled data. The (Euclidean) distance ($d$) between two points $(x_1, y_1)$ and $(x_2, y_2)$ was defined for you as $$d = \sqrt{(x_1 - x_2)^2 + (y_1 - y_2)^2}.$$ Use this distance to define the distance between two penguins.

From one of the recent lectures, we created a `distance` function and a `row_distance` function for you. The `distance` function does the above calculation for arrays and the `row_distance` function more specifically calculates the distance between two rows of a table where each row contains only numerical data. 

Run the following cell to define these functions.

```python
def distance(pt1, pt2):
    """Return the distance between two points, represented as arrays"""
    return np.sqrt(sum((pt1 - pt2)**2))

def row_distance(row1, row2):
    """Return the distance between two numerical rows of a table"""
    return distance(np.array(row1), np.array(row2))
```

---

## Training, Validation, and Testing Data

---

The kNN algorithm that you will implement is tied to a specific set of data, referred to as a training set. The training set is formed by using a certain percentage of the sample data. A common percentage is 70\%. 

If 70\% of the sample data is used for the nearest neighbors classification process, then what is the remaining 30\% used for? Half of the remaining sample data, is known as the validation set, and the remaining data as the test set. Each are 15\% of the original data. The validation set is use to "fine-tune" the classifier by choosing the optimal `k`, and the test set is reserved for the evaluation of the classifier. You'll go through these processes throughout the lab.

---

### Task 04 📍

Split the `penguins` data into a training set, and a validation and testing set. (In task 05 you will split the second set further).
- The training set should reflect the training data and consist of a random sample of 70\% of the rows of `penguins`
- The validation and testing set should reflect the remaining rows of `penguins`.
Assign the training set to `train_pg` and the testing set to `validation_and_test_pg`.

**Notes:** 
* Although we can complete the task by shuffling penguins and using the `take` table method, the autograder will only work if you use the `split` table method.
* To help make sure you are doing the tasks correctly, we'll make sure that you always get the same random training and test sets as everyone else. This is done with the code `np.random.seed(2025)`.

```python
np.random.seed(2025) # Keep this for grading consistency.
row_70th = round( ... )
...
# Inspect train table, and validation and test table
print(f'The following training set contains {round(train_pg.num_rows / penguins.num_rows * 100)}% of the rows of the penguins table.') 

display(train_pg)
print(f'\nThe following validation and test set contains {round(validation_and_test_pg.num_rows / penguins.num_rows * 100)}% of the rows of the penguins table.')
display(validation_and_test_pg)
```

```python
grader.check("task_04")
```

---

### Task 05 📍

Split the `validation_and_test_pg` data into a validation set, and a testing set.
- The validation set should consist of a random sample of 15\% of the rows of the original table, `penguins`
- The test sest should consist of a random sample of 15\% of the rows of the original table, `penguins`.
- Assign the validation set to `validation_pg` and the testing set to `test_pg`.

**Notes:**
* Make sure to use the `split` table method, like we did in task 04.
* To help make sure you are doing the tasks correctly, we'll make sure that you always get the same random training and test sets as everyone else. This is done with the code `np.random.seed(2025)`.

```python
np.random.seed(2025) # Keep this for grading consistency.
row_50th = round( ... )  
...
# Inspect validation table and test table
print(f'The following validation set contains {round(validation_pg.num_rows / penguins.num_rows * 100)}% of the rows of the penguins table.') 

display(validation_pg)
print(f'\nThe following test set contains {round(test_pg.num_rows / penguins.num_rows * 100)}% of the rows of the penguins table.')
display(test_pg)
```

```python
grader.check("task_05")
```

---

## An Example Point

---

Your goal is to be able to classify a penguin as Adélie or not based on the labeled data in Gorman's data set. Let's use the first penguin listed in the `test` table as an example to use for a few tasks. That penguin should have a culmen length of 36.7 mm and a flipper length of 193 mm. You will now make and use a kNN classifier to predict the Adélie label for that penguin. You will be focusing on using table rows to do this. We know the true label for this penguin is `True`, but for now, let's ignore it and see how the classifier works. So an unlabeled example row for that penguin would be:

```python
example_penguin = test_pg.drop('Adelie').row(0)
example_penguin
```

Notice where this example penguin fits within the collected data in `penguin`.

```python
penguins.scatter('Culmen Length (mm)', 'Flipper Length (mm)', group='Adelie')
plt.scatter(example_penguin.item(0), example_penguin.item(1), marker='*', s=100, color='red')
plt.show()
```

Visually, it seems that the unlabeled example is an Adélie penguin because it is near the points of other Adélie penguins. Now, teach the computer how to see that as well!

---

## Building a Classifier

---

### Task 06 📍

For this task, create a function called `distances`. It should:
* Take the following as input:
    * `unlabeled_example` (`Row`): a row object representing an unlabeled example of features,
    * `labeled_features_tbl` (`Table`): a table containing labeled features, like `penguins`, with a column containing `True`/`False` labels, and
    * `label_col` (`str`): a string representing the column in the labeled features table containing the labels.
* It should produce a table (`Table`) as output that contains the same information in `labeled_features_tbl` along with an additional column labeled `'Distance'` containing the row distances between a row of features called `unlabeled_example` and every row in `labeled_features_tbl`. 

```python
def distances(unlabeled_example, labeled_features_tbl, label_col):
    distances_array = make_array()
    features = labeled_features_tbl.drop(...)
   
    for row in features.rows:
        distances_array = np.append(distances_array, ...)
    
    labeled_features_tbl = labeled_features_tbl.with_column('Distance', ...)
    return ...

# Here is the results for the unlabeled example:
distances(example_penguin, penguins, 'Adelie').show(3)
```

```python
grader.check("task_06")
```

---

### Task 07 📍

With a way to calculate distances between penguins based on their numerical attributes, the next part of the kNN process is to find the `k` nearest points. 

Define a function called `closest`.
1. The function should have 4 input arguments:
    * `labeled_features_tbl`: a table of labeled features (e.g. `penguins`)
    * `unlabeled_example`: an example row (e.g. `example_penguin`)
    * `label_col`: a column label that indicates which columns contain the labels (e.g. `'Adelie'`)
    * `k`: an odd integer that indicates the number of nearest points to consider (e.g. `5`)
2. If the value of `k` is even, the function should return an error. Otherwise, the function should return a table with the `k` nearest rows in the `labeled_features_tbl` table to the row `unlabeled_example`.

**Note:** We'll help you raise the error since that is not part of this class; however the if/then programming logic is part of this class. 🤓

```python
def closest(unlabeled_example, labeled_features_tbl, label_col, k):
    '''
    Return a table of the k (odd) closest neighbors to the unlabeled example
    '''
    if ...:
        # Raise an error since k should be odd
        raise ValueError("The value of k should be odd.")
    else:
        # Use the value of k to make a table of the k closest rows
        distances_tbl = ...
        distances_tbl_sorted = ...
        nearest_k = ...
        return ...

# Test even and odd values of k to see if your function is working

k = ... # Try using k = 4 and k = 5
try:
    result = closest(example_penguin, penguins, 'Adelie', k)
    display(result)
except ValueError as the_error:
    display(the_error)
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Using the fact that `closest` returns the `k` closest rows in the data table to the provided example, create a function called `majority_class` that takes a table (`top_k_closest`) as input and returns the majority label associated with those `k` rows.

```python
def majority_class(top_k_closest):
    ...
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Finally, using the `majority_class` and `closest` functions, create your kNN classifier as a function called `knn_classifier`. 

1. The function should have 4 input arguments:
    * `labeled_features_tbl`: a table of labeled features (e.g. `penguins`)
    * `unlabeled_example`: an example row (e.g. `example_penguin`)
    * `label_col`: a column label that indicates which columns contain the labels (e.g. `'Adelie'`)
    * `k`: an integer value that indicates the number of nearest points to consider (e.g. `5`)
2. The function should return a `bool` value for whether the unlabeled example is an Adélie penguin or not using the kNN algorithm.

```python
def knn_classifier(unlabeled_example, labeled_features_tbl, label_col, k):
    ....
```

```python
grader.check("task_09")
```

---

## Classifier Accuracy

---

Great work building a classifier! It did a great job classifying the first row of the test set, but how good of a job does it do in general?

We used `k=5` for our classifier above, but is that the optimal value? We will answer that question in task 10 using the validation set. Then, for an optimal `k` we want to test to see how well it works.

A standard way to analyze how well a classifier does is to measure its accuracy in predicting the labels of the data in the test set. You can do this by:

1. Using the classifier (including the training data) to predict the labels for the data in the test set.
2. Define the accuracy of the model to be the proportion of the correctly labeled data points in the testing data set.

Over the next few tasks, we'll help you define the accuracy of your classifier.

---

### Task 10 📍

Now, write a function to calculate the accuracy of your kNN classifer `knn_classifier`, on a training table called `train`, a testing table called `test`, and an integer k-value called `k`. Assume that the columns of `train` contain a label column called `Adelie`. We'll guide you through this in a few steps:
1. Use the `train` set with the `knn_classifier` and `k` to create an array `predicted_labels` of the predicted labels for all the rows of the `test` table.
2. Create a table called `test_with_predictions` that is copy of `test` with an added column called `'Predicted Adelie'` containing the predictions in `predicted_labels`.
3. Create an array called `success` that contains `bool` values showing whether or not the predicted label was correct for each row of the test set.
4. Assign `accuracy` to the proportion of `True` values in `success`.
5. The function should output the 'accuracy'.

The code cell will print the accuracy of our kNN classifier on our `train_pg` and `test_pg` sets with `k=3`.

```python
def evaluate_accuracy(train, test, k):

    # Step 1
    predicted_labels = make_array()

    for ...:
        ...

    # Step 2
    test_with_predictions = ...

    # Step 3
    success = test_with_predictions.column('Adelie') == ...

    # Step 4
    accuracy = ...

    ...

#Do not change anything below this comment
accuracy_k_3 = evaluate_accuracy(train_pg, test_pg, 3)
print(f'The accuracy of your kNN classifer (with k = 3) using your randomly \
generated train and test sets is: {accuracy_k_3*100:.02f}%.')
```

```python
grader.check("task_10")
```

---

Now, when you apply the classifier you built to a new data point (pair of culmen lengths and flipper length measurements), you can more confidently believe it will correctly label the species of the penguin.

---

Fine-tuning existing classifiers to improve performance metrics like accuracy is a big part of machine learning. Since kNN classifiers use all the available data to make a prediction, there are not many parameters to adjust to improve accuracy. 

For this task, we will try the following idea to improve the accuracy of your classifier: Finding an optimal value for `k` on the validation set. In this class, that means try some different values of `k` for the classifier using the validation set.

* Start by running the following cell. In it we define a function `whick_k` that inputs a `training` table, a `validation` table, and an array of integers called `k_values`. 
* The function evaluates the accuracy of the classifier on the validation set, for each of the integers in `k_values`. 
* The function will output a table with two columns. The first column is the array of `k_values`, called `'k'`, and the second column is the corresponding accuracy, called `'Accuracy'`.
* The output of the code cell should be your function acting on the sets `train_pg`, `validation_pg`, and the provided array called `some_odd_k_s`. The code cell should also output a line plot of the table.

```python
#Just run this cell and don't change it
def which_k(training, validation, k_values):

    accuracies = make_array()

    for i in np.arange(np.size(k_values)):

        accuracy = evaluate_accuracy(training, validation, k_values.item(i))
        accuracies = np.append(accuracies, accuracy)

    k_table = Table().with_columns("k", k_values,
                                   "Accuracy", accuracies)
                             
    return k_table
    
some_odd_k_s=make_array(1,3,5,7,9,11,13,15,17,19)
which_k(train_pg, validation_pg, some_odd_k_s).show()
which_k(train_pg, validation_pg, some_odd_k_s).plot('k')
```

### Task 11 📍

Based on the line plot and table above, which value of `'k'` is optimal to use?

Assign `optimal_k` to one of the values of `k` that are optimal choices.

```python
optimal_k = ...
```

```python
grader.check("task_11")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Citations

1. Palmer Station Antarctica LTER and K. Gorman. 2020. Structural size measurements and isotopic signatures of foraging among adult male and female Adélie penguins (Pygoscelis adeliae) nesting along the Palmer Archipelago near Palmer Station, 2007-2009 ver 5. Environmental Data Initiative. https://doi.org/10.6073/pasta/98b16d7d563f265cb52372c8ca99e60f (Accessed 2024-05-02).

2. Palmer Station Antarctica LTER and K. Gorman. 2020. Structural size measurements and isotopic signatures of foraging among adult male and female Chinstrap penguins (Pygoscelis antarcticus) nesting along the Palmer Archipelago near Palmer Station, 2007-2009 ver 8. Environmental Data Initiative. https://doi.org/10.6073/pasta/ce9b4713bb8c065a8fcfd7f50bf30dde (Accessed 2024-05-02).

3. Palmer Station Antarctica LTER and K. Gorman. 2020. Structural size measurements and isotopic signatures of foraging among adult male and female gentoo penguins (Pygoscelis papua) nesting along the Palmer Archipelago near Palmer Station, 2007-2009 ver 7. Environmental Data Initiative. https://doi.org/10.6073/pasta/9fc8f9b5a2fa28bdca96516649b6599b (Accessed 2024-05-02).

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_02 (points: None)
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash({a_good_pair, 1, 2})
'190cda75e2bd9dc669b09abe9f1131c5'
```
failure_message: ❌ We don't agree with your choice. There are 2 other pairs of features that better show separation in the species. Ask for help seeing this.
success_message: ✅ We agree with your choice! We don't feel there is a clear winner, but there are two pairs that do a good job at showing the separation of the species types.


Tests for task_03 (points: None)
```
>>> is_adelie('Adelie Penguin (Pygoscelis adeliae)')
True
```
failure_message: ❌ is_adelie doesn't correctly label an Adélie penguin. Try checking the species value to see if it contains the substring 'Adelie'.
success_message: ✅ is_adelie correctly labels an Adélie penguin
```
>>> is_adelie('Chinstrap penguin (Pygoscelis antarctica)')
False
```
failure_message: ❌ is_adelie doesn't correctly label a chinstrap penguin. If the species name doesn't contain 'Adelie', then you want to return False.
success_message: ✅ is_adelie correctly labels an chinstrap penguin
```
>>> is_adelie('Gentoo penguin (Pygoscelis papua)')
False
```
failure_message: ❌ is_adelie doesn't correctly label a gentoo penguin. If the species name doesn't contain 'Adelie', then you want to return False.
success_message: ✅ is_adelie correctly labels an gentoo penguin
```
>>> 'Adelie' in penguins.labels
True
```
failure_message: ❌ You have not added the Adelie column to the penguins table. Try using the apply table method to get an array of True/False values and then using the with_column table method.
success_message: ✅ You've added the Adelie column to the penguins table.
```
>>> np.sum(penguins.column('Adelie')) == 151
np.True_
```
failure_message: ❌ Your updated penguins table doesn't seem correct. Ask for help.
success_message: ✅ Your updated penguins table seems correct.


Tests for task_04 (points: None)
```
>>> isinstance(train_pg, Table)
True
```
failure_message: ❌ train_pg is not assigned to a Table.
success_message: ✅ train_pg is assigned to a Table.
```
>>> isinstance(validation_and_test_pg, Table)
True
```
failure_message: ❌ validation_and_test_pg is not assigned to a Table.
success_message: ✅ validation_and_test_pg is assigned to a Table.
```
>>> np.isclose(train_pg.num_rows / penguins.num_rows, 0.698, atol=0.001)
np.True_
```
failure_message: ❌ The training set should contain approximately 70 percent of the data.
success_message: ✅ Your training set contains approximately 70 percent of the data.
```
>>> np.isclose(validation_and_test_pg.num_rows / penguins.num_rows, 0.301, atol=0.001)
np.True_
```
failure_message: ❌ The validation and test set should contain approximately 30 percent of the data.
success_message: ✅ Your validation and test set contains approximately 30 percent of the data.
```
>>> train_pg
Adelie | Culmen Length (mm) | Flipper Length (mm)
False  | 42.4               | 181
False  | 46.8               | 215
True   | 36.6               | 184
True   | 40.1               | 188
False  | 49.8               | 229
True   | 37.6               | 194
False  | 49.1               | 212
False  | 48.7               | 208
True   | 41.4               | 202
True   | 37.7               | 198
... (229 rows omitted)
```
failure_message: ❌ train_pg doesn't seems correct. Make sure not to change np.random.seed(2025)
success_message: ✅ train_pg seems correct.
```
>>> validation_and_test_pg
Adelie | Culmen Length (mm) | Flipper Length (mm)
False  | 49.1               | 220
False  | 46.5               | 217
False  | 51.3               | 197
True   | 40.2               | 193
False  | 50.2               | 218
False  | 49.7               | 195
True   | 40.9               | 191
False  | 50.2               | 202
False  | 45.8               | 210
True   | 42                 | 190
... (93 rows omitted)
```
failure_message: ❌ validation_and_test_pg doesn't seems correct. Make sure not to change np.random.seed(2025).
success_message: ✅ validation_and_test_pg seems correct


Tests for task_05 (points: None)
```
>>> isinstance(validation_pg, Table)
True
```
failure_message: ❌ validation_pg is not assigned to a Table.
success_message: ✅ validation_pg is assigned to a Table.
```
>>> isinstance(test_pg, Table)
True
```
failure_message: ❌ test_pg is not assigned to a Table.
success_message: ✅ test_pg is assigned to a Table.
```
>>> np.isclose(validation_pg.num_rows / penguins.num_rows, 0.152, atol=0.001)
np.True_
```
failure_message: ❌ The validation set should contain approximately 15 percent of the original data.
success_message: ✅ Your validation set contains approximately 15 percent of the original data.
```
>>> np.isclose(test_pg.num_rows / penguins.num_rows, 0.1491, 0.001)
np.True_
```
failure_message: ❌ The test set should contain approximately 15 percent of the original data.
success_message: ✅ Your test set contains approximately 15 percent of the original data.
```
>>> validation_pg
Adelie | Culmen Length (mm) | Flipper Length (mm)
True   | 36.5               | 182
True   | 37                 | 185
False  | 48.4               | 213
True   | 38.7               | 195
True   | 39                 | 191
True   | 35.9               | 189
True   | 40.9               | 191
False  | 46.8               | 189
True   | 38.9               | 190
False  | 51.3               | 197
... (42 rows omitted)
```
failure_message: ❌ validation_pg doesn't seems correct. Make sure not to change np.random.seed(2025)
success_message: ✅ validation_pg seems correct.
```
>>> test_pg
Adelie | Culmen Length (mm) | Flipper Length (mm)
True   | 36.7               | 193
True   | 41.1               | 205
False  | 45.8               | 210
False  | 48.6               | 230
True   | 41.1               | 182
False  | 44.4               | 219
True   | 37.2               | 184
True   | 39                 | 185
False  | 43.6               | 217
False  | 50.2               | 202
... (41 rows omitted)
```
failure_message: ❌ test_pg doesn't seems correct. Make sure not to change np.random.seed(2025).
success_message: ✅ test_pg seems correct


Tests for task_06 (points: None)
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> isinstance(distances(example_penguin, penguins.take(0), 'Adelie'), Table)
True
```
failure_message: ❌ Your distances function doesn't return an array. Try returning distances_array.
success_message: ✅ Your distances function returns an array.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> distances(penguins.drop('Adelie').row(0), penguins.take(0), 'Adelie').row(0).item('Distance') == 0
np.True_
```
failure_message: ❌ Your distances function doesn't correctly show that two data points with the same features have a distance of 0. Try using the row_distance function above.
success_message: ✅ Your distances function correctly shows that two data points with the same features have a distance of 0.
```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(np.sum(distances(example_penguin, penguins, 'Adelie').column('Distance')))
'923e65c0e0e4a8aacd6b135f3df2b9f2'
```
failure_message: ❌ Your function doesn't seem to be working with example_penguin. Ask for help before moving on.
success_message: ✅ Your function seems to be working with example_penguin.


Tests for task_07 (points: None)
```
>>> try:
...     result = closest(example_penguin, penguins, 'Adelie', 4)
...     print(result)
... except ValueError as e:
...     print(e)
The value of k should be odd.
```
failure_message: ❌ Your function shouldn't accept even values of k. Remember that 8 % 2 equals 0 and 8 is an even number.
success_message: ✅ Your function doesn't accept even values of k.
```
>>> isinstance(closest(example_penguin, penguins, 'Adelie', 5), Table)
True
```
failure_message: ❌ Your function should return a table. Try running printing type(...) on what you return to see if it is a Table.
success_message: ✅ Your function returns a table.
```
>>> closest(example_penguin, penguins, 'Adelie', 5).num_rows == 5
True
```
failure_message: ❌ Your function should return a table with 5 rows. Try using take
success_message: ✅ Your function returns a table with 5 rows.
```
>>> import numpy as np
>>> def is_sorted(arr):
...     return np.all(arr[:-1] <= arr[1:])
>>> test_closest = closest(example_penguin, penguins, 'Adelie', 5).column('Distance')
>>> is_sorted(test_closest)
np.True_
```
failure_message: ❌ Your function doesn't return a table where the Distance values are sorted in ascending order. Try using the sort table method and use descending=False (the default value).
success_message: ✅ Your function returns a table where the Distance values are sorted in ascending order.
```
>>> closest(example_penguin, penguins, 'Adelie', 5)
Adelie | Culmen Length (mm) | Flipper Length (mm) | Distance
True   | 36.7               | 193                 | 0
True   | 36.8               | 193                 | 0.1
True   | 37.8               | 193                 | 1.1
True   | 37.3               | 192                 | 1.16619
True   | 37.9               | 193                 | 1.2
```
failure_message: ❌ Your function doesn't seem to be working correctly for example_penguin. Ask for help.
success_message: ✅ Your function seems to be working correctly for example_penguin.


Tests for task_08 (points: None)
```
>>> test_result_majority_class = majority_class(closest(example_penguin, penguins, 'Adelie', 5))
>>> test_result_majority_class in penguins.column('Adelie')
True
```
failure_message: ❌ Your majority_class function doesn't return a label from the data. 
success_message: ✅ Your majority_class function returns a label from the data.
```
>>> majority_class(closest(example_penguin, penguins, 'Adelie', 5))
True
```
failure_message: ❌ The kNN classifier with k=5 and this data shouldn't return False.
success_message: ✅ The kNN classifier with k=5 and this data should return True.


Tests for task_09 (points: None)
```
>>> knn_classifier(example_penguin, penguins, 'Adelie', 5)
True
```
failure_message: ❌ knn_classifier should not classify example_penguin as not an Andelie penguin. Ask for help.
success_message: ✅ knn_classifier classifies example_penguin as an Andelie penguin.
```
>>> another_unlabeled_example = Table().with_columns('Culmen Length (mm)', [40.5], 'Flipper Length (mm)', [192]).row(0)
>>> knn_classifier(another_unlabeled_example, penguins, 'Adelie', 5)
True
```
failure_message: ❌ knn_classifier doesn't classify our other text example as an Andelie penguin. Ask for help.
success_message: ✅ knn_classifier classifies our other text example as an Andelie penguin.


Tests for task_10 (points: None)
```
>>> callable(evaluate_accuracy)
True
```
failure_message: ❌ evaluate_accuracy is not a function.
success_message: ✅ evaluate_accuracy is a function.
```
>>> isinstance(evaluate_accuracy(train_pg, test_pg, 3), (float, np.floating))
True
```
failure_message: ❌ evaluate_accuracy does not have the correct type of output.
success_message: ✅ evaluate_accuracy has the correct type of output.
```
>>> np.isclose(accuracy_k_3, 0.961, atol=0.001)
np.True_
```
failure_message: ❌ The accuracy you calculated should be approximately 96%.
success_message: ✅ The accuracy you calculated is approximately%


Tests for task_11 (points: None)
```
>>> optimal_k in {3, 5, 7}
True
```
failure_message: ❌ Incorrect. Look back at the table to see which k-values give the highest accuracy.
success_message: ✅ Correct. Any of the values of k=3, 5, or 7 are reasonable choices since they have the highest accuracy.
