<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 10 - Sampling Distributions

---

## References

* [Sections 14.3 - 14.6 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/14/3/SD_and_the_Normal_Curve.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab10.ipynb')
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Normal Distributions

When we visualize the distribution of a sample, we are often interested in the mean and the standard deviation of the sample (for the rest of this lab, we will abbreviate "standard deviation" as "SD"). These two sample statistics can give us a bird's eye view of the distribution -- by letting us know where the distribution sits on the number line and how spread out it is, respectively. 

First, load in the `births` table from lecture, which is a large random sample of US births and includes information about mother-child pairs.

```python
births = Table.read_table('baby.csv')
births
```

---

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

Write a line of code to plot the distribution of mother's ages from the `births` table. Don't change the last line, which will plot the mean of the sample on the distribution itself. 

**Note**: Since this task does not have an auto-grader, check your solution with a classmate, tutor, or the instructor before moving on.

```python
...

# Do not change anything below this line
plt.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plt.ylim(-0.005, 0.07);
```

<!-- END QUESTION -->

---

### Center of Gravity

In the plot above, we can see that the mean is the center of gravity (or balance point) of the distribution. If you imagine cutting the distribution out of cardboard and placing your finger below the mean (the red dot), the distribution would balance perfectly on your finger. Since the distribution is right-skewed (i.e., it has a long right tail), the mean of the distribution is larger than the median (the "halfway" point of the data). Conversely, if the distribution had been left-skewed, the mean would be smaller than the median.

---

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

Run the following cell to plot the mean (red) and median (blue) on the distribution of the mothers' ages. Then write a sentence describing the relationship between the mean and the median in this distribution.

**Note**: Since this task does not have an auto-grader, check your solution with a classmate, tutor, or the instructor before moving on.

```python
births.hist("Maternal Age")
plt.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50, label="Mean")
plt.scatter(np.median(births.column("Maternal Age")), 0, color='blue', zorder=2, s=50, label="Median")
plt.legend()
plt.ylim(-0.005, 0.07);
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Standard Deviation

We are also interested in the standard deviation (SD) of mother's ages. The SD gives us a sense of how variable mothers' ages are around the average mothers' age. If the SD is larger, then this indicates that the mothers' ages are more spread out around the average than if the SD is smaller. A smaller SD would indicate a more tightly clustered distribution around the average mother's age.

**Recall that the SD of an array is defined as the "root mean square of the deviations", where the deviations are the difference between the original elements and their average.**

You can read more about Variability and SD in [Section 14.2](https://ccsf-math-108.github.io/textbook/chapters/14/2/Variability.html). This chapter also goes a bit more into detail about how "root mean square of the deviations" is calculated. Fun fact! $\sigma$ (Greek letter sigma) is used to represent the SD and $\mu$ (Greek letter mu) is used for the mean. 

---

### Task 03 📍

Assign `age_mean` to the mean of the maternal ages and `age_sd` to the SD of the maternal ages. Then run the cell below to see the width of one SD (blue) from the sample mean (red) plotted on the histogram of maternal ages.

**Hint**: `np.std` might be useful here.

```python
age_mean = ...
age_sd = ...
births.hist("Maternal Age")

plt.scatter(age_mean, 0, color='red', zorder=2, s=50);
plt.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plt.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plt.ylim(-0.005, 0.07);
```

```python
grader.check("task_03")
```

---

### Standard Deviations and Inflection Points

In the histogram above, the standard deviation is not easy to identify if you were just looking at the graph.

However, the distributions of some variables allow us to easily spot the standard deviation on the plot. For example, if a sample follows a **[normal distribution](https://ccsf-math-108.github.io/textbook/chapters/14/3/SD_and_the_Normal_Curve.html)**, the standard deviation is easily spotted using the point of inflection (the point where the curve begins to change the direction of its curvature) of the distribution.

---

### Task 04 📍

Fill in the following code to examine the distribution of maternal heights, which is roughly normally distributed. We'll plot the standard deviation on the histogram, as before. Notice where one standard deviation (blue) away from the mean (red) falls on the plot. Do you see the points of inflection?

```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plt.scatter((height_mean), 0, color='red', zorder=2, s=50);
plt.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plt.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plt.ylim(-0.01, 0.16);
```

```python
grader.check("task_04")
```

---

### Central Limit Theorem

We don't always know how a variable will be distributed, and making assumptions about whether or not a variable will follow a normal distribution is dangerous. However, the Central Limit Theorem (CLT) gives us distributions that always follow a normal distribution. The CLT is a powerful statistical tool because it tells us that the distribution of sample statistics (like means, proportions, and sums) becomes normally distributed as the sample size increases, regardless of the original population's distribution. Remember that the CLT refers to the distribution of a statistic calculated from a distribution, not the distribution of the original sample or population. If this is confusing, ask a staff member!

The next section explores distributions of sample means and how the standard deviation of these distributions depends on sample size.

---

## Variability of the Sample Mean

By the [Central Limit Theorem](https://ccsf-math-108.github.io/textbook/chapters/14/4/Central_Limit_Theorem.html), the probability distribution of the mean of a large random sample drawn with replacement is roughly normal. The bell curve is centered at the population mean. Some of the sample means are higher and some are lower, but the deviations from the population mean are roughly symmetric on either side, as we have seen repeatedly.

In our simulations, we also noticed that the means of larger samples tended to be more tightly clustered around the mean of the distribution from which we sampled than those of smaller samples. Note that in this lab, the data we have access to and sample from is the population (which we rarely have access to), while when we bootstrap, the distribution we sample from is our original sample.

In this section, we will quantify the [variability of the sample mean](https://ccsf-math-108.github.io/textbook/chapters/14/5/Variability_of_the_Sample_Mean.html) and develop a relation between the variability and the sample size.

Let's look at the salaries of 85,782 employees of the City of San Francisco in 2025. **Note that in this lab, this dataset is the population.**

This data was collected from [data.sfgov.org](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd/about_data) where the `salary` column is calculated by summing up regular, overtime, and other pay.

_Note: If you get stuck on any part of this lab, please refer to [Chapter 14 of the textbook](https://ccsf-math-108.github.io/textbook/chapters/14/Why_the_Mean_Matters.html)._

```python
salaries = Table.read_table('sf_salaries_2025.csv')
salaries.show(5)
```

```python
salary_mean = np.mean(salaries.column('salary'))
print('Mean salary of San Francisco city employees in 2025: ', salary_mean.round(2))
```

```python
max(salaries.column(0))
```

```python
salaries.hist('salary', bins=np.arange(0, 400000, 10000))
plt.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plt.title('2025 salaries of city of SF employees')
plt.ylim(-0.0000005, 0.000011);
```

Clearly, the population _does not_ follow a Normal distribution. Keep that in mind as we progress through these exercises.

Let's take random samples **with replacement** and look at the probability distribution of the sample mean. As usual, we will use simulation to get an empirical approximation to this distribution.

---

### Task 05 📍

Define a function `one_sample_mean`. Its arguments should be `table` (the name of a table), `label` (the label of the column containing the variable), and `sample_size`(the number of employees in the sample). It should sample with replacement from the table and
return the mean of the `label` column of the sample.

_Note:_ The function should reference the arguments and be able to work with any table.

```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Use `one_sample_mean` to define a function `simulate_sample_mean`. The arguments are the name of the `table`, the `label` of the column containing the variable, the `sample_size`, and the number of simulations (`repetitions`).

The function should sample with replacement from the table and calculate the mean of each sample. It should save the sample means in an array called `means`. The remaining code in the function displays an empirical histogram of the sample means.

_Note:_ We round all values to 2 decimal places in the provided code because we are working with salaries.

```python
def simulate_sample_mean(table, label, sample_size, repetitions):
    
    means = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means = ...

    sample_means = Table().with_column('Sample Means', means)
    
    # Display empirical histogram and print all relevant quantities – don't change this!
    sample_means.hist(bins=20)
    plt.xlabel('Sample Means')
    plt.title('Sample Size {sample_size}; {repetitions} Resamples'.format(sample_size=sample_size, repetitions=repetitions))
    print(f"Sample size: {sample_size}")
    print(f"Population mean: {np.mean(table.column(label)).round(2)}")
    print(f"Average of sample means: {np.mean(means).round(2)}")
    print(f"Population SD: {np.std(table.column(label)).round(2)}")
    print(f"SD of sample means: {np.std(means).round(2)}")
    return np.std(means)
```

```python
grader.check("task_06")
```

---

### Verify `simulate_sample_mean`

Verify with a peer, a tutor, or the instructor that you've implemented the function above correctly. If you haven't implemented it correctly, the rest of the lab won't work properly, so this step is crucial.  In the following cell, we will create a sample of size 100 from `salaries` and graph it using our new `simulate_sample_mean` function.
 
*Hint: You should see a distribution similar to something we've been talking about. If not, check your function.*

```python
simulate_sample_mean(salaries, 'salary', 100, 10000) 
plt.xlim(130000, 200000);
```

Notice that our distribution of *sample means* looks approximately normal! Did it matter if our original distribution was normally distributed?

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Simulate two sample means, one for a sample of 400 salaries and one for a sample of 625 salaries. In each case, perform 10,000 repetitions. 

**Notes:** 
* Since this task does not have an auto-grader, check your solution with a classmate, tutor, or the instructor before moving on.
* You are not responsible for the `plt.xlim` or `plt.ylim` lines - they just make sure that all of the plt have the same x-axis and y-axis, respectively.


```python
simulate_sample_mean(..., ..., ..., ...)
plt.xlim(150000, 180000)
plt.ylim(0, 0.00014)
plt.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plt.xlim(150000, 180000)
plt.ylim(0, 0.00014)
plt.show();
```

<!-- END QUESTION -->

---

### Task 08 📍

Assign `task8` to an array of numbers corresponding to true statement(s) about the plots from [Task 07](#Task-07-📍🔎). 

*Hint*: If a distribution is bell-shaped, what type of curve does it follow? 

1. We see the Central Limit Theorem (CLT) in action because the distributions of the sample means are bell-shaped.
2. We see the Law of Averages in action because the distributions of the sample means look like the distribution of the population.
3. One of the conditions for CLT is that we have to draw a small random sample with replacement from the population.
4. One of the conditions for CLT is that we have to draw a large random sample with replacement from the population.
5. One of the conditions for CLT is that the population must be normally distributed.


```python
task8 = ...
task8
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Assign `task9` to an array of numbers corresponding to true statement(s) about the plots from [Task 07](#Task-07-📍🔎).

1. Both plots in Task 07 are roughly centered around the population mean.
2. Both plots in Task 07 are roughly centered around the mean of a particular sample.
3. The distribution of sample means for sample size 625 has less variability than the distribution of sample means for sample size 400.
4. The distribution of sample means for sample size 625 has more variability than the distribution of sample means for sample size 400.

```python
task9 = ...
task9
```

```python
grader.check("task_09")
```

---

### Increasing Number of Resamples

Below, we'll look at what happens when we take an **increasing number of resamples of a fixed sample size.** Notice what number in the code changes, and what stays the same. How does the distribution of the resampled means change?

```python
simulate_sample_mean(salaries, 'salary', 100, 500)
plt.xlim(130000, 200000)
plt.ylim(0, 0.00006)
```

```python
simulate_sample_mean(salaries, 'salary', 100, 5000)
plt.xlim(130000, 200000)
plt.ylim(0, 0.00006)
```

```python
simulate_sample_mean(salaries, 'salary', 100, 10000)
plt.xlim(130000, 200000)
plt.ylim(0, 0.00006)
```

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

What did you notice about the distributions of sample means in the three histograms above? Discuss with a peer, a tutor, or your instructor before moving on.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Task 11 📍

Assign the variable `SD_of_sample_means` to the integer corresponding to your answer to the following question:

When I increase the number of _resamples_ that I take, for a _fixed_ sample size, the SD of my sample means will...

1. Increase
2. Decrease
3. Stay about the same
4. Vary widely

```python
SD_of_sample_means = ...
SD_of_sample_means
```

```python
grader.check("task_11")
```

---

### Increasing Sample Size

### Task 12 📍

Let's think about how the relationships between population SD, sample SD, and SD of sample means change with varying sample size. Which of the following is true? Assign the variable `pop_vs_sample` to an array of integer(s) that correspond to true statement(s).

1. Sample SD gets smaller with increasing sample size.
2. Sample SD gets larger with increasing sample size.
3. Sample SD becomes more consistent with population SD with increasing sample size.
4. SD of sample means gets smaller with increasing sample size.
5. SD of sample means gets larger with increasing sample size.
6. SD of sample means stays the same with increasing sample size.

**Hints**: 
* The sample SD is different from the SD of the sample means.
* [Section 14.5](https://ccsf-math-108.github.io/textbook/chapters/14/5/Variability_of_the_Sample_Mean.html) might be helpful for answering this question.

```python
pop_vs_sample = ...
pop_vs_sample
```

```python
grader.check("task_12")
```

---

### Task 13 📍

Is there a relationship between the sample size and the standard deviation of the sample means? Assign `task13` to the number corresponding to the statement that answers this question.

1. The SD of the sample means is inversely proportional to the square root of sample size.
2. The SD of the sample means is directly proportional to the square root of sample size.

_Hint:_ [Chapter 14.5](https://ccsf-math-108.github.io/textbook/chapters/14/5/Variability_of_the_Sample_Mean.html) of the textbook may be helpful.

```python
task13 = ...
task13
```

```python
grader.check("task_13")
```

---

### Bootstrap Samples

**Throughout this lab, we have been taking many random samples from a population** (in contrast to re-sampling from a sample). However, all of these principles hold for bootstrapped resamples from a single sample. The bootstrap works because it's like drawing from the original population, assuming the sample is representative. If your original sample is relatively large, all of your resamples will also be relatively large, and so the SD of resampled means will be relatively small.

It is also important to keep in mind that when doing the bootstrap, our histogram of resample means will be centered around the **original sample mean**, rather than the population mean (as we usually don't have access to the population mean)!

**In order to change the variability of your sample mean, you'd have to change the size of the original sample from which you are taking bootstrapped resamples.**

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_03 (points: None)
```
>>> np.isclose(age_mean, 27.228279386712096)
np.True_
```
failure_message: "❌ age_mean does not seem correct. 
Try using np.mean with the Maternal Age values."
success_message: "✅ age_mean seems correct."

```
>>> np.isclose(age_sd, 5.815360404190897)
np.True_
```
failure_message: "❌ age_sd does not seem correct. 
Try using np.std with the Maternal Age values."
success_message: "✅ age_sd seems correct."



Tests for task_04 (points: None)
```
>>> np.isclose(height_mean, 64.04940374787053)
np.True_
```
failure_message: "❌ height_mean does not seem correct. 
Try using np.mean with the Maternal Height values."
success_message: "✅ height_mean seems correct."

```
>>> np.isclose(height_sd, 2.5250254409674375)
np.True_
```
failure_message: "❌ height_sd does not seem correct. 
Try using np.std with the Maternal Height values."
success_message: "✅ height_sd seems correct."



Tests for task_05 (points: None)
```
>>> callable(one_sample_mean)
True
```
failure_message: "❌ one_sample_mean is not callable. 
Make sure you are using a def statement to create this function."
success_message: "✅ one_sample_mean is callable."

```
>>> import numbers
>>> np.random.seed(1010101)
>>> test_sample_mean = one_sample_mean(salaries, 'salary', 100)
>>> isinstance(test_sample_mean, numbers.Number)
True
```
failure_message: "❌ one_sample_mean should return a number. 
Try calculating a mean in your function body."
success_message: "✅ one_sample_mean returns a number."

```
>>> np.random.seed(1010101)
>>> test_sample_mean = one_sample_mean(salaries, 'salary', 100)
>>> np.isclose(test_sample_mean, 163439.6862)
np.True_
```
failure_message: "❌ one_sample_mean does not seem to be working correctly.
Did you resample the table with replacement and calculate 
the mean using the provided label?"
success_message: "✅ one_sample_mean seems to be working correctly."



Tests for task_06 (points: None)
```
>>> callable(simulate_sample_mean)
True
```
failure_message: "❌ simulate_sample_mean is not callable. 
Make sure you are using a def statement to define this function."
success_message: "✅ simulate_sample_mean is callable."



Tests for task_08 (points: None)
```
>>> isinstance(task8, np.ndarray)
True
```
failure_message: "❌ task8 should be assigned to an array."
success_message: "✅ task8 is assigned to an array"

```
>>> isinstance(task8.item(0), (int, np.integer))
True
```
failure_message: "❌ task8 should be assigned to an array of integers. 
There is at least one correct statement."
success_message: "✅ task8 is assigned to an array of integers"

```
>>> np.all(task8 <= 5) and np.all(task8 >= 1)
np.True_
```
failure_message: "❌ task8 should be assigned to an array of integers between 1 and 5, inclusive."
success_message: "✅ task8 is assigned to an array of integers between 1 and 5, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(set(np.array(task8).astype(int))) == '5a4cba74899c2ad8a3fb49fc72c02e33'
True
```
failure_message: "❌ You have not assigned the correct values to task8."
success_message: "✅ You've assigned the correct values to task8."



Tests for task_09 (points: None)
```
>>> isinstance(task9, np.ndarray)
True
```
failure_message: "❌ task9 should be assigned to an array."
success_message: "✅ task9 is assigned to an array"

```
>>> isinstance(task9.item(0), (int, np.integer))
True
```
failure_message: "❌ task9 should be assigned to an array of integers. 
There is at least one correct statement."
success_message: "✅ task9 is assigned to an array of integers"

```
>>> np.all(task9 <= 5) and np.all(task9 >= 1)
np.True_
```
failure_message: "❌ task9 should be assigned to an array of integers between 1 and 4, inclusive."
success_message: "✅ task9 is assigned to an array of integers between 1 and 4, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(set(np.array(task9).astype(int))) == '5f0ee4e363a0fe5e49f385f38358bb1c'
True
```
failure_message: "❌ You have not assigned the correct values to task9."
success_message: "✅ You've assigned the correct values to task9."



Tests for task_11 (points: None)
```
>>> isinstance(SD_of_sample_means, (int, np.integer))
True
```
failure_message: "❌ You have not assigned SD_of_sample_means to an integer."
success_message: "✅ You've assigned SD_of_sample_means to an integer."

```
>>> SD_of_sample_means in {1, 2, 3, 4}
True
```
failure_message: "❌ You have not assigned SD_of_sample_means to an integer from 1, 2, 3, 4."
success_message: "✅ You've assigned SD_of_sample_means to an integer from 1, 2, 3, 4."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(SD_of_sample_means) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
True
```
failure_message: "❌ You have not assigned the correct values to SD_of_sample_means."
success_message: "✅ You've assigned the correct values to SD_of_sample_means."



Tests for task_12 (points: None)
```
>>> isinstance(pop_vs_sample, np.ndarray)
True
```
failure_message: "❌ pop_vs_sample should be assigned to an array."
success_message: "✅ pop_vs_sample is assigned to an array"

```
>>> isinstance(pop_vs_sample.item(0), (int, np.integer))
True
```
failure_message: "❌ pop_vs_sample should be assigned to an array of integers. 
There is at least one correct statement."
success_message: "✅ pop_vs_sample is assigned to an array of integers"

```
>>> np.all(pop_vs_sample <= 6) and np.all(pop_vs_sample >= 1)
np.True_
```
failure_message: "❌ pop_vs_sample should be assigned to an array of integers between 1 and 6, inclusive."
success_message: "✅ pop_vs_sample is assigned to an array of integers between 1 and 6, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(set(np.array(pop_vs_sample).astype(int))) == '1598dfd5bc2078fe85358ea6ac2d5d08'
True
```
failure_message: "❌ You have not assigned the correct values to pop_vs_sample."
success_message: "✅ You've assigned the correct values to pop_vs_sample."



Tests for task_13 (points: None)
```
>>> isinstance(task13, (int, np.integer))
True
```
failure_message: "❌ You have not assigned task13 to an integer."
success_message: "✅ You've assigned task13 to an integer."

```
>>> task13 in {1, 2}
True
```
failure_message: "❌ You have not assigned task13 to the integer 1 or 2."
success_message: "✅ You've assigned task13 to the integer 1 or 2."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(task13) == 'c4ca4238a0b923820dcc509a6f75849b'
True
```
failure_message: "❌ You have not assigned the correct values to task13."
success_message: "✅ You've assigned the correct values to task13."

