<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 12: Regression Inference

## References
* [Inference for Regression](https://ccsf-math-108.github.io/textbook/chapters/16/Inference_for_Regression.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp26/resources/code-reference.html)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw12.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Regression Inference

---

Previously, you've used confidence intervals to quantify uncertainty about estimates. You can also run hypothesis tests using a confidence interval under the following procedure:

1. Define a null and alternative hypothesis (they must be of the form "The parameter is X" and "The parameter is not X").
2. Choose a p-value cutoff, and call it $q$.
3. Construct a $(100-q)\%$ interval using bootstrap sampling (for example, if your p-value cutoff is 0.01, or 1%, then construct a 99% confidence interval).
4. Using the confidence interval, determine if your data are more consistent with your null or alternative hypothesis:
   * If the null hypothesis parameter X is in your confidence interval, the data are more consistent with the null hypothesis.
   * If the null hypothesis parameter X is *not* in your confidence interval, the data are more consistent with the alternative hypothesis.

More recently, we've discussed the use of linear regression to make predictions based on correlated variables. For example, we can predict the height of children based on the heights of their parents.

We can combine these two topics to make powerful statements about our population by using the following techniques:

- Bootstrapped interval for the true slope
- Bootstrapped prediction interval for y (given a particular value of x)

This part of the homework explores constructing an interval for the true slope.

---

### The Data

The [Snowy Plover](https://www.audubon.org/field-guide/bird/snowy-plover) is a tiny bird that lives on the coast in parts of California and elsewhere. It is so small that it is vulnerable to many predators, including people and dogs that don't look where they are stepping when they go to the beach. It is considered endangered in many parts of the U.S.

The data are about the eggs and newly-hatched chicks of the Snowy Plover. Here's a picture of [a parent bird incubating its eggs](http://cescos.fau.edu/jay/eps/articles/snowyplover.html).

<img src="plover_and_eggs.jpeg" alt="Plover and Eggs">

The data were collected at the Point Reyes National Seashore by a former [student at Berkeley](https://openlibrary.org/books/OL2038693M/BLSS_the_Berkeley_interactive_statistical_system). The goal was to see how the size of an egg could be used to predict the weight of the resulting chick. The bigger the newly-hatched chick, the more likely it is to survive.

<img src="plover_and_chick.jpeg" alt="Plover and Chick">

Each row of the table below corresponds to one Snowy Plover egg and the resulting chick. Note how tiny the bird is:

- `Egg Length` and `Egg Breadth` (widest diameter) are measured in millimeters
- `Egg Weight` and `Bird Weight` are measured in grams; for comparison, a standard paper clip weighs about one gram

```python
birds = Table.read_table('snowy_plover.csv')
birds
```

In this investigation, we will be using the egg weight to predict bird weight. Run the cell below to create a scatter plot of the egg weights and bird weights, along with their line of best fit.

```python
# Just run this cell and examine the scatter plot.
birds.scatter('Egg Weight', "Bird Weight", fit_line=True)
```

---

### Residual Plot

It seems like a linear model might be appropriate to apply to this situation, but you should check the residual plot before moving forward. You'll likely find the following functions, `standard_units` and `correlation`, useful.

```python
def standard_units(data):
    return (data - np.mean(data)) / np.std(data)

def correlation(x, y):
    return np.mean(standard_units(x) * standard_units(y))
```

---

#### Task 01 📍

Create a function called `fit_line`.  It should take a table (e.g. `birds`) and the column names associated with *x* and *y* as its arguments and return an *array* containing the slope and intercept of the regression line (in that order) that predicts the *y* column in the table using the *x* column.

_Points:_ 3

```python
def fit_line(tbl, x_col, y_col):
    x_vals = tbl.column(x_col)
    y_vals = tbl.column(y_col)
    r = ...
    slope = ...
    intercept = ...
    return make_array(slope, intercept)

fit_line(birds, "Egg Weight", "Bird Weight")
```

```python
grader.check("task_01")
```

---

Now, run the following cell to plot the line produced by calling `fit_line` on the `birds` table.

**Note:** You are not responsible for the code in the cell below, but make sure that your `fit_line` function generates a reasonable line for the data. Compare the line with the line produced in the scatter plot above with the `fit_line=True` parameter.

```python
# Just run this cell
snowy_slope, snowy_intercept = fit_line(birds, "Egg Weight", "Bird Weight")
birds.scatter("Egg Weight", "Bird Weight")
plt.plot([min(birds.column("Egg Weight")), 
          max(birds.column("Egg Weight"))], 
         [snowy_slope*min(birds.column("Egg Weight"))+snowy_intercept, 
          snowy_slope*max(birds.column("Egg Weight"))+snowy_intercept])
plt.show()
```

---

#### Task 02 📍

With a linear regression model, update the `birds` table with a column of predicted birth weights using the regression line and the associated residuals.

1. Create an array called `predicted_bird_weights` that contains the predicted bird weights from the regression line.
2. Create an array called `residuals` that contains the residuals associated with those predictions.
3. Update the `birds` table to include columns `'Predicted Bird Weight'` and `'Residual'` with their respective data.

**Hint:** Consider using the variables `snowy_slope`, `snowy_intercept` we defined above.

_Points:_ 3

```python
predicted_bird_weights = ...
residuals = ...
birds = birds.with_columns(
    'Predicted Bird Weight', ...,
    'Residual', ...)
```

```python
grader.check("task_02")
```

---

#### Task 03 📍🔎

<!-- BEGIN QUESTION -->

Now that you have the residuals for these predictions:
1. Create a residual plot showing egg weights on the horizontal axis and the residuals on the vertical axis.
2. In the space provided, provide a short justification as to why it is reasonable to predict bird weight from egg weight using the linear model. Make sure to point out features from the residual plot relevant to [Section 15.5](https://ccsf-math-108.github.io/textbook/chapters/15/5/Visual_Diagnostics.html#regression-diagnostics) that back up your justification. Additionally, it is reasonable to express concerns about applying a model to a given situation, so you are welcome to point out any concerns in relation to the content from Section 15.5.

_Points:_ 2

_Type your answer here, replacing this text._

```python
# Generate the residual plot here
birds.scatter('Egg Weight', 'Residual')
```

<!-- END QUESTION -->

---

### Finding the Bootstrap Confidence Interval for the True Slope

```python
# Just run this cell
birds.scatter('Egg Weight', "Bird Weight", fit_line=True)
```

Looking at the scatter plot of our sample, we observe a linear relationship between egg weight and bird weight. However, relationships that appear in a sample might not exist in the population from which the sample was taken.

We want to know if there truly exists a linear relationship between egg weight and bird weight for Snowy Plovers. If there is no linear relationship between the two variables, then we'd expect a correlation of 0. Consequently, the slope of the regression line would also be 0. 

We highly recommend reviewing section [16.2](https://ccsf-math-108.github.io/textbook/chapters/16/2/Inference_for_the_True_Slope.html) before attempting this part of the homework

---

#### Task 04 📍

Let's run a hypothesis test using confidence intervals to see if there is a linear relationship between egg weight and bird weight. Define the null and alternative hypotheses that will allow you to conduct this test.

- **Null Hypothesis: The true slope of the regression line relating egg weight and bird weight for Snowy Plovers is ____.**
- **Alternative Hypothesis: The true slope of the regression line relating egg weight and bird weight for Snowy Plovers is not ____.**

Assign `snowy_slope_null_value` to the numerical value that completes both of the above statements to form an appropriate pair of hypotheses for this test.

_Points:_ 2

```python
snowy_slope_null_value = ...
```

```python
grader.check("task_04")
```

---

Next, you will create a confidence interval that quantifies our uncertainty about the true relationship between egg weight and bird weight.

---

#### Task 05 📍

Create an array called `resampled_slopes` that contains the slope of the best fit line for 1,000 bootstrap resamples of `birds`. Plot the distribution of these slopes.

_Points:_ 2

```python
resampled_slopes = ...

for i in np.arange(1_000): 
    birds_bootstrap = ...
    bootstrap_line = ...
    bootstrap_slope = ...
    resampled_slopes = ...
    
# DO NOT CHANGE THIS LINE
Table().with_column("Slope estimate", resampled_slopes).hist()
```

```python
grader.check("task_05")
```

---

#### Task 06 📍

Use your resampled slopes to construct an 95% confidence interval for the true value of the slope.

_Points:_ 3

```python
lower_end = ...
upper_end = ...
print("95% confidence interval for slope: [{:g}, {:g}]".format(lower_end, upper_end))
```

```python
grader.check("task_06")
```

---

#### Task 07 📍🔎

<!-- BEGIN QUESTION -->

In your response, make to address the following:
1. Based on your confidence interval, would you **reject** or **fail to reject** the null hypothesis that the true slope is 0? 
2. Provide an explanation Why?  
3. What p-value cutoff are you using?

**Hint:** Re-read the [Regression Inference section](#Regression-Inference) of this homework.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Finding the Bootstrap Prediction Interval

Suppose we're visiting Point Reyes and stumble upon some Snowy Plover eggs; we'd like to know how heavy they'll be once they hatch. In other words, we want to use our regression line to make predictions about a bird's weight based on the weight of the corresponding egg.

However, just as we're uncertain about the slope of the true regression line, we're also uncertain about the predictions made based on the true regression line.

---

#### Task 08 📍

Define the function `fitted_value`.  It should take in four arguments:

1. `table`: a table like `birds`.  We'll be predicting the values in the second column using the first.
2. `x_col`: the name of our x-column within the input `table`
3. `y_col`: the name of our y-column within the input `table`
2. `given_x`: a number, the value of the predictor variable for which we'd like to make a prediction.
  
The function should return the line's prediction for the given x.

*Hint:* Make sure to use the `fit_line` function you defined in [Task 1](#Task-01-📍).

_Points:_ 2

```python
def fitted_value(table, x_col, y_col, given_x):
    line = ...
    slope = ...
    intercept = ...
    ...

# Here's an example of how fitted_value is used. The code below
# computes the prediction for the bird weight, in grams, based on
# an egg weight of 8 grams.
egg_weight_eight = fitted_value(birds, "Egg Weight", "Bird Weight", 8)
egg_weight_eight
```

```python
grader.check("task_08")
```

---

#### Task 09 📍

Raymond, the resident Snowy Plover expert at Point Reyes, tells us that the egg he has been carefully observing has a weight of 9 grams. Using `fitted_value` above, assign the variable `experts_egg` to the predicted bird weight for Raymond's egg.

_Points:_ 2

```python
experts_egg = ...
experts_egg
```

```python
grader.check("task_09")
```

---

Let's look at the number of rows in the birds table.

```python
birds.num_rows
```

A fellow parkgoer raises the following objection to your prediction:

> "Your prediction depends on your sample of 44 birds.  Wouldn't your prediction change if you had a different sample of 44 birds?"

Having read section [16.3](https://inferentialthinking.com/chapters/16/3/Prediction_Intervals.html) of the textbook, you know just the response! Had the sample been different, the regression line would have been different too. This would ultimately result in a different prediction. To see how good our prediction is, we must get a sense of how variable the prediction can be.

---

#### Task 10 📍

Define a function `compute_resampled_line` that takes in a table `tbl`and two column names, `x_col` and `y_col`, and returns an array containing the parameters of the best fit line (slope and intercept) for one bootstrapped resample of the table.

_Points:_ 2

```python
def compute_resampled_line(tbl, x_col, y_col):
    resample = ...
    resampled_line = ...
    ...
```

```python
grader.check("task_10")
```

---

**Run** the following cell below in order to define the function `bootstrap_lines`.  It takes in four arguments:
1. `tbl`: a table like `birds`
2. `x_col`: the name of our x-column within the input `tbl`
3. `y_col`: the name of our y-column within the input `tbl`
4. `num_bootstraps`: an integer, a number of bootstraps to run.

It returns a *table* with one row for each bootstrap resample and the following two columns:
1. `Slope`: the bootstrapped slopes 
2. `Intercept`: the corresponding bootstrapped intercepts 

```python
# Just run this cell
def bootstrap_lines(tbl, x_col, y_col, num_bootstraps):
    resampled_slopes = make_array()
    resampled_intercepts = make_array() 
    for i in np.arange(num_bootstraps): 
        resampled_line = compute_resampled_line(tbl, x_col, y_col) 
        resampled_slope = resampled_line.item(0) 
        resampled_intercept = resampled_line.item(1) 
        resampled_slopes = np.append(resampled_slopes,resampled_slope)
        resampled_intercepts = np.append(resampled_intercepts,resampled_intercept)
    tbl_lines = Table().with_columns('Slope', resampled_slopes, 'Intercept', resampled_intercepts)
    return tbl_lines

regression_lines = bootstrap_lines(birds, "Egg Weight", "Bird Weight", 1000)
regression_lines
```

---

#### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Create an array called `predictions_for_eight` that contains the predicted bird weights based on an egg of weight 8 grams for each regression line in `regression_lines`.

_Points:_ 2

```python
predictions_for_eight = ...

# This will make a histogram of your predictions:
table_of_predictions = Table().with_column('Predictions at Egg Weight=8', predictions_for_eight)
table_of_predictions.hist('Predictions at Egg Weight=8', bins=20)
```

<!-- END QUESTION -->

---

#### Task 12 📍

Create an approximate 95% confidence interval for the predictions in `predictions_for_eight`.

_Points:_ 3

```python
lower_bound = ...
upper_bound = ...

print('95% Confidence interval for predictions for x=8: (', lower_bound,",", upper_bound, ')')
```

```python
grader.check("task_12")
```

---

#### Task 13 📍

Set `plover_statements` to an array of integer(s) that correspond to statement(s) that are true.

1. The 95% confidence interval covers 95% of the bird weights for eggs that had a weight of eight grams in the population of snowy plovers.

2. The 95% confidence interval quantifies the uncertainty in our estimate of what the true line would predict.

_Points:_ 2

```python
plover_statements = ...
```

```python
grader.check("task_13")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 3)
```
>>> type(fit_line(Table().with_columns('x', make_array(0, 1), 'y', make_array(1, 3)), 'x', 'y')) == np.ndarray
True
```
failure_message: ❌ Your fit_line function should produce an array.
success_message: ✅ Your fit_line function produces an array.
```
>>> len(fit_line(Table().with_columns('x', make_array(0, 1), 'y', make_array(1, 3)), 'x', 'y')) == 2
True
```
failure_message: ❌ The array from your fit_line function should contain 2 items.
success_message: ✅ The array from your fit_line function contains 2 items.
```
>>> np.allclose(np.round(fit_line(Table().with_columns('x', make_array(0, 1), 'y', make_array(1, 3)), 'x', 'y'), 5), np.array([2, 1]))
True
```
failure_message: ❌ Your fit_line function did not produce the correct two values for our test data.
success_message: ✅ Your fit_line function produced the correct two values for our test data.


Tests for task_02 (points: 3)
```
>>> set(birds.labels) == {'Egg Length', 'Egg Breadth', 'Egg Weight', 'Bird Weight', 'Predicted Bird Weight', 'Residual'}
True
```
failure_message: ❌ The birds table does not have the correct labels.
success_message: ✅ The birds table has the correct labels.
```
>>> np.isclose(np.sum(predicted_bird_weights), 270.4000000000001)
np.True_
```
failure_message: ❌ The sum of your predicted bird weights doesn't seem correct.
success_message: ✅ The sum of your predicted bird weights seems correct.
```
>>> np.isclose(np.average(residuals), 0)
np.True_
```
failure_message: ❌ The average residual calculated from your array residuals should be (approximately) 0.
success_message: ✅ The average residual calculated from your array residuals is (approximately) 0.
```
>>> np.isclose(correlation(birds.column('Egg Weight'), birds.column('Residual')), 0)
np.True_
```
failure_message: ❌ The correlation between the egg weights and residuals calculated from your birds table should be (approximately) 0.
success_message: ✅ The correlation between the egg weights and residuals calculated from your birds table is (approximately) 0.


Tests for task_04 (points: 2)
```
>>> isinstance(snowy_slope_null_value, (int, np.float32, np.float64))
True
```
failure_message: ❌ snowy_slope_null_value should be assigned to a number.
success_message: ✅ snowy_slope_null_value is assigned to a number.


Tests for task_05 (points: 2)
```
>>> isinstance(resampled_slopes, np.ndarray)
True
```
failure_message: ❌ resampled_slopes should be an array.
success_message: ✅ resampled_slopes is an array.
```
>>> isinstance(resampled_slopes[0], (int, np.float32, np.float64)) and len(resampled_slopes) == 1000
True
```
failure_message: ❌ resampled_slopes should contain 1,000 numerical values.
success_message: ✅ resampled_slopes contains 1,000 numerical values.


Tests for task_06 (points: 3)
```
>>> lower_end > 0
np.True_
```
failure_message: ❌ The lower_end value should be greater than 0.
success_message: ✅ The lower_end value is greater than 0.
```
>>> upper_end < 4
np.True_
```
failure_message: ❌ The upper_end value should be less than 4.
success_message: ✅ The upper_end value is less than 4.


Tests for task_08 (points: 2)
```
>>> type(egg_weight_eight) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ 
success_message: ✅ 


Tests for task_09 (points: 2)
```
>>> type(experts_egg) in set([float, np.float32, np.float64])
True
```
failure_message: ❌ 
success_message: ✅ 


Tests for task_10 (points: 2)
```
>>> results = compute_resampled_line(Table().with_columns('x', make_array(0, 1), 'y', make_array(1, 3)), 'x', 'y')
>>> isinstance(results, np.ndarray) and len(results) == 2
True
```
failure_message: ❌ Your function should return an array with 2 items.
success_message: ✅ Your function returns an array with 2 items.


Tests for task_12 (points: 3)
```
>>> lower_bound > 5
np.True_
```
failure_message: ❌ The lower_bound value should be greater than 5.
success_message: ✅ The lower_bound value is greater than 5.
```
>>> upper_bound < 6.5
np.True_
```
failure_message: ❌ The upper_bound value should be less than 6.5.
success_message: ✅ The upper_bound value is less than 6.5.


Tests for task_13 (points: 2)
```
>>> type(plover_statements) == np.ndarray
True
```
failure_message: ❌ plover_statements should be an array.
success_message: ✅ plover_statements is an array.
