<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 03: Tables

---

## Textbook References

* [Sequences](https://ccsf-math-108.github.io/textbook/chapters/05/Sequences.html)
* [Tables](https://ccsf-math-108.github.io/textbook/chapters/06/Tables.html)

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the follwing code cell to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")

# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## Fruit

To begin this homework, you are going to practice with some of the `Table` methods you learned about this week.

### Task 01 📍

Suppose you have a small fruit stand that has an inventory of 4 apples, 3 oranges, and 1 pineapple.

Create a table called `fruits` that contains this information:
* It should have two columns: `'fruit name'` and `'amount'`.
* It should have three rows, one for each type of fruit.
* The values for `'fruit name'` should be the strings `'apple'`, `'orange'`, and `'pineapple'`
* The values for `'amount'` should be integers `4`, `3`, and `1`.

*Note:* Use lower-case and singular words for the name of each fruit, like `"apple"`.

_Points:_ 2

```python
fruits = ...
fruits
```

```python
grader.check("task_01")
```

---

### Task 02 📍

The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. 

**Note**: `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.

_Points:_ 2

```python
inventory = ...
inventory
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit from the other boxes or to `False` if multiple boxes contain the same fruit.

*Hint:* You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.

_Points:_ 2

```python
all_different = ...
all_different
```

```python
grader.check("task_03")
```

---

### Task 04 📍

The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called `price per fruit ($)` that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.

_Points:_ 2

```python
sales = ...
sales
```

```python
grader.check("task_04")
```

---

### Task 05 📍

How many fruits did the store sell in total on that day (last Saturday)?

_Points:_ 2

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("task_05")
```

---

### Task 06 📍

What was the store's total revenue (the total value of all fruits sold) on that day? Assign that value to `total_revenue`.

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales. It should be `355 * 0.06`, which gives `21.3`.

_Points:_ 2

```python
total_revenue = ...
total_revenue
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's **original** count, so that the `count` column is **updated to be** the amount of fruit remaining after Saturday.

**Notes**:
* As a check, since `355` grapes were sold from the original `517`, then `162` grapes should be in the remaining inventory once you're done.
* We recommend that you do not do the subtractions individually. Let Python (with the help of arrays) do the hard work for you!
* Don't forget that you can always add extra lines of code and intermediate variables. The auto-grader is only checking the value of `remaining_inventory`.

_Points:_ 3

```python
remaining_inventory = ...
remaining_inventory
```

```python
grader.check("task_07")
```

---

## Unemployment


The [Federal Reserve Bank of St. Louis](https://fred.stlouisfed.org/categories/33509) publishes data about jobs in the US. We've imported data from this site on unemployment rates in the United States. There are many ways of defining unemployment rates and our dataset includes two.
1. Non-Employment Index (NEI): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job. 
2. NEI+PTER (Non-Employment Index and 'Part-Time for Economic Reasons'): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job or are only working at a part-time job.
The data that you will use in this assignment contains [quarterly average NEI percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIM156SFRBRIC) and [quarterly average NEI+PTER percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIPTERM156SFRBRIC)

*In this assignment, you will do some preliminary analysis of these unemployment metrics. In a future assignment, you will work on creating visualizations and summarizing your analysis.*


---

### Task 08 📍

The data are in a CSV file called `unemployment.csv`.  Load the data in that file into a `Table` and assign the table to the name `unemployment`.


_Points:_ 1

```python
unemployment = ...
unemployment
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Dates can be tricky to work with. When you loaded the data into the table `unemployment`, the dates were read as strings. Run the following cell to see this is true.

```python
type(unemployment.column('DATE').item(0))
```

Strings do not always make the best sense to use for dates since we organize dates by months, years, etc. and this is different than the alphabetical order we typically use for strings. A common date type for dates is [`datetime`](https://docs.python.org/3/library/datetime.html). In this class, you are not responsible for knowing how to work directly with the `datetime` format. Instead, we will set up the code to help you with that date type.

**Run the following code cell to convert all the dates in the `'DATE'` column of `unemployment` to `datetime` format.**

_Points:_ 0

```python
# Just run this cell
import datetime
from datetime import datetime, date

def update_date_format(date_string):
    date_object = datetime.strptime(date_string, '%Y-%m-%d').date()
    return date_object

if not isinstance(unemployment.column('DATE').item(0),  date):
    unemployment = unemployment.with_column(
        'DATE', unemployment.apply(update_date_format, 'DATE'))
    print("The date type for the values in the 'DATE' column have been updated.\n")
else:
    print("The date type is correct for the values in the 'DATE' column.\n")
    
display(unemployment)
```

```python
grader.check("task_09")
```

---

### Task 10 📍

For this task:

1. Sort the data in unemployment by NEI in descending order naming the sorted table by_nei.
2. Create another table called by_nei_pter that's sorted by NEI+PTER in descending order instead.

_Points:_ 2

```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("task_10")
```

---

### Task 11 📍

For this task:
1. Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.
2. Assign that table to the name `greatest_nei`. Keep in mind that `greatest_nei` should be sorted in descending order of `NEI`.

**Note**: Each row of `unemployment` represents a quarter.


_Points:_ 2

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("task_11")
```

---

### Task 12 📍

It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI+PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  

Using the original `unemployment` table, create an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

_Points:_ 2

```python
pter = ...
pter
```

```python
grader.check("task_12")
```

---

### Task 13 📍

For this task,
1. Add `pter` as a column to `unemployment` (named `"PTER"`).
2. Make sure the data is sorted by the `'DATE'` values in ascending order (`1994-01-01` should be the first date).
3. Assign this table to the name `unemployment` to update the original table.

**Note**: You do not need to do this all in one line of code.

_Points:_ 2

```python
unemployment = ...
unemployment
```

```python
grader.check("task_13")
```

---

## NHANES Demographic Data

The National Health and Nutrition Examination Survey (NHANES) collects data about the health of adults and children in the United States. The CDC's [National Center for Health Statistics (NCHS)](https://www.cdc.gov/nchs/nhanes/index.html) conducts NHANES. It is a national health survey that includes health exams, laboratory tests, and dietary interviews for participants of all ages. In addition to health measures, the survey also collects demographic data.

We will only consider the following demographic variables for each individual for the remaining tasks:

- `ID`: A unique identifier for each individual in the data set
- `Gender`: `male` or `female` (no other gender identifications are used)
- `AGE`: age of the participant
- `EDUC_LEV`: education level (`less than 9th grade`, `9 to 11th grade`, `HS graduate`, `some college`, `college graduate`)
- `MAR_STATUS`: marital status (`married or lives with partner`, `widowed or divorced or separated`, `never married`)
- `HH_SIZE`: household size
- `RATIO`: the ratio of the participants income to the national poverty level

Run the following code cell to import a data set from the NHAES that has 7738 adult respondents. 

```python
nhaes_adult = Table().read_table('nhanes_adult.csv')
nhaes_adult
```

### Task 14 📍

Look carefully at the table above, and assign `num_attributes` to an array of the number(s) of each variable below that are numerical attributes.

1. `GENDER`
2. `AGE`
3. `EDUC_LEV`
4. `MAR_STATUS`
5. `HH_SIZE`
6. `RATIO`

_Points:_ 4

```python
num_attributes = ...
```

```python
grader.check("task_14")
```

---

### Task 15 📍

Which visualization will we use if we want to analyze the relationship between visualize the individual's age and their ratio?

Assign `task_15_viz` to the number of the correct choice below:

1. Bar chart
2. Histogram
3. Line Plot
4. Scatterplot

_Points:_ 2

```python
task_15_viz = ...
```

```python
grader.check("task_15")
```

---

### Task 16 📍

Which visualization will we use if we want to analyze the distribution of household size?

Assign `task_16_viz` to the number of the correct choice below:

1. Bar chart
2. Histogram
3. Line Plot
4. Scatterplot

_Points:_ 2

```python
task_16_viz = ...
```

```python
grader.check("task_16")
```

---

### Task 17📍

Which visualization will we use if we want to analyze the distribution of marital status?

Assign `task_17_viz` to the number of the correct choice below:

1. Bar chart
2. Histogram
3. Line Plot
4. Scatterplot

_Points:_ 2

```python
task_17_viz = ...
```

```python
grader.check("task_17")
```

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

1. **Review the Rubric:** View the rubric on the Canvas assignment page to understand the scoring criteria.  
2. **Review Manually Graded Tasks:** Verify that you have responded to all manually graded tasks marked with 📍🔎.  
3. **Review Auto-Graded Tasks:** Inspect the output from the `grade.check_all()` command above. Ensure you have passed all available tests for the auto-graded tasks in this notebook marked with 📍.  
   - This command executes all auto-grader tests sequentially.  
   - Note that Homework and Project assignments include hidden auto-grader tests that are not visible to you. Passing all visible tests does **not** guarantee that your solution is fully correct.  
4. **Save Your Work:** In the notebook toolbar, go to `File → Save Notebook` to save your work and create a checkpoint.  
5. **Download the Notebook:** In the notebook toolbar, go to `File → Download IPYNB` to download the notebook (`.ipynb`) file.  
6. **Upload to Canvas:** Upload the downloaded `.ipynb` file to the appropriate Canvas assignment page.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> isinstance(fruits, Table)
True
```
failure_message: ❌ fruits should be a Table. Try using Table().with_columns(...) to build the table.
success_message: ✅ fruits is a Table.
```
>>> fruits.sort(0)
fruit name | amount
apple      | 4
orange     | 3
pineapple  | 1
```
failure_message: ❌ fruits doesn't look correct. Double-check the capitalization, spelling, and counts.
success_message: ✅ fruits looks correct!


Tests for task_02 (points: 2)
```
>>> isinstance(inventory, Table)
True
```
failure_message: ❌ inventory should be a Table.
success_message: ✅ inventory is a Table.
```
>>> inventory.sort(0)
box ID | fruit name | count
25274  | apple      | 20
26187  | strawberry | 255
43566  | peach      | 40
48800  | orange     | 35
52357  | strawberry | 102
53686  | kiwi       | 45
57181  | strawberry | 123
57930  | grape      | 517
```
failure_message: ❌ inventory doesn't look correct.
success_message: ✅ inventory looks correct.


Tests for task_03 (points: 2)
```
>>> isinstance(all_different, bool)
True
```
failure_message: ❌ all_different should be assigned to True or False. Double check your spelling/capitalization of True or False.
success_message: ✅ all_different is assigned to True or False.


Tests for task_04 (points: 2)
```
>>> isinstance(sales, Table)
True
```
failure_message: ❌ sales should be a Table.
success_message: ✅ sales is a Table.
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```
failure_message: ❌ sales doesn't look correct.
success_message: ✅ sales looks correct.


Tests for task_05 (points: 2)
```
>>> isinstance(total_fruits_sold, (int, np.integer))
True
```
failure_message: ❌ total_fruits_sold should be an integer.
success_message: ✅ total_fruits_sold is an integer.
```
>>> total_fruits_sold > 10
np.True_
```
failure_message: ❌ total_fruits_sold is not a possible quantity.
success_message: ✅ total_fruits_sold is a possible quantity.


Tests for task_06 (points: 2)
```
>>> 50 <= total_revenue <= 150
np.True_
```
failure_message: ❌ total_revenue is not a possible value. It should be a number between 50 and 150.
success_message: ✅ total_revenue is a possible value.
```
>>> isinstance(total_revenue, (float, np.float64))
True
```
failure_message: ❌ total_revenue should be a float.
success_message: ✅ total_revenue is a float.


Tests for task_07 (points: 3)
```
>>> isinstance(remaining_inventory, Table)
True
```
failure_message: ❌ remaining_inventory should be a Table.
success_message: ✅ remaining_inventory is a Table.
```
>>> remaining_inventory.where(1, 'grape')
box ID | fruit name | count
57930  | grape      | 162
```
failure_message: ❌ The updated grape inventory doesn't look correct.
success_message: ✅ The updated grape inventory looks good.


Tests for task_08 (points: 1)
```
>>> isinstance(unemployment, Table)
True
```
failure_message: ❌ It doesn't seem that you have assigned the name unemployment to a table. Try using the command Table().read_table(...)
success_message: ✅ It seems that you have assigned the name unemployment to a table.
```
>>> np.isclose(np.sum(unemployment.column('NEI')), 1044.46519909577)
np.True_
```
failure_message: ❌ The contents in unemployment don't seem correct. Double-check that you are using the unemployment.csv
success_message: ✅ The contents in unemployment seem correct.


Tests for task_09 (points: 0)
```
>>> isinstance(unemployment.column('DATE').item(0), date)
True
```
failure_message: ❌ It doesn't seem like you've run the cell to convert the date data types. Make sure to run the above code cell.
success_message: ✅ You've converted all the dates.


Tests for task_10 (points: 2)
```
>>> by_nei.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei.
success_message: ✅ It seems like you have the correct values in by_nei.
```
>>> by_nei_pter.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei_pter.
success_message: ✅ It seems like you have the correct values in by_nei_pter.


Tests for task_11 (points: 2)
```
>>> greatest_nei.num_rows == 10
True
```
failure_message: ❌ There are not 10 rows (quarters) in your table.
success_message: ✅ There are 10 rows (quarters) in your table.


Tests for task_12 (points: 2)
```
>>> isinstance(pter, np.ndarray)
True
```
failure_message: ❌ pter doesn't seem to be an array.
success_message: ✅ pter is an array.
```
>>> round(pter.item(6), 4) != -1.1446
True
```
failure_message: ❌ Be careful of subtracting the arrays in the wrong order.
success_message: ✅ You didn't seem to subtract the correct arrays in the wrong order.


Tests for task_13 (points: 2)
```
>>> unemployment.num_columns == 4
True
```
failure_message: ❌ Your table doesn't have the correct number of columns.
success_message: ✅ Your table has the correct number of columns.
```
>>> np.all(unemployment.sort('DATE').column('DATE') == unemployment.column('DATE'))
np.True_
```
failure_message: ❌ Your table doesn't appear to be sorted, in ascending order, by the date values.
success_message: ✅ Your table appears to be sorted, in ascending order, by the date values.


Tests for task_14 (points: 4)
```
>>> isinstance(num_attributes, np.ndarray)
True
```
failure_message: ❌ num_attributes should be an array.
success_message: ✅ num_attributes is an array.
```
>>> isinstance(num_attributes.item(0), int)
True
```
failure_message: ❌ num_attributes should contain at least one integer.
success_message: ✅ num_attributes contains at least one integer.


Tests for task_15 (points: 2)
```
>>> task_15_viz in [1, 2, 3, 4]
True
```
failure_message: ❌ 1, 2, 3, or 4 should be assigned to task_15_viz.
success_message: ✅ This is a possible value for task_15_viz.


Tests for task_16 (points: 2)
```
>>> task_16_viz in [1, 2, 3, 4]
True
```
failure_message: ❌ 1, 2, 3, or 4 should be assigned to task_16_viz.
success_message: ✅ This is a possible value for task_16_viz.


Tests for task_17 (points: 2)
```
>>> task_17_viz in [1, 2, 3, 4]
True
```
failure_message: ❌ 1, 2, 3, or 4 should be assigned to task_17_viz.
success_message: ✅ This is a possible value for task_17_viz.
