<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 5: Simulation and Chance

## Textbook References

* [Sections 8.4 and 8.5](https://ccsf-math-108.github.io/textbook/chapters/08/4/Joining_Tables_by_Columns.html)
* [Chapter 9](https://ccsf-math-108.github.io/textbook/chapters/09/Randomness.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cells to set up the notebook.

_You are not responsible for understanding what this code does._

_Points:_ 0

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

```python
# Configure this Notebook
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
```

<!-- END QUESTION -->

---

## 🏀 2024-25 CCSF Men's Basketball Season 

The CCSF men's basketball team [won the 32CA State Championship](https://ccsfathletics.com/sports/mbkb/2024-25/releases/20250316ryn5ow) in the 2024-25 season - the sixth state title in school history. For the rest of the tasks we will analyze some performance measures for their season. A men's college basketball game is divided into two periods, called halves. The number of points CCSF scored in each half (e.g. `CCSF 1H`), and the number of points their opponent scored in each half  (e.g. `Opp 1H`) are stored in a table called `game_summaries`. If the game is tied at the end of the 2nd half, the game goes into an additional period called overtime.

**Notes**:
* The 2024-25 season data was collected using the Game Log from [the CCSF Athletics website](https://ccsfathletics.com/sports/mbkb/2024-25/teams/sanfrancisco?view=gamelog).
* CCSF played some opponents more than once.
* Games 1 - 28 were regular season games, and games 29 - 33 were playoff games. The championship game is game 33.
                                                                                             

```python
game_summaries = Table().read_table("ccsf_bb_24_25_summary.csv")
game_summaries.show()
```

Let's start by finding the total points each team scored in a game.

---

### Task 01 📍

Write a function called `sum_scores`.  It should take three arguments, where each argument is the team's half or overtime score. It should return the team's total score for that game.

_Points:_ 2

```python
def sum_scores(...):
    """
    Returns the total score calculated by adding up the score of each half and overtime.
    
    Examples
    --------
    >>> sum_scores(45, 38, 12)
    95
    >>> sum_scores(50, 47, 0)
    97
    """
    ...

sum_scores(25, 32, 9) #DO NOT CHANGE THIS LINE
```

```python
grader.check("task_01")
```

---

### Task 02 📍

Create a new table `final_scores` with four columns in this *specific* order: `Game`, `Opponent`, `CCSF Score`, `Opponent Score`. You will have to create the `CCSF Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem.

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating final_scores, check that `ccsf_scores` and `opponent_scores` output what you want. Also, you will more than likely want to create more steps/intermediate variables.



_Points:_ 3

```python
final_scores = ...
final_scores
```

```python
grader.check("task_02")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `(n+1)`st row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. 

Run the following cells to an example:

```python
game_summaries.row(9)
```

```python
game_summaries.row(9).item("CCSF 2H")
```

---

### Task 03 📍

We want to see for a particular game whether or not CCSF won. Write a function called `did_ccsf_win`.  It should take one argument: a **row object** from the `final_scores` table. It should return either `True` if CCSF's score was greater than the Opponent's score, and `False` otherwise.

*Note 1*: "Row object" means a row from the table extracted (behind the scenes) using `tbl.row(index)` that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!




_Points:_ 3

```python
def ...(row):
    '''
    Return a boolean indicating whether CCSF won the game.
    
    Examples
    --------
    >>> did_ccsf_win(game_summaries.row(6))
    True
    >>> did_ccsf_win(game_summaries.row(15))
    False
    '''
    ...
```

```python
grader.check("task_03")
```

---

### Task 04 📍

As State (and Conference) Champtions, CCSF won a lot of games during the 2024-25 season. Using the `final_scores` table, assign `results` to an array of `True` and `False` values that correspond to whether or not CCSF won. Add the `results` array to the `final_scores` table, and assign this to `final_scores_with_results`. Then, respectively assign the number of wins and losses CCSF had to `ccsf_wins` and `ccsf_losses`.

*Hint*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`



_Points:_ 2

```python
results = ...
final_scores_with_results = ...
ccsf_wins = ...
ccsf_losses = ...

# Don't delete or edit the following line:
print(f"In the 2024-25 Season, CCSF Men's Basketball won {ccsf_wins} games and lost {ccsf_losses} game(s).")
```

```python
grader.check("task_04")
```

---

### Task 05 📍

Sometimes in basketball, the two teams are equally matched and the game is quite close. Other times, it is a blowout, where the winning team wins by a large margin of victory. Let's define a **big win** to be a game in which the winning team won by more than 15 points. 

Create a function called `is_big_win`.
* The function should accept a single row (`datascience.tables.Row` data type) from a table like the `final_scores` table.
* The function should output `True` if CCSF won by **more than** 15 points. Otherwise, it should return `False`.

Test your function on the sixth row of the table. `final_scores.row(5)` should be the row object `Row(Opponent='LA Southwest', CCSF Score=125, Opponent Score=89)`. Since CCSF's score is more than 15 points larger than LA Southwest's score, then `is_big_win` applied to that row should output `True`.

_Points:_ 3

```python
def ...
    '''
    Return a boolean to describe whether or not a game (row) is a big win

    Examples
    --------
    >>> is_big_win(final_scores.row(0))
    True
    >>> is_big_win(final_scores.row(4))
    False
    '''
    ...
    
# Test your function 
row6 = final_scores.row(5)
is_big_win(row6)
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Use your `final_scores` table with your `is_big_win` function to assign `big_wins` to an array of team names that CCSF had big wins against during the 2024-25 basketball season.

_Points:_ 3

```python
big_wins = ...

for row_index ...
    row = ...
    opponent = ...
    if ...:
        big_wins = np.append(big_wins, ...)

big_wins
```

```python
grader.check("task_06")
```

---

### Joining Tables

If you run the code cell below you will import and display a new table called `game_stats` showing a collection of game performance statistics from the CCSF men's basketball 24/25 season. Below are a description of the variables:

- String valued - `opponent` played; `venue` played
- Float valued - field-goal percentage (`fg pct`), 3-point percentage (`3pt pct`), and free-throw percentage (`ft pct`) representing the percentage of each type of basket that was successfully scored.
- Integer valued - game number (`gm`), offensive rebounds (`off reb`), defensive rebounds (`def reb`), total rebounds (`total reb`), assists (`ast`), turnovers (`to`), steals (`st`), and blocks (`blk`) representing the count of each of these actions.


```python
game_stats = Table().read_table('ccsf_bb_24_25_stats.csv')
game_stats.show()
```

---

### Task 07 📍

Complete the following code to join the data in the `final_scores` table with the data in the `game_stats` table. Make sure the final score columns appear before the game stats columns. If two columns have duplicated information, remove the column that comes second. The resulting table should have a row for each game and 15 columns.

_Points:_ 3

```python
games = final_scores...
games
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Create the appproptiate visualization analyzing the relationship between field goal percentage (`fg pct`) and `CCSF Score`, then assign `task_08_answer` to the number of the most appropriate response below.

    1. Based on the visualization, there is no relationship between the field goal percentage and how many points CCSF scored.
    2. Based on the visualization, there is a positive relationship - games where the field goal percentage was higher tend to have higher CCSF scoring.
    3. Based on the visualization, there is a negative relationship - games where the field goal percentage was higher tend to have lower CCSF scoring.

_Points:_ 2

```python
# Make your visualization in this cell, and enter your answer in the following code.
...
```

```python
task_08_answer = ...
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Create the appproptiate visualization analyzing the average field goal percentage (`fg pct`) based on whether the game was at home, away, or at a neutral `venue`. Then assign `task_09_answer` to the number of the most appropriate response below.

    1. Based on the visualization, CCSF tends to have similar field goal percentage at home, away, and neutral venues.
    2. Based on the visualization, CCSF tends to have significantly higher field goal percentage at home then either away or neutral venues.
    3. Based on the visualization, CCSF tends to have significantly lower field goal percentage at home then either away or neutral venues.

_Points:_ 2

```python
# Make your visualization in this cell, and enter your answer in the following code.
...
```

```python
task_09_answer = ...
```

```python
grader.check("task_09")
```

---

## Roulette

---

### Chance

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ..., 36. Pockets 0 and 00 are green, and the other pockets are alternating red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="./roulette_wheel.jpeg" alt="roulette wheel" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

Before you do the following tasks, make sure you understand the logic behind all the examples in [Section 9.5](https://ccsf-math-108.github.io/textbook/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen. Listing all the distinct ways the event can occur.
- Seeing if it is easier to find the chance that the event does not happen.

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green.

---

### Task 10 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The winning pocket is black on all of the first three spins.

_Points:_ 2

```python
first_three_black = ...
```

```python
grader.check("task_10")
```

---

### Task 11 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The color green never wins in the first 10 spins.

_Points:_ 2

```python
no_green = ...
```

```python
grader.check("task_11")
```

---

### Task 12 📍

_Provide an expression that Python evaluates to the chance of the event described._

Event: The color green wins **at least** once on the first 10 spins.

_Points:_ 2

```python
at_least_one_green = ...
```

```python
grader.check("task_12")
```

---

### Comparing Chances

In each of the following two tasks, two events A and B are described. Answer the question 'Which of the events is more likely?' by choosing from the following three options and setting the variable in the task to the single integer that corresponds to the option you selected.

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time.

---

### Task 13 📍

A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo.

- Event A: all four different animals are picked, assuming the child picks without replacement
- Event B: all four different animals are picked, assuming the child picks with replacement

_Points:_ 2

```python
toys_option = ...
```

```python
grader.check("task_13")
```

---

### Task 14 📍

In a lottery, two numbers are drawn at random without replacement from the integers 1 through 1000.

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


_Points:_ 2

```python
lottery_option = ...
```

```python
grader.check("task_14")
```

---

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

_Points:_ 0

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

1. **Review the Rubric:** View the rubric on the Canvas assignment page to understand the scoring criteria.  
2. **Review Manually Graded Tasks:** Verify that you have responded to all manually graded tasks marked with 📍🔎.  
3. **Review Auto-Graded Tasks:** Inspect the output from the `grade.check_all()` command above. Ensure you have passed all available tests for the auto-graded tasks in this notebook marked with 📍.  
   - This command executes all auto-grader tests sequentially.  
   - Note that Homework and Project assignments include hidden auto-grader tests that are not visible to you. Passing all visible tests does **not** guarantee that your solution is fully correct.  
4. **Save Your Work:** In the notebook toolbar, go to `File → Save Notebook` to save your work and create a checkpoint.  
5. **Download the Notebook:** In the notebook toolbar, go to `File → Download IPYNB` to download the notebook (`.ipynb`) file.  
6. **Upload to Canvas:** Upload the downloaded `.ipynb` file to the appropriate Canvas assignment page.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: 2)
```
>>> sum_scores(45, 38, 12) == 95
True
```
failure_message: ❌ Your function doesn't seem to work correctly with the first example data.
success_message: ✅ Your function works with the first example data.
```
>>> sum_scores(50, 47, 0) == 97
True
```
failure_message: ❌ Your function doesn't seem to work correctly with the second example data.
success_message: ✅ Your function works with the second example data.


Tests for task_02 (points: 3)
```
>>> final_scores.num_columns == 4
True
```
failure_message: ❌ final_scores does not have the correct number of columns.
success_message: ✅ final_scores has the correct number of columns.
```
>>> set(['Game', 'Opponent', 'CCSF Score', 'Opponent Score']) == set(final_scores.labels)
True
```
failure_message: ❌ final_scores does not have the correct labels.
success_message: ✅ final_scores has the correct labels.


Tests for task_03 (points: 3)
```
>>> 'did_ccsf_win' in globals() and callable(did_ccsf_win)
True
```
failure_message: ❌ You didn't create a function called did_ccsf_win.
success_message: ✅ You've created a function called did_ccsf_win.
```
>>> bool(did_ccsf_win(final_scores.row(6)) == True)
True
```
failure_message: ❌ Your function doesn't that CCSF won against Contra Costa.
success_message: ✅ Your function shows that CCSF won against Contra Costa.
```
>>> bool(did_ccsf_win(final_scores.row(15)) == False)
True
```
failure_message: ❌ Your function doesn't show that CCSF lost against Las Positas.
success_message: ✅ Your function shows that CCSF lost against Las Positas


Tests for task_04 (points: 2)
```
>>> bool((0 <= ccsf_wins <= 33) == True)
True
```
failure_message: ❌ There were only 33 games, so CCSF won 0, 1, 2, ..., or 33 games.
success_message: ✅ ccsf_wins is a possible value
```
>>> bool((0 <= ccsf_losses <= 33) == True)
True
```
failure_message: ❌ There were only 33 games, so CCSF lost 0, 1, 2, ..., or 33 games.
success_message: ✅ ccsf_losses is a possible value


Tests for task_05 (points: 3)
```
>>> callable(is_big_win)
True
```
failure_message: ❌ It doesn't seem like you defined a function called is_big_win.
success_message: ✅ It seems like you defined a function called is_big_win.
```
>>> bool(is_big_win(final_scores.row(0)) == True)
True
```
failure_message: ❌ Your function doesn't produce the correct output for the first row of final_scores.
success_message: ✅ Your function produces the correct output for the first row of final_scores.
```
>>> bool(is_big_win(final_scores.row(4)) == False)
True
```
failure_message: ❌ Your function doesn't produce the correct output for the fifth row of final_scores.
success_message: ✅ Your function produces the correct output for the fifth row of final_scores.


Tests for task_06 (points: 3)
```
>>> isinstance(big_wins, np.ndarray)
True
```
failure_message: ❌ big_wins should be a NumPy array.
success_message: ✅ big_wins is a NumPy array.
```
>>> type(big_wins.item(0)) == str
True
```
failure_message: ❌ The first item in big_wins is not a string.
success_message: ✅ The first item in big_wins is a string.
```
>>> bool((0 <= len(big_wins) <= 33) == True)
True
```
failure_message: ❌ There is not a possible value for the number items in big_wins.
success_message: ✅ There is a possible value for the number items in big_wins.


Tests for task_07 (points: 3)
```
>>> isinstance(games, Table)
True
```
failure_message: ❌ games should be a Table. Are you using the join method?
success_message: ✅ games is a Table.


Tests for task_08 (points: 2)
```
>>> task_08_answer in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to task_08_answer.
success_message: ✅ This is a possible value for task_08_answer.


Tests for task_09 (points: 2)
```
>>> task_09_answer in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to task_09_answer.
success_message: ✅ This is a possible value for task_09_answer.


Tests for task_10 (points: 2)
```
>>> bool((0 <= first_three_black <= 1) == True)
True
```
failure_message: ❌ first_three_black should be a numerical value between 0 and 1, inclusive.
success_message: ✅ first_three_black is a numerical value between 0 and 1, inclusive.


Tests for task_11 (points: 2)
```
>>> bool((0 <= no_green <= 1) == True)
True
```
failure_message: ❌ no_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ no_green is a numerical value between 0 and 1, inclusive.


Tests for task_12 (points: 2)
```
>>> bool((0 <= at_least_one_green <= 1) == True)
True
```
failure_message: ❌ at_least_one_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ at_least_one_green is a numerical value between 0 and 1, inclusive.


Tests for task_13 (points: 2)
```
>>> toys_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to toys_option.
success_message: ✅ This is a possible value for toys_option.


Tests for task_14 (points: 2)
```
>>> lottery_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to lottery_option.
success_message: ✅ This is a possible value for lottery_option.
