<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 03: Tables and Visualizations

---

## References

* [Sections 6.0 - 6.4 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/06/Tables.html)
* [Sections 7.0 - 7.4 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/07/Visualization.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Markdown Cheat Sheet](https://www.markdownguide.org/cheat-sheet/)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Set Up Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the follwing code cell to set up the notebook.

_You are not responsible for understanding what this code does._

```python
# Initialize Otter
import otter
grader = otter.Notebook('lab03.ipynb')

# Configure this Notebook
import numpy as np
from datascience import *
import warnings
warnings.simplefilter(action='ignore', category=UserWarning)
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

<!-- END QUESTION -->

---

## Loading Data

---

### CSV Files

The CSV file `farmers_markets.csv` contains data on farmers' markets in the United States.  (The was data collected from [the USDA's website](https://apps.ams.usda.gov/FarmersMarketsExport/ExcelExport.aspx)). CSV (comma-separated values) refers to how the information in the file is organized. If you run the following code cell, you'll see the first 3 lines of the file and probably notice **how challenging it is to read**.

```python
!head -n 3 farmers_markets.csv
```

To better display that information and to provide you with a collection of tools to work with the data, we will guide you in storing that CSV into a `Table` format.

---

### `read_table`

The `read_table` function from the `datascience` library helps you load the contents from a CSV file and store the information within a `Table`.

Run the next cell to load the `farmers_markets` table.

```python
farmers_markets = Table.read_table('farmers_markets.csv')
```

After running that code cell, `farmers_markets` represents the table of information. In this lab and some future assignments, you are going to focus on working with these tables. If you run the following code cell, you'll see the contents of the table. 

* By default, Jupyter displays the first 10 rows of a table.
* Notice that it shows `... (1671 rows omitted)` below the displayed table.

```python
farmers_markets
```

Hopefully, that is a little easier to visualize compared to looking at the direct contents of the related CSV file! Now, you didn't do all of that just to visualize the information, you unlocked a collection of tools for you to engage with that data.

---

## Attributes and Methods

---

When you create something in Python, such as a `Table` (or any object), you gain access to a collection of properties (attributes) and functions (methods) associated with that object's data type. These properties and methods provide ways to interact with, manipulate, and retrieve information from the object.

*There are often many attributes and methods available, so it's common to use documentation and tools like artificial intelligence to explore and understand them. In this course, you will focus on a specific subset of these tools. We encourage you to use our reference materials and commit some of this syntax to memory rather than relying on artificial intelligence for learning.*

---

## Some `Table` Attributes and Methods

---

### `show`

Earlier, you were able to preview the `farmers_market` table by using its name, but how can you look at a specific number of rows or the entire table? The `show` table method in the `datascience` library is used to *display* a subset of rows from a `Table`. This method is handy for quickly inspecting the data.

* By default, `tbl.show()` displays all of the rows from the table `tbl`.
* You can specify the number of rows to display by passing a number as an argument.

For example, the following code shows the first 3 rows of the `farmers_market` table. This method is for visual inspection only and does not return a `Table` data type.

```python
farmers_markets.show(3)
```

---

#### Task 01 📍🔎

Use the method `show` to display the first 5 rows of `farmers_markets`.

Make sure to check your visualization with a classmate, a tutor, or the instructor before moving on since there is no auto-grader for this lab task.

**Important Note:** Be careful to not to call `.show()` without an argument. There is a lot of information in the farmers market table and showing it all will crash your kernel!

<!-- BEGIN QUESTION -->



```python
...
```

<!-- END QUESTION -->

Notice that some of the values in this table are missing, as denoted by `nan`. This means either that the value is not available (e.g. if we don’t know the market’s street address) or not applicable (e.g. if the market doesn’t have a street address). You'll also notice that the table has a large number of columns in it!

---

### `num_columns`

The table attribute `num_columns` returns the number of columns in a table. Running `tbl.num_columns` would provide you with the number of columns in the table `tbl`. Notice that `num_columns` is not a function, so you don't use `()` at the end. Think of `num_columns` as a variable name with specific stored information about the table.

---

#### Task 02 📍

Use `num_columns` to find the number of columns in our farmers' markets dataset.

Assign the number of columns to `num_farmers_markets_columns`.

```python
num_farmers_markets_columns = ...
print("The table has", num_farmers_markets_columns, "columns in it!")
```

```python
grader.check("task_02")
```

---

### `num_rows`

Similarly, the attribute `num_rows` tells you how many rows are in a table. Run the following cell to see how that table property can be accessed.

```python
num_farmers_markets_rows = farmers_markets.num_rows
print("The table has", num_farmers_markets_rows, "rows in it!")
```

---

### `select` and `drop`

Most of the columns in `farmers_markets` are about particular products -- whether the market sells tofu, pet food, etc.  If we're not interested in that information, it just makes the table difficult to read and potentially slows down the computer if we work with the table. This comes up more than you might think, because people who collect and publish data may not know ahead of time what people will want to do with it.

In such situations, we can use the table method `select` to choose only the columns that we want in a particular table. It takes any number of arguments. Each should be the name of a column in the table. It returns a new table with only those columns in it. The columns are in the order *in which they were listed as arguments*.

For example, the value of `farmers_markets.select("MarketName", "State")` is a table with only the name and the state of each farmers' market in `farmers_markets`.

---

#### Task 03 📍

Use `select` to create a table with only the name, city, county, state, longitude (`x`), and latitude (`y`) of each market, in that order. Assign that new table to the name `farmers_markets_locations`. 

**Note:** We didn't create the column names (and their format), they come directly from the United States Department of Agriculture data resource. Make sure to be exact when using column names with `select`. Double-check capitalization! Also, in this task, the order of the columns matters.

```python
farmers_markets_locations = ...
farmers_markets_locations
```

```python
grader.check("task_03")
```

---

### `drop`

`drop` serves the similar purpose as `select`, but it takes away the columns that you provide rather than the ones that you don't provide. Like `select`, `drop` returns a new table.

For example, the following code would create a copy of the `farmers_markets` table without the `FMID` and `updateTime` columns.

```python
farmers_markets.drop("FMID", "updateTime")
```

---

### `where`

Now let's say we want a table of all farmers' markets in California. We can use the table method `where` to do this. 

Run the following cell to filter the table to include only California farmers' markets.

```python
california_farmers_markets = farmers_markets_locations.where('State', are.equal_to('California'))
california_farmers_markets
```

Ignore the syntax for the moment.  Instead, try to read that line like this:

> Assign the name **`california_farmers_markets`** to a table whose rows are the rows in the **`farmers_markets_locations`** table **`where`** the **`'State'`**s **`are` `equal` `to` `California`**.

Let's dive into the details a bit more.  `where` takes 2 arguments:

1. The name of a column.  `where` finds rows where that column's values meet some criterion.
2. A predicate that describes the criterion that the column needs to meet.

The predicate in the example above called the function `are.equal_to` with the value we wanted, 'California'.  We'll see other predicates soon.

`where` returns a table that's a copy of the original table, but **with only the rows that meet the given predicate**.

---

#### Task 04 📍

Use `california_farmers_markets` to create a table called `sf_markets` containing farmers' markets in San Francisco, California. 

**Note:** These aren't all of the markets in the city, but they are the markets listed in USDA's data set.

```python
sf_markets = ...
sf_markets
```

```python
grader.check("task_04")
```

So far we've only been using `where` with the predicate that requires finding the values in a column to be *exactly* equal to a certain value. However, there are many other predicates. Here are a few:

|Predicate|Example|Result|
|-|-|-|
|`are.equal_to`|`are.equal_to(50)`|Find rows with values equal to 50|
|`are.not_equal_to`|`are.not_equal_to(50)`|Find rows with values not equal to 50|
|`are.above`|`are.above(50)`|Find rows with values above (and not equal to) 50|
|`are.above_or_equal_to`|`are.above_or_equal_to(50)`|Find rows with values above 50 or equal to 50|
|`are.below`|`are.below(50)`|Find rows with values below 50|
|`are.between`|`are.between(2, 10)`|Find rows with values above or equal to 2 and below 10|
|`are.containing`|`are.containing('Data Science')`|Find rows with values that contain the substring `'Data Science'`|

Next, you are going to practice using some of those predicates to filter the data in the table. For additional predicates, see the Code Reference [Table Filtering Predicates section](https://ccsf-math-108.github.io/materials-sp25/resources/code-reference.html#table-filtering-predicates).

---

### `take`

You might recall from early in this notebook that the `tbl.show(5)` will display the first 5 rows of the table `tbl`, but it doesn't output a table. So, what should you use if you want to actually make a table with only the first 5 rows of some other table?

The `take` table method can help you do this. If you run `tbl.take(5)`, then you won't end up with the first 5 rows of `tbl`. Instead, you'll produce a table with the 6th row of the `tbl` because it assumed that `5` refers to the row index 5 (aka the 6th row). 

Run the following code cell to see this.

```python
farmers_markets.take(5)
```

In order to get the first 5 rows, you need to provide an array of row indices as an argument for `take`. Specifically, you need the row indices 0, 1, 2, 3, 4. Thankfully, you learned about `np.arange` previously as a way to generate a sequence of numbers like this (i.e. `np.arange(5)`).

#### Task 05 📍

Assign `first_5_markets` to a table containing the information from the first 5 rows of `farmers_markets`.

```python
first_5_markets = ...
first_5_markets
```

```python
grader.check("task_05")
```

---

### `group`

The `farmers_markets` table is initially set up to have one market per row, but what if you want to analyze the markets from the State perspective? For example, it could be helpful to count how many markets there are in each state.

The `group` table method in the `datascience` library is a powerful tool for summarizing and aggregating data based on the unique values within one or more columns. It works by grouping rows in a table based on the distinct values in one or more specified columns.

Run the following cell which takes `farmers_market_locations` and creates a new table that will group the farmers' markets by `'State'` and list the counts for the number of farmers' markets in each state. We will go into greater depth with the `group` method later in the course. Notice that the column of counts automatically gets the label of `count`.

```python
farmers_markets_locations.group('State')
```

#### Task 06 📍

For this task, use `group` to assign `markets_by_county` to a table containing a row for each county in `california_farmers_markets` with a column (`'County'`) for the name of the county and a column (`'count'`) showing the count of markets within that county.

```python
markets_by_county = ...
markets_by_county
```

```python
grader.check("task_06")
```

---

### `sort`

Notice that `farmers_markets_locations.group('State')` sorted the information such that the states are in alphabetical order. If you want to re-organize the table to show the states with the most markets at the top, then you'd want to use the `sort` table method.

You need to provide `sort` with the column name (or index) you want to sort by and the information will be sorted in ascending order by default.

Run the following cell to find the states in the data set with the least farmers' markets.

```python
farmers_markets_locations.group('State').sort('count')
```

If you want the largest counts at the top of the table, you'd need to sort in descending order. You can do this by adjusting the `descending` argument to have a value of `True`. Specifically, you could type `.sort('count', descending=True)`.

---

#### Task 07 📍

Assign `sorted_markets_by_county` to a table containing the information in `markets_by_county` sorted such that the counts are in descending order and the counties with the largest counts are at the top of the table.

```python
sorted_markets_by_county = ...
sorted_markets_by_county
```

```python
grader.check("task_07")
```

---

## Visualizations and Attributes

Visualizing data is an essential step in gaining insights from the vast and complex datasets that permeate our modern world. There exists a myriad of techniques and tools to transform raw data into comprehensible, meaningful representations. Among these techniques, a set of standard visualizations has emerged as go-to options, each with its unique strengths and applications. The choice of which standard visualization to employ hinges on various factors, a key factor among them being the attribute type of the data under investigation. It's important to note that this attribute type may not always align with the data type in which information is stored, making the selection of an appropriate visualization an artful and pivotal decision in the data analysis process.

To streamline our understanding of attribute types for data visualization, we can simplify them into two broad categories: **numerical** and **categorical**. 
* Numerical attributes encompass data that consists of continuous or discrete numeric values. These attributes are typically quantitative and can be operated on mathematically. Examples of numerical attributes include variables like age, temperature, or income.
* Categorical attributes deal with data that fall into distinct categories or labels. They represent qualitative information where mathematical operations typically do not have clear meanings. Examples of categorical attributes include gender (male, female, nonbinary), color (red, blue, green), or product categories (electronics, clothing, food).

By classifying attributes into these two fundamental types, we can better tailor our choice of visualization methods to the nature of the data, allowing us to extract more valuable insights from it.

---

### Numbers for Categories

Just because an attribute has values that are numbers, does not mean you should treat the attribute as numerical. Postal codes are numbers. However, the attribute type for postal codes is categorical rather than numerical. Postal codes represent specific geographical regions and are not meant for mathematical operations like addition or subtraction. For example, `90210` (Beverly Hills) and `10001` (Manhattan) are categorical values representing different locations. Choosing an appropriate visualization method for postal codes would involve treating them categorically, not numerically, despite their data type.

---

### Task 08 📍

Which of the following attributes are categorical in nature? Assign `categorical_attributes` to an array with the numbers for the variables that represent categorical attributes.

1. Height in centimeters
2. Eye color (e.g., blue, brown, green)
3. Temperature in degrees Celsius
4. Years of education completed
5. Vehicle make and model (e.g., Toyota Camry)
6. Employee identification number
7. Blood type (e.g., A, B, AB, O)
8. Stock prices
9. Time of day (e.g., morning, afternoon, evening)
10. Mobile phone number

```python
categorical_attributes = ...
```

```python
grader.check("task_08")
```

---

## Bar Charts

When it comes to visualizing categorical data, one of the most effective tools is a bar chart. Bar charts display frequencies or proportions of each of the categories using bars of different lengths (or heights).

In Python, you can create bar charts easily using libraries like `matplotlib` or `datascience`. Specifically, the `datascience` library provides the `bar` and `barh` table methods, which simplify the process of generating bar charts. The `bar` method is used for vertical bar charts, while the `barh` method is employed for horizontal bar charts.

Run the following code cell to create the table `car_inventory`.

```python
car_inventory = Table().with_columns(
    'Car Type', ['Sedan', 'SUV', 'Truck', 'Hatchback'],
    'Count', [25, 15, 12, 8]
)
car_inventory
```

Using the `barh` table method, you can create a horizontal bar chart to visualize the distribution of car types in `car_inventory`.

Run the following code cell to generate the bar chart.

```python
car_inventory.barh('Car Type', 'Count')

# Optional Customization
plt.title('Distribution of Car Types')
plt.show()
```

This visualization tool is invaluable for allowing us to quickly see the distribution and comparison of the various categories in a categorical dataset.

The `barh` method takes a few arguments, but the first two are the most important.

* The first argument, `column_for_categories`, specifies the column in the table to be used for the categorical values in the bar chart. In the example above, `column_for_categories='Car Type'`.

* The second argument, `select`, specifies which column in the table will be used to determine the length of the bars. If you don't specify this argument, it will use all the columns it can to create multiple bar charts. In the example above, `select='Count'`.

There are other arguments that you can experiment with, but you will typically just need to work with these two.

### Using `group`

It is common for data to be stored in a way such that each line of the data represents one observation, so as an analyst, you usually need to create a table like `car_inventory` that has summary information. One tool to help with this creation is the `group` table method.

`tbl.group(col_name)` creates a new table from `tbl`, where each row corresponds to a unique value from the specified column. By default, the `group` function also adds a column to the new table that shows the count of occurrences for each unique value in `tbl`.

For example, the following table `customers` shows purchases for various customers where each line shows a single transaction.

```python
customers = Table().with_columns(
    'Customer ID', [101, 102, 101, 104, 105, 106, 107, 
                    108, 109, 110, 101, 105, 109, 106, 102],
    'Purchase', ['Laptop', 'Phone', 'Tablet', 'Laptop', 
                 'Phone', 'Tablet', 'Headphones', 'Laptop', 
                 'Phone', 'Headphones', 'Phone', 'Laptop', 
                 'Tablet', 'Laptop', 'Headphones'],
    'Amount ($)', [1200, 800, 400, 1300, 850, 420, 150, 
                   1100, 900, 200, 700, 1250, 450, 1150, 180]
)
customers
```

Customer 101 has made three purchases. Using `group` and `sort`, it becomes easier to see this. Run the following cell to see the grouped table.

```python
customers_by_ID = customers.group('Customer ID').sort('count', True)
customers_by_ID 
```

With the data organized in this way, you can visualize the customer purchase counts. Run the following code cell to see the results.

```python
customers_by_ID.barh('Customer ID')

plt.title('Customer Purchase Counts')
plt.show()
```

---

### `movies_by_year` Table

The `top_movies_1995_2022.csv` dataset has information about movie sales in recent years. Run the following cell to load that data as the Table `movies_by_year`.

```python
movies_by_year = Table.read_table("top_movies_1995_2022.csv")
movies_by_year
```

---

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

Using the [`movies_by_year` table](#movies_by_year-Table), create a bar chart showing the distribution of the movie distributors where the length of the bar reflects the number of movies in the table for the given distributor. Your bars should be sorted such that the longest bar is at the top of the graphic.

_Make sure to check your visualization with a classmate, a tutor, or the instructor before moving on since there is no auto-grader for this lab task._

```python
# Generate your chart in this cell
movies_by_distributor = movies_by_year.group(...).sort(...)
...

plt.title('Movie Distributor Counts')
plt.show()
```

<!-- END QUESTION -->

---

## Density Histograms

When it comes to visualizing numerical data, one of the default and fundamental techniques is to use a histogram. Histograms provide a graphical representation of the distribution of numerical values within a dataset, allowing you to observe patterns, central tendencies, and variations. 

In the `datascience` library, you can create histograms conveniently using the `hist` table method. 

When working with histograms, it's essential to consider the choice of bin sizes or intervals, as this can impact the interpretation of the data. 

The `hist` method defaults to displaying data density on the vertical axis rather than raw counts. This means that the height of each bar in the histogram represents the density of data within that bin, and the area of the bar, not the count, reflects the amount of data. This distinction is crucial for accurately understanding the distribution of numerical data and is a core concept in data visualization and analysis.

Run the following code cell to create a table called `ages` which contains the ages of a group of people.

```python
ages = Table().with_column(
    'Age', [12, 15, 18, 20, 22, 25, 26, 28, 30, 32, 35, 36, 38, 40, 45, 50, 55]
)
ages
```

To visualize this distribution, you can use the code `ages.hist('Age')`. The method hist has several arguments. The first argument `columns` identifies the column(s) that contain the numerical data for the histogram. Run the cell code below.

```python
ages.hist('Age')

# Optional Customization
plt.title('Distribution of Ages')
plt.show()
```

The `hist` table method offers two other arguments that are worth mentioning for this class: `bins` and `unit`. These arguments play a pivotal role in customizing the appearance and interpretation of the histogram. 

* The `bins` argument allows you to specify the number of bins or intervals into which the data range will be divided. A well-chosen number of bins can significantly affect the visual representation of the data, influencing the granularity of the histogram. By default, there is an algorithm that attempts to generate "good" bins, but you might need to specifically define the bins with an array or the number of bins with an integer to get the histogram to look good for your situation.
* The `unit` argument provides a way to provide labels to the horizontal and vertical axes as a reminder of what the units of the data are.

The ages are most likely measured in years and it might make sense to bin these ages by creating bins that are 10 years wide. You can achieve this with the parameters `unit="Years"` and `bins=np.arange(10, 61, 10)`. 

Run the following code cell to see the results.

```python
ages.hist('Age', unit="Years", bins=np.arange(10, 61, 10))

# Optional Customization
plt.title('Distribution of Ages')
plt.show()
```

Notice how the shape of the histogram changes! The same numerical data can look very different in a histogram depending on how it is binned. This can be used as a tool for analysis and inquiry, but it can also be used as a tool to misguide.

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Use the `hist` method with the [`movies_by_year` table](#movies_by_year-Table) to show the distribution of tickets sold. The default settings for `hist` make a reasonable graphic, but you can try adjusting the `bins` parameter to see what happens.

_Make sure to check your visualization with a classmate, a tutor, or the instructor before moving on, since there is no auto-grader for this lab task._

```python
# Generate your chart in this cell
...

plt.title('Distribution of Salaries')
plt.show()
```

<!-- END QUESTION -->

---

## Line Plots and Scatter Plots

When it comes to visualizing numerical relationships in data, scatter plots and line plots are two fundamental tools that provide valuable insights. 

In the `datascience` library, the `scatter` method is used to create scatter plots, while the `plot` method is employed to generate line plots. 

These visualization techniques share a conceptual similarity: both display data points on a two-dimensional plane, typically with one numerical variable on the x-axis and another on the y-axis. However, the key distinction lies in the purpose and interpretation.

* Scatter plots are versatile and are primarily used to showcase the association and general pattern between two numerical variables. They are excellent for revealing relationships, correlations, and outliers in the data.
* Line plots are best suited when the horizontal axis represents sequential data, such as time or distance. These plots connect the data points with lines, making them ideal for visualizing trends and showing how a numerical variable changes over a continuous range, as in the case of tracking revenue over time.

In essence, scatter plots excel at depicting associations, while line plots are tailor-made for illustrating trends and sequential relationships in numerical data.

For example, run the following code cell to generate a table called `company_data` showing revenue and profit data for the last ten years for some **fake** company.

```python
# Generate random data
np.random.seed(0)
years = np.arange(2014, 2024)
revenue = np.random.poisson(14, 10) * 2_500
profit = revenue * np.random.normal(0.08, 0.002, 10)

company_data = Table().with_columns('Year', years, 'Revenue', revenue, 'Profit', profit)
company_data
```

A line plot would be a standard choice to visualize the trend of profit over time. This can be done with the command `company_data.plot('Year', 'Profit')`.

Run the following code cell to see the results.

```python
company_data.plot('Year', 'Profit')

# Optional Customization
plt.title('Profits Over Time')
plt.show()
```

The trend of this line shows that there was something very significant that happened around 2018 to make the company very profitable. After a short period of time of high profits, there was a steep decline in profit levels that were lower than in years previous to 2018. For the last few years, the company's profits seem to lack stability. This is likely due to economic instability surrounding the pandemic (but remember that this profit data was made up).

It might be nice to compare the trends of two numerical distributions over the same horizontal axis. This would be a great time to try an overlaid line plot. For example, you could plot the lines for both profit and revenue over time.

The `plot` method can handle this by making sure the table only contains the variables you are interested in `'Year'`, `'Profit'`, and `'Revenue'` and just specifying the horizontal axis in the `plot` method. For example, just use `company_data.plot('Year')`.

Run the following cell to see that a line is created for every numerical column in the table other than `'Year'`.

```python
company_data.plot('Year')

# Optional Customization
plt.title('Revenue and Profits Over Time')
plt.show()
```

Revenue looks much less stable on this graph because of the scale of the values. Profits were hovering around 8% of revenue, so putting both lines on the y-axis doesn't offer a fair comparison.

How do you better see the relationship between revenue and profit? Since the data are not sequential and you are just looking to visualize the association, use a scatter plot. The `scatter` method would help out with this. Since profit follows from revenue, it is standard practice to have the horizontal axis reflect revenue values. Use the command `company_data.scatter('Revenue', 'Profit')` to make this happen.

Run the following cell to see the results.

```python
company_data.scatter('Revenue', 'Profit')

# Optional Customization
plt.title('Revenue vs. Profit')
plt.show()
```

This shows a pretty strong (linear) positive relationship between revenue and profit. Dividing the profit values by the revenue shows a pretty stable profit margin of roughly 8%.

```python
profit_margins = company_data.column('Profit') / company_data.column('Revenue')
print('The profit margins (profile/revenue) are:\n')
display(profit_margins)
average_profit_margin = "{:.2%}".format(np.average(profit_margins))
print(f'\nThe average profit margin over this period is {average_profit_margin}')
```

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

For this task, return to the [`movies_by_year` Table](#movies_by_year-Table) and create a line plot using the `plot` table method to visualize the trend of the number of tickets sold over time. What do you notice about the line graph in terms of trends in the number of tickets sold over time?

_Make sure to check your visualization with a classmate, a tutor, or the instructor before moving on since there is no auto-grader for this lab task._

```python
""" # BEGIN PROMPT
# Generate your chart in this cell
...

# Customization
plt.title('Tickets Sold Over Time')
plt.show()
"""; # END PROMPT
# BEGIN SOLUTION NO PROMPT
movies_by_year.plot('Year', 'Tickets Sold')
plt.title('Tickets Sold Over Time')
plt.show()
# END SOLUTION
```

<!-- END QUESTION -->

---

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

For this task return to the [`movies_by_year` table](#movies_by_year-Table). Using `movies_by_year`, create a scatterplot showing the relationship between the tickets sold and the gross revenue.
    
_Make sure to check your response with a classmate, a tutor, or the instructor before moving on since there is no auto-grader for this lab task._

```python
""" # BEGIN PROMPT
# Generate your chart in this cell
...
# Customization
plt.title('Tickets Sold vs. Total Gross')
plt.show()
"""; # END PROMPT
# BEGIN SOLUTION NO PROMPT
movies_by_year.scatter('Tickets Sold','Total Gross')
# Customization
plt.title('Tickets Sold vs. Total Gross')
plt.show()
# END SOLUTION
```

<!-- END QUESTION -->

Great work! You now have some practice with the basics of data visualization.

---

## Final Task 📍

<!-- BEGIN QUESTION -->

You're almost done! Run the cell below to check that you passed all of the non-hidden auto-graded tasks, and then follow the instructions below to submit your assignment to Canvas.

```python
grader.check_all()
```

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` above, to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

Tests for task_02 (points: None)
```
>>> num_farmers_markets_columns == 59
True
```
failure_message: ❌ It doesn't seem that you have the correct value. Make sure you using .num_columns directly after the name of the table.
success_message: ✅ Great work!


Tests for task_03 (points: None)
```
>>> sorted(farmers_markets_locations.labels) == ['County', 'MarketName', 'State', 'city', 'x', 'y']
True
```
failure_message: ❌ It doesn't seem that you have the correct columns. Make sure you are using the names of the 5 column, sepearted by commas, as the input for .select().
success_message: ✅ Great work. It seems you have the correct columns!
```
>>> farmers_markets_locations.num_rows == 1681 and farmers_markets_locations.sort('State').row(10).item(3) == 'Alabama'
True
```
failure_message: ❌ It doesn't seem that you all of the information from the table for those specific columns
success_message: ✅ Your table seems to have the correct information.


Tests for task_04 (points: None)
```
>>> sf_markets.num_rows == 3
True
```
failure_message: ❌ It doesn't seem that you have the correct number of markets. Make sure you are using the where method with the 'city' column as the first arugment and and correct city as the second argument
success_message: ✅ It seems you have the correct number of markets.
```
>>> list(sf_markets.column('city')) == ['San Francisco', 'San Francisco', 'San Francisco']
True
```
failure_message: ❌ It doesn't seem that you have filtered the table to just San Francisco markets. Make sure to double-check your spelling.
success_message: ✅ You've filtered the table to just San Francisco markets.


Tests for task_05 (points: None)
```
>>> isinstance(first_5_markets, Table)
True
```
failure_message: ❌ first_5_markets doesn't reflect a Table. Are you using the take table method?
success_message: ✅ first_5_markets reflects a Table.
```
>>> first_5_markets.num_rows == 5
True
```
failure_message: ❌ first_5_markets doesn't have the correct number of rows. Try using np.arange(5) as the argument to take.
success_message: ✅ first_5_markets has the correct number of rows.
```
>>> first_5_market_names = set(first_5_markets.column('MarketName'))
>>> actual_first_5_market_names = {' Caledonia Farmers Market Association - Danville', "18th Street Farmer's Market", "26th Annual Highlands Business Partnership's Farmers Market", "29 Palms Farmers' Market", '2nd Street Market - Five Rivers MetroPark'}
>>> first_5_market_names == actual_first_5_market_names
True
```
failure_message: ❌ first_5_markets doesn't seem correct.
success_message: ✅ first_5_markets seems correct.


Tests for task_06 (points: None)
```
>>> isinstance(markets_by_county, Table)
True
```
failure_message: ❌ It doesn't seem that you made a table with the correct name.
success_message: ✅ You've made a table with the correct name.
```
>>> set(markets_by_county.labels) == {'County', 'count'}
True
```
failure_message: ❌ markets_by_counts doesn't have the correct columns.
success_message: ✅ markets_by_counts has the correct columns.
```
>>> sum(markets_by_county.column('count'))
np.int64(136)
```
failure_message: ❌ markets_by_counts doesn't seem correct. Your counts are not adding up to the correct number.
success_message: ✅ markets_by_counts seems correct.


Tests for task_07 (points: None)
```
>>> sorted_markets_by_county.column('count').item(0)
28
```
failure_message: ❌ sorted_markets_by_county doesn't have the correct first row. Are you using descending=True with the sort method?
success_message: ✅ sorted_markets_by_county has the correct first row.


Tests for task_08 (points: None)
```
>>> isinstance(categorical_attributes, np.ndarray)
True
```
failure_message: ❌ categorical_attributes should be an array. Try using make_array to make an array with your choice of numbers 1 - 10.
success_message: ✅ categorical_attributes is an array.
```
>>> set(categorical_attributes).issubset(set(np.arange(1, 11)))
True
```
failure_message: ❌ categorical_attributes should contain numbers 1 - 10 to reflect the 10 possible attributes.
success_message: ✅ categorical_attributes contains numbers 1 - 10.
```
>>> set(categorical_attributes) == {2, 5, 6, 7, 9, 10}
True
```
failure_message: ❌ It seems like some of your choices are incorrect and are more likely to be considered as numerical attributes.
success_message: ✅ Great work on identifying all of the categorical attributes!
